Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.conSem
% (10322)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 174
% Passive clauses: 3720
% Generated clauses: 20110
% Final active clauses: 164
% Final passive clauses: 3348
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 680
% Trivial inequalities: 513
% Fw subsumption resolutions: 2143
% Fw demodulations: 3132
% Bw demodulations: 6
% 
% Simple tautologies: 534
% Equational tautologies: 56
% Forward subsumptions: 8973
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 319
% Factoring: 1
% Forward superposition: 6818
% Backward superposition: 5977
% Self superposition: 144
% Equality factoring: 90
% Equality resolution: 51
% 
% Split clauses: 218
% Split components: 448
% SAT solver clauses: 519
% SAT solver unit clauses: 11
% SAT solver binary clauses: 96
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 218
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 518
% TWLsolver calls for satisfiability: 159
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.conSem
% (10340)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 144
% Passive clauses: 1397
% Generated clauses: 14551
% Final active clauses: 134
% Final passive clauses: 1213
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 415
% Trivial inequalities: 369
% Fw subsumption resolutions: 962
% Fw demodulations: 1806
% Bw demodulations: 42
% Condensations: 755
% Global subsumptions: 8
% 
% Simple tautologies: 260
% Equational tautologies: 45
% Forward subsumptions: 6398
% 
% Binary resolution: 192
% Factoring: 1
% Forward superposition: 4702
% Backward superposition: 3580
% Self superposition: 111
% Equality factoring: 57
% Equality resolution: 35
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 2265
% SAT solver unit clauses: 155
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 49
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2231
% TWLsolver calls for satisfiability: 668
% 
% Memory used [KB]: 18166
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.conSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.conSem
% SZS output start Proof for thm.conSem
fof(f112796,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f429,f427,f260,f428,f426,f431,f1384,f1905,f1908,f1910,f1916,f1918,f1921,f1923,f1926,f1928,f1930,f1932,f1934,f1936,f1938,f1954,f1956,f1958,f1960,f1964,f1966,f1968,f1970,f1972,f1974,f1977,f2044,f2045,f2046,f2047,f5078,f5085,f5391,f5392,f5395,f5396,f6336,f6345,f6338,f6535,f6553,f6551,f8444,f8455,f8454,f12980,f13054,f13130,f13131,f13132,f13133,f13770,f13772,f13780,f13782,f13790,f13791,f13792,f13826,f13827,f13828,f13831,f13836,f13844,f13866,f13868,f13904,f13947,f13950,f13951,f13952,f13953,f15568,f15618,f19802,f19875,f18184,f17447,f4741,f1906,f4761,f6337,f20139,f20140,f2265,f1395,f1394,f22936,f23035,f13053,f112670,f112687])).
fof(f112687,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f112686,f429])).
fof(f112686,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f101559,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.conSem.v_|49||49||split|4','_i0_ax')).
fof(f101559,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f309,f95900])).
fof(f95900,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f86075,f327])).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.conSem.v_CASE_7'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.conSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.conSem.v_CASE_7'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.conSem.v'),X0),X7)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.conSem.v_CASE_7'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.conSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.conSem.v_CASE_7'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),X9)),X11),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp',X9))),X12),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.conSem.v'),X9),X15)))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_case_def|split|4')).
fof(f86075,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.conSem.v_CASE_7'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X1),s(fun('type.list.list'('type.conSem.v'),X0),X6)))) ) | $spl0),
  inference(superposition,[],[f327,f426])).
fof(f112670,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f112669,f429])).
fof(f112669,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f101534,f309])).
fof(f101534,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f95900,f309])).
fof(f13053,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) | $spl30),
  inference(cnf_transformation,[],[f13053_D])).
fof(f13053_D,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f23035,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) | $spl44),
  inference(cnf_transformation,[],[f23035_D])).
fof(f23035_D,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f22936,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(duplicate_literal_removal,[],[f22763])).
fof(f22763,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f19951,f20464])).
fof(f20464,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(trivial_inequality_removal,[],[f20449])).
fof(f20449,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f19949,f445])).
fof(f445,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.conSem.v_|49||49||split|4',aHL_BOOL_CASES)).
fof(f19949,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26 | $spl34)),
  inference(global_subsumption,[],[f19892,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1497,f1391,f1390,f1389,f1388,f1387,f1533,f1386,f1385,f319,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2228,f2229,f2230,f2231,f2232,f2524,f2730,f2704,f2803,f3321,f3322,f3415,f3633,f4107,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4507,f4508,f4669,f4590,f4600,f4606,f4673,f4743,f4829,f4873,f4872,f4866,f5054,f5079,f5082,f5078,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5389,f5391,f5393,f5394,f5400,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5688,f5689,f5700,f5844,f5849,f5860,f5862,f5866,f5872,f5886,f5928,f5933,f5944,f5946,f5950,f5956,f6219,f6336,f6345,f6480,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6789,f6790,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6819,f6820,f6831,f6832,f6833,f6836,f6837,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8444,f8454,f8520,f8521,f8609,f8610,f8620,f8795,f8841,f8845,f8846,f8848,f8849,f8850,f8851,f8852,f9225,f9229,f9230,f9232,f9233,f9234,f9235,f9236,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9466,f9428,f9429,f9431,f9432,f9433,f9434,f9435,f9561,f9761,f9947,f9948,f9858,f9863,f9872,f9878,f9884,f10049,f10011,f10012,f10014,f10015,f10016,f10017,f10018,f10148,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11387,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11955,f12150,f12155,f12437,f12406,f12547,f12714,f12719,f12720,f12722,f12723,f12724,f12725,f12726,f12727,f13130,f13132,f8395,f8411,f8445,f8450,f13136,f264,f275,f276,f472,f513,f540,f633,f686,f743,f909,f1378,f1488,f2371,f2417,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2529,f2530,f2531,f2689,f2913,f2915,f2917,f2949,f2951,f2965,f3320,f3328,f3337,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3718,f3728,f3729,f4073,f4078,f4081,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4506,f4514,f4515,f4557,f4558,f4568,f4570,f4571,f4575,f4576,f4586,f4667,f4766,f4826,f4848,f4918,f4975,f5381,f5388,f5390,f5402,f5687,f5695,f5696,f5808,f5809,f5819,f5821,f5822,f5826,f5827,f5837,f5893,f5894,f5904,f5906,f5907,f5911,f5912,f5922,f5968,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6229,f6342,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6490,f6567,f6575,f6576,f6608,f6612,f6616,f6787,f6788,f6796,f6797,f6818,f6826,f6827,f6834,f6835,f6843,f6844,f6864,f6878,f6882,f8319,f8327,f8328,f8334,f8342,f8343,f8377,f8383,f8391,f8414,f13954,f8524,f8798,f9470,f9745,f9753,f9754,f9793,f9795,f9796,f9800,f9801,f9813,f9903,f10054,f11823,f11824,f11825,f11826,f11829,f11834,f11904,f11905,f11965,f12033,f12145,f12334,f12443,f12661,f3330,f3329,f1376,f6568,f6569,f6621,f1374,f2661,f2663,f5372,f9430,f10013,f12401,f317,f320,f321,f346,f347,f1904,f4589,f4668,f4684,f4763,f6246,f6326,f6328,f6329,f6331,f6332,f6333,f6334,f6339,f6556,f6557,f6566,f6595,f6597,f6598,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6613,f9938,f12029,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12289,f12290,f12300,f12318,f12319,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12330,f12331,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12616,f12617,f12627,f12645,f12646,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12657,f12658,f345,f348,f5806,f5807,f5810,f5812,f5813,f5818,f5820,f5887,f5888,f5891,f5892,f5895,f5897,f5898,f5903,f5905,f5923,f6314,f9426,f10008,f11280,f11441,f11681,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12018,f12021,f12022,f12024,f12025,f12026,f12027,f12030,f12396,f5889,f8365,f12023,f11988,f11989,f12020,f12017,f12019,f11999,f5804,f6313,f6236,f6237,f6325,f6327,f9937,f6330,f311,f420,f2994,f2995,f2997,f2999,f3000,f3001,f3002,f3003,f4592,f6316,f6346,f6622,f8320,f8321,f8335,f8336,f339,f8367,f8369,f8370,f8372,f8373,f8374,f8375,f8388,f8318,f9776,f9746,f9747,f9777,f9778,f9814,f9817,f10010,f12398,f5845,f5843,f5929,f5927,f9775,f5848,f5932,f9940,f8308,f8309,f9953,f2998,f8371,f5138,f5803,f5838,f9425,f10007,f11281,f11442,f11682,f12395,f8366,f2996,f8368,f1889,f398,f349,f350,f354,f355,f356,f4779,f4789,f4823,f4928,f4938,f4972,f8470,f8483,f8745,f8758,f5852,f5854,f8473,f8748,f5936,f5938,f8474,f8749,f5851,f5853,f8460,f8735,f5935,f5937,f8461,f8736,f1484,f1486,f1909,f2227,f5370,f9427,f10009,f12397,f318,f829,f830,f831,f832,f833,f834,f930,f4518,f4519,f4521,f4523,f4524,f4525,f4526,f4527,f4591,f4671,f4672,f11301,f11462,f11702,f4522,f2699,f5931,f5847,f9832,f9828,f5116,f5113,f4520,f1515,f4555,f4556,f4559,f4561,f4562,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f4567,f4569,f5377,f11291,f11452,f11692,f4843,f4742,f5368,f5369,f5371,f5373,f5374,f5375,f5376,f14235,f14236,f15383,f15388,f15389,f15391,f15392,f15393,f15394,f15395])).
fof(f15395,plain,(
  ( ! [X132,X133,X131] : (s(bool,X131) = s(bool,sK7('const.pair.FST_0',X133,bool)) | s(bool,'T_0') = s(bool,sK4(X131,X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f836])).
fof(f15394,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,X129) = s(bool,sK7('const.pair.FST_0',X130,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X129)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f342])).
fof(f15393,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) = s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,X125)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f336])).
fof(f15392,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X121) = s(bool,sK7('const.pair.FST_0',X123,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,X121),s(X120,X122)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f341])).
fof(f15391,plain,(
  ( ! [X118,X119,X117] : (s(bool,X118) = s(bool,sK7('const.pair.FST_0',X119,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X117,bool),'const.sum.INR_1'(s(bool,X118)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f337])).
fof(f15389,plain,(
  ( ! [X111,X112,X110] : (s(bool,X110) = s(bool,sK7('const.pair.FST_0',X112,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f316])).
fof(f15388,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X108) = s(bool,sK7('const.pair.FST_0',X109,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X106),bool),sK18(X107,X108,X106,bool)),s('type.list.list'(X106),'const.list.NIL_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f357])).
fof(f15383,plain,(
  ( ! [X83,X84] : (s(bool,X83) = s(bool,sK7('const.pair.FST_0',X84,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X83)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14235,f309])).
fof(f14236,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14181,f264])).
fof(f14181,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f271,f13954])).
fof(f14235,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14174,f432])).
fof(f14174,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13954,f445])).
fof(f5376,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4843,f342])).
fof(f5375,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,X138)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X139),bool),sK18(X140,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X138)))),X139,bool))))) )),
  inference(superposition,[],[f4843,f336])).
fof(f5374,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4843,f341])).
fof(f5373,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,X129)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK18(X131,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X129)))),X130,bool))))) )),
  inference(superposition,[],[f4843,f337])).
fof(f5371,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK18(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) )),
  inference(superposition,[],[f4843,f316])).
fof(f5369,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK18(X114,happ(s(fun('type.list.list'(X110),bool),sK18(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4843,f357])).
fof(f5368,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4843,f309])).
fof(f4742,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK18(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4668,f357])).
fof(f4843,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4831,f264])).
fof(f4831,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4742,f270])).
fof(f11692,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2231,f836])).
fof(f11452,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2229,f836])).
fof(f11291,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | $spl10),
  inference(superposition,[],[f2225,f836])).
fof(f5377,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK18(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4843,f836])).
fof(f4569,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f494,f836])).
fof(f4567,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f447,f836])).
fof(f4593,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (s(fun(X166,X167),X168) != s(fun(X166,X167),X170) | s(X167,happ(s(fun(X166,X167),X170),s(X166,X171))) = s(X167,happ(s(fun(X166,X167),sK4(X168,X169,fun(X166,X167))),s(X166,X171))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X166,X167),bool),X169)))) )),
  inference(superposition,[],[f352,f836])).
fof(f6783,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f6817,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f4598,plain,(
  ( ! [X189,X187,X188] : (s('type.ast.lit',X187) != s('type.ast.lit',X189) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X189))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK4(X187,X188,'type.ast.lit')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X188)))) )),
  inference(superposition,[],[f335,f836])).
fof(f5305,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X48),s('type.list.list'('type.conSem.v'),X50))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X51),s('type.list.list'('type.conSem.v'),X52))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X51) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK4(X48,X49,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),bool),X49)))) )),
  inference(superposition,[],[f394,f836])).
fof(f10695,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s('type.list.list'('type.conSem.v'),X72) != s('type.list.list'('type.conSem.v'),X74) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X76),s('type.list.list'('type.conSem.v'),X74))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X75),s('type.list.list'('type.conSem.v'),sK4(X72,X73,'type.list.list'('type.conSem.v'))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X75) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X76) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.conSem.v'),bool),X73)))) )),
  inference(superposition,[],[f396,f836])).
fof(f11543,plain,(
  ( ! [X101,X99,X97,X102,X100,X98,X96] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X96),s('type.list.list'('type.string.char'),X98),s('type.conLang.exp',X99))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X100),s('type.list.list'('type.string.char'),X101),s('type.conLang.exp',X102))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X100) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK4(X96,X97,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),bool),X97)))) )),
  inference(superposition,[],[f399,f836])).
fof(f4562,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f418,f836])).
fof(f4561,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f333,f836])).
fof(f4559,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f330,f836])).
fof(f4556,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f272,f836])).
fof(f4555,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f271,f836])).
fof(f1515,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f318,f338])).
fof(f4520,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f836,f338])).
fof(f5113,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0x1_2'(s(fun(X17,bool),X19),s(fun(X16,X17),X18))))) | s(X16,X20) = s(X16,sK4(X20,happ(s(fun(fun(X17,bool),fun(X16,bool)),'_dstsel_0_x01_1'(s(fun(X16,X17),X18))),s(fun(X17,bool),X19)),X16))) )),
  inference(superposition,[],[f836,f368])).
fof(f5116,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun(X31,X32),X33)))))) | s(X31,X35) = s(X31,sK4(X35,happ(s(fun(fun(X32,bool),fun(X31,bool)),'_dstsel_0_x01_1'(s(fun(X31,X32),X33))),s(fun(X32,bool),X34)),X31))) )),
  inference(superposition,[],[f318,f368])).
fof(f9828,plain,(
  ( ! [X412,X414,X416,X413,X415,X411,X417] : (p(s(bool,'|dollar|exists_1'(s(fun(X411,bool),happ(s(fun(X413,fun(X411,bool)),happ(s(fun(X412,fun(X413,fun(X411,bool))),X414),s(X412,X415))),s(X413,X416)))))) | s(X411,X417) = s(X411,sK4(X417,happ(s(fun('type.pair.prod'(X412,X413),fun(X411,bool)),sK19(X414,X413,X412,fun(X411,bool))),s('type.pair.prod'(X412,X413),'const.pair.,_2'(s(X412,X415),s(X413,X416)))),X411))) )),
  inference(superposition,[],[f318,f367])).
fof(f9832,plain,(
  ( ! [X445,X447,X441,X443,X444,X446,X442] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X441,bool),happ(s(fun(X443,fun(X441,bool)),happ(s(fun(X442,fun(X443,fun(X441,bool))),X444),s(X442,X445))),s(X443,X446))))) | s(X441,X447) = s(X441,sK4(X447,happ(s(fun('type.pair.prod'(X442,X443),fun(X441,bool)),sK19(X444,X443,X442,fun(X441,bool))),s('type.pair.prod'(X442,X443),'const.pair.,_2'(s(X442,X445),s(X443,X446)))),X441))) )),
  inference(superposition,[],[f836,f367])).
fof(f5847,plain,(
  ( ! [X356,X358,X354,X360,X357,X359,X355] : (p(s(bool,'|dollar|exists_1'(s(fun(X354,bool),happ(s(fun(X357,fun(X354,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X354,bool)),X356),s(fun(X357,X355),X358))),s(X357,X359)))))) | s(X354,X360) = s(X354,sK4(X360,happ(s(fun(X355,fun(X354,bool)),X356),s(X355,happ(s(fun(X357,X355),X358),s(X357,X359)))),X354))) )),
  inference(superposition,[],[f318,f369])).
fof(f5931,plain,(
  ( ! [X350,X344,X346,X349,X345,X347,X348] : (p(s(bool,'|dollar|exists_1'(s(fun(X344,bool),happ(s(fun(X346,fun(X344,bool)),X347),s(X346,happ(s(fun(X345,X346),X348),s(X345,X349)))))))) | s(X344,X350) = s(X344,sK4(X350,happ(s(fun(X345,fun(X344,bool)),'_dst_x0x1_2'(s(fun(X346,fun(X344,bool)),X347),s(fun(X345,X346),X348))),s(X345,X349)),X344))) )),
  inference(superposition,[],[f318,f369])).
fof(f2699,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,'|dollar|exists_1'(s(fun(X139,bool),happ(s(fun('type.num.num',fun(X139,bool)),X141),s('type.num.num',X142)))))) | s(X139,X143) = s(X139,sK4(X143,'const.ind_type.FCONS_3'(s(fun(X139,bool),X140),s(fun('type.num.num',fun(X139,bool)),X141),s('type.num.num','const.num.SUC_1'(s('type.num.num',X142)))),X139))) )),
  inference(superposition,[],[f318,f324])).
fof(f4522,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))) | s(X16,X20) = s(X16,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))) )),
  inference(superposition,[],[f836,f324])).
fof(f11702,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f11662])).
fof(f11662,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2231,f836])).
fof(f11462,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f11422])).
fof(f11422,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2229,f836])).
fof(f11301,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f11261])).
fof(f11261,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(superposition,[],[f2225,f836])).
fof(f4672,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4620,f262])).
fof(f4620,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f418,f836])).
fof(f4671,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4619,f262])).
fof(f4619,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f333,f836])).
fof(f4591,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f318,f836])).
fof(f4527,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f836,f836])).
fof(f4526,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X37))) | s(X34,X38) = s(X34,sK4(X38,'const.pair.SND_1'(s('type.pair.prod'(X35,fun(X34,bool)),'const.pair.,_2'(s(X35,X36),s(fun(X34,bool),X37)))),X34))) )),
  inference(superposition,[],[f836,f342])).
fof(f4525,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X32))) | s(X30,X33) = s(X30,sK4(X33,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X30,bool),X31),'const.sum.INL_1'(s(fun(X30,bool),X32)))),X30))) )),
  inference(superposition,[],[f836,f336])).
fof(f4524,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f836,f341])).
fof(f4523,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTR_1'(s('type.sum.sum'(X22,fun(X21,bool)),'const.sum.INR_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f836,f337])).
fof(f4521,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f836,f316])).
fof(f4519,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f836,f357])).
fof(f4518,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f836,f309])).
fof(f930,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK18(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f318,f357])).
fof(f834,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X23)))) | s(X20,X24) = s(X20,sK4(X24,'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))) )),
  inference(superposition,[],[f318,f342])).
fof(f833,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | s(X16,X19) = s(X16,sK4(X19,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X16,bool),X17),'const.sum.INL_1'(s(fun(X16,bool),X18)))),X16))) )),
  inference(superposition,[],[f318,f336])).
fof(f832,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f318,f341])).
fof(f831,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X9)))) | s(X7,X10) = s(X7,sK4(X10,'const.sum.OUTR_1'(s('type.sum.sum'(X8,fun(X7,bool)),'const.sum.INR_1'(s(fun(X7,bool),X9)))),X7))) )),
  inference(superposition,[],[f318,f337])).
fof(f830,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X6) = s(X3,sK4(X6,'const.ind_type.FCONS_3'(s(fun(X3,bool),X4),s(fun('type.num.num',fun(X3,bool)),X5),s('type.num.num','const.num.|48|_0')),X3))) )),
  inference(superposition,[],[f318,f316])).
fof(f829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f318,f309])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : ! [X34,X24] : (! [X38] : (s(X9,X24) = s(X9,X38) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X34)))))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.JRH_INDUCT_UTIL')).
fof(f12397,plain,(
  ( ! [X83,X82] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),X83)))) ) | $spl10),
  inference(superposition,[],[f11955,f338])).
fof(f10009,plain,(
  ( ! [X127,X125,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X127),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X125),X126))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X125),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X125),X126)))) )),
  inference(superposition,[],[f6480,f338])).
fof(f9427,plain,(
  ( ! [X127,X125,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X127,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X125),X126))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X125),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X125),X126)))) )),
  inference(superposition,[],[f6219,f338])).
fof(f5370,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK18(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4843,f338])).
fof(f2227,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(forward_demodulation,[],[f2192,f2070])).
fof(f2192,plain,(
  ( ! [X23,X22] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(superposition,[],[f2168,f338])).
fof(f1909,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28))))) )),
  inference(forward_demodulation,[],[f1890,f338])).
fof(f1890,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f317,f338])).
fof(f1486,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f338,f445])).
fof(f1484,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f338,f445])).
fof(f8736,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f8461,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5937,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X393,X394),happ(s(fun(X396,fun(X393,X394)),X397),s(X396,happ(s(fun(X395,X396),X398),s(X395,X399))))),s(fun(X394,X393),X400)))) | s(X393,X401) = s(X393,happ(s(fun(X394,X393),X400),s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X396,fun(X393,X394)),X397),s(fun(X395,X396),X398))),s(X395,X399))),s(X393,X401)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5935,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X375,X376),happ(s(fun(X378,fun(X375,X376)),X379),s(X378,happ(s(fun(X377,X378),X380),s(X377,X381))))),s(fun(X376,X375),X382)))) | s(X376,X383) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),'_dst_x0x1_2'(s(fun(X378,fun(X375,X376)),X379),s(fun(X377,X378),X380))),s(X377,X381))),s(X375,happ(s(fun(X376,X375),X382),s(X376,X383)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f8735,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f8460,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5853,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X411,X404,X406] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X403,X404),happ(s(fun(X407,fun(X403,X404)),'_dst_x0x1_2'(s(fun(X405,fun(X403,X404)),X406),s(fun(X407,X405),X408))),s(X407,X409))),s(fun(X404,X403),X410)))) | s(X403,X411) = s(X403,happ(s(fun(X404,X403),X410),s(X404,happ(s(fun(X403,X404),happ(s(fun(X405,fun(X403,X404)),X406),s(X405,happ(s(fun(X407,X405),X408),s(X407,X409))))),s(X403,X411)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5851,plain,(
  ( ! [X389,X391,X385,X387,X393,X388,X390,X386,X392] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X385,X386),happ(s(fun(X389,fun(X385,X386)),'_dst_x0x1_2'(s(fun(X387,fun(X385,X386)),X388),s(fun(X389,X387),X390))),s(X389,X391))),s(fun(X386,X385),X392)))) | s(X386,X393) = s(X386,happ(s(fun(X385,X386),happ(s(fun(X387,fun(X385,X386)),X388),s(X387,happ(s(fun(X389,X387),X390),s(X389,X391))))),s(X385,happ(s(fun(X386,X385),X392),s(X386,X393)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f8749,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f8474,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5938,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X403,X402),X409),s(fun(X402,X403),happ(s(fun(X405,fun(X402,X403)),X406),s(X405,happ(s(fun(X404,X405),X407),s(X404,X408)))))))) | s(X403,X410) = s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),'_dst_x0x1_2'(s(fun(X405,fun(X402,X403)),X406),s(fun(X404,X405),X407))),s(X404,X408))),s(X402,happ(s(fun(X403,X402),X409),s(X403,X410)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5936,plain,(
  ( ! [X389,X391,X385,X387,X388,X390,X384,X386,X392] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X385,X384),X391),s(fun(X384,X385),happ(s(fun(X387,fun(X384,X385)),X388),s(X387,happ(s(fun(X386,X387),X389),s(X386,X390)))))))) | s(X384,X392) = s(X384,happ(s(fun(X385,X384),X391),s(X385,happ(s(fun(X384,X385),happ(s(fun(X386,fun(X384,X385)),'_dst_x0x1_2'(s(fun(X387,fun(X384,X385)),X388),s(fun(X386,X387),X389))),s(X386,X390))),s(X384,X392)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f8748,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f8473,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5854,plain,(
  ( ! [X412,X414,X420,X416,X418,X413,X415,X417,X419] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X413,X412),X419),s(fun(X412,X413),happ(s(fun(X416,fun(X412,X413)),'_dst_x0x1_2'(s(fun(X414,fun(X412,X413)),X415),s(fun(X416,X414),X417))),s(X416,X418)))))) | s(X413,X420) = s(X413,happ(s(fun(X412,X413),happ(s(fun(X414,fun(X412,X413)),X415),s(X414,happ(s(fun(X416,X414),X417),s(X416,X418))))),s(X412,happ(s(fun(X413,X412),X419),s(X413,X420)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5852,plain,(
  ( ! [X397,X399,X395,X401,X396,X398,X394,X400,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X395,X394),X401),s(fun(X394,X395),happ(s(fun(X398,fun(X394,X395)),'_dst_x0x1_2'(s(fun(X396,fun(X394,X395)),X397),s(fun(X398,X396),X399))),s(X398,X400)))))) | s(X394,X402) = s(X394,happ(s(fun(X395,X394),X401),s(X395,happ(s(fun(X394,X395),happ(s(fun(X396,fun(X394,X395)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400))))),s(X394,X402)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f8758,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f8745,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f8483,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f8470,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f4972,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4941,f262])).
fof(f4941,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f355,f445])).
fof(f4938,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f4928,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f4823,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4792,f262])).
fof(f4792,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f354,f445])).
fof(f4789,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f836])).
fof(f4779,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f836])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0)))))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X31,X9] : ! [X10,X37] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) <=> (! [X38] : s(X31,X38) = s(X31,happ(s(fun(X9,X31),X10),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))) & ! [X43] : s(X9,X43) = s(X9,happ(s(fun(X31,X9),X37),s(X31,happ(s(fun(X9,X31),X10),s(X9,X43)))))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.ISO')).
fof(f355,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f354,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X31] : ! [X37,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) => (! [X34] : (! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X34] : (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X6,X39] : (s(X9,X6) = s(X9,happ(s(fun(X31,X9),X37),s(X31,X39))) <=> s(X31,X39) = s(X31,happ(s(fun(X9,X31),X10),s(X9,X6))))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.ISO_USAGE')).
fof(f349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f398,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X31,X44,X9,X45] : ! [X46,X37,X47,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X45),X10),s(fun(X45,X9),X47)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X31,X44),X37),s(fun(X44,X31),X46))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X31),fun(X45,X44)),'_dst_x00x11_2'(s(fun(X31,X44),X37),s(fun(fun(X9,X31),fun(X45,X31)),'_dstsel_0_x01_1'(s(fun(X45,X9),X47))))),s(fun(fun(X45,X44),fun(X9,X31)),'_dst_x00x11_2'(s(fun(X44,X31),X46),s(fun(fun(X45,X44),fun(X9,X44)),'_dstsel_0_x01_1'(s(fun(X9,X45),X10)))))))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.ISO_FUN')).
fof(f1889,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.min.@_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)))))))) )),
  inference(superposition,[],[f317,f357])).
fof(f8368,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29))),s('type.pair.prod'(X26,X27),X30)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29))),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),fun('type.pair.prod'(X26,X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29)),X27,X26)),s(X27,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),fun('type.pair.prod'(X26,X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29)),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f8310,f338])).
fof(f8310,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29))),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),fun('type.pair.prod'(X26,X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29)),X27,X26)),s(X27,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),fun('type.pair.prod'(X26,X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29)),X27,X26))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),fun('type.pair.prod'(X26,X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X26,X27),bool),X28),X29))),s('type.pair.prod'(X26,X27),X30))))) )),
  inference(superposition,[],[f339,f338])).
fof(f2996,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f2980,f338])).
fof(f2980,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f338])).
fof(f8366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),X5)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f8306,f368])).
fof(f8306,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4))),s('type.pair.prod'(X0,X1),X5))))) )),
  inference(superposition,[],[f339,f368])).
fof(f12395,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X72,bool),X73),s(fun(X74,X72),X75))),s(X74,X76))))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76)))))) ) | $spl10),
  inference(superposition,[],[f11955,f369])).
fof(f11682,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X126),'const.sum.INL_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2231,f369])).
fof(f11442,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2229,f369])).
fof(f11281,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun(X75,bool),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X75,bool),X76),s(fun(X74,X75),X77))),s(X74,X78)))))) ) | $spl10),
  inference(superposition,[],[f2225,f369])).
fof(f10007,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6480,f369])).
fof(f9425,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6219,f369])).
fof(f5838,plain,(
  ( ! [X292,X294,X296,X293,X295,X291,X297] : (~p(s(bool,happ(s(fun(X293,bool),'_dst_x0x1_2'(s(fun(X291,bool),X292),s(fun(X293,X291),X294))),s(X293,X295)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X296),bool),sK18(X297,happ(s(fun(X291,bool),X292),s(X291,happ(s(fun(X293,X291),X294),s(X293,X295)))),X296,bool))))) )),
  inference(superposition,[],[f4843,f369])).
fof(f5803,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f4668,f369])).
fof(f5138,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36)))))) )),
  inference(forward_demodulation,[],[f5117,f368])).
fof(f5117,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36))))) | ~p(s(bool,happ(s(fun('type.list.list'(X36),bool),happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39))),s('type.list.list'(X36),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f368])).
fof(f8371,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),X47)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)))))))) )),
  inference(forward_demodulation,[],[f8313,f324])).
fof(f8313,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))),s('type.pair.prod'(X42,X43),X47))))) )),
  inference(superposition,[],[f339,f324])).
fof(f2998,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f2982,f324])).
fof(f2982,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f324])).
fof(f9953,plain,(
  ( ! [X613,X615,X609,X611,X616,X612,X614,X610] : (p(s(bool,happ(s(fun('type.pair.prod'(X609,X610),bool),happ(s(fun(X612,fun('type.pair.prod'(X609,X610),bool)),happ(s(fun(X611,fun(X612,fun('type.pair.prod'(X609,X610),bool))),X613),s(X611,X614))),s(X612,X615))),s('type.pair.prod'(X609,X610),X616)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X609,X610),bool),happ(s(fun(X612,fun('type.pair.prod'(X609,X610),bool)),happ(s(fun(X611,fun(X612,fun('type.pair.prod'(X609,X610),bool))),X613),s(X611,X614))),s(X612,X615))),s('type.pair.prod'(X609,X610),'const.pair.,_2'(s(X609,sK7(happ(s(fun('type.pair.prod'(X611,X612),fun('type.pair.prod'(X609,X610),bool)),sK19(X613,X612,X611,fun('type.pair.prod'(X609,X610),bool))),s('type.pair.prod'(X611,X612),'const.pair.,_2'(s(X611,X614),s(X612,X615)))),X610,X609)),s(X610,sK8(happ(s(fun('type.pair.prod'(X611,X612),fun('type.pair.prod'(X609,X610),bool)),sK19(X613,X612,X611,fun('type.pair.prod'(X609,X610),bool))),s('type.pair.prod'(X611,X612),'const.pair.,_2'(s(X611,X614),s(X612,X615)))),X610,X609)))))))) )),
  inference(forward_demodulation,[],[f9851,f367])).
fof(f9851,plain,(
  ( ! [X613,X615,X609,X611,X616,X612,X614,X610] : (~p(s(bool,happ(s(fun('type.pair.prod'(X609,X610),bool),happ(s(fun(X612,fun('type.pair.prod'(X609,X610),bool)),happ(s(fun(X611,fun(X612,fun('type.pair.prod'(X609,X610),bool))),X613),s(X611,X614))),s(X612,X615))),s('type.pair.prod'(X609,X610),'const.pair.,_2'(s(X609,sK7(happ(s(fun('type.pair.prod'(X611,X612),fun('type.pair.prod'(X609,X610),bool)),sK19(X613,X612,X611,fun('type.pair.prod'(X609,X610),bool))),s('type.pair.prod'(X611,X612),'const.pair.,_2'(s(X611,X614),s(X612,X615)))),X610,X609)),s(X610,sK8(happ(s(fun('type.pair.prod'(X611,X612),fun('type.pair.prod'(X609,X610),bool)),sK19(X613,X612,X611,fun('type.pair.prod'(X609,X610),bool))),s('type.pair.prod'(X611,X612),'const.pair.,_2'(s(X611,X614),s(X612,X615)))),X610,X609))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X609,X610),bool),happ(s(fun('type.pair.prod'(X611,X612),fun('type.pair.prod'(X609,X610),bool)),sK19(X613,X612,X611,fun('type.pair.prod'(X609,X610),bool))),s('type.pair.prod'(X611,X612),'const.pair.,_2'(s(X611,X614),s(X612,X615))))),s('type.pair.prod'(X609,X610),X616))))) )),
  inference(superposition,[],[f339,f367])).
fof(f8309,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,sK7(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18)),s(X19,sK8(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24))),s('type.pair.prod'(X18,X19),X25))))) )),
  inference(superposition,[],[f339,f369])).
fof(f8308,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X14,fun('type.pair.prod'(X10,X11),bool)),'_dst_x0x1_2'(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s('type.pair.prod'(X10,X11),'const.pair.,_2'(s(X10,sK7(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10)),s(X11,sK8(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16))))),s('type.pair.prod'(X10,X11),X17))))) )),
  inference(superposition,[],[f339,f369])).
fof(f9940,plain,(
  ( ! [X455,X457,X459,X460,X456,X458] : (~p(s(bool,happ(s(fun('type.list.list'(X455),bool),happ(s(fun(X457,fun('type.list.list'(X455),bool)),happ(s(fun(X456,fun(X457,fun('type.list.list'(X455),bool))),X458),s(X456,X459))),s(X457,X460))),s('type.list.list'(X455),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X455),bool),happ(s(fun(X457,fun('type.list.list'(X455),bool)),happ(s(fun(X456,fun(X457,fun('type.list.list'(X455),bool))),X458),s(X456,X459))),s(X457,X460))),s('type.list.list'(X455),sK2(happ(s(fun('type.pair.prod'(X456,X457),fun('type.list.list'(X455),bool)),sK19(X458,X457,X456,fun('type.list.list'(X455),bool))),s('type.pair.prod'(X456,X457),'const.pair.,_2'(s(X456,X459),s(X457,X460)))),X455)))))) )),
  inference(forward_demodulation,[],[f9834,f367])).
fof(f9834,plain,(
  ( ! [X455,X457,X459,X460,X456,X458] : (p(s(bool,happ(s(fun('type.list.list'(X455),bool),happ(s(fun(X457,fun('type.list.list'(X455),bool)),happ(s(fun(X456,fun(X457,fun('type.list.list'(X455),bool))),X458),s(X456,X459))),s(X457,X460))),s('type.list.list'(X455),sK2(happ(s(fun('type.pair.prod'(X456,X457),fun('type.list.list'(X455),bool)),sK19(X458,X457,X456,fun('type.list.list'(X455),bool))),s('type.pair.prod'(X456,X457),'const.pair.,_2'(s(X456,X459),s(X457,X460)))),X455))))) | ~p(s(bool,happ(s(fun('type.list.list'(X455),bool),happ(s(fun('type.pair.prod'(X456,X457),fun('type.list.list'(X455),bool)),sK19(X458,X457,X456,fun('type.list.list'(X455),bool))),s('type.pair.prod'(X456,X457),'const.pair.,_2'(s(X456,X459),s(X457,X460))))),s('type.list.list'(X455),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f367])).
fof(f5932,plain,(
  ( ! [X356,X352,X354,X351,X353,X355] : (p(s(bool,happ(s(fun('type.list.list'(X351),bool),happ(s(fun(X353,fun('type.list.list'(X351),bool)),X354),s(X353,happ(s(fun(X352,X353),X355),s(X352,X356))))),s('type.list.list'(X351),sK2(happ(s(fun(X352,fun('type.list.list'(X351),bool)),'_dst_x0x1_2'(s(fun(X353,fun('type.list.list'(X351),bool)),X354),s(fun(X352,X353),X355))),s(X352,X356)),X351))))) | ~p(s(bool,happ(s(fun('type.list.list'(X351),bool),happ(s(fun(X352,fun('type.list.list'(X351),bool)),'_dst_x0x1_2'(s(fun(X353,fun('type.list.list'(X351),bool)),X354),s(fun(X352,X353),X355))),s(X352,X356))),s('type.list.list'(X351),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f5848,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (p(s(bool,happ(s(fun('type.list.list'(X361),bool),happ(s(fun(X364,fun('type.list.list'(X361),bool)),'_dst_x0x1_2'(s(fun(X362,fun('type.list.list'(X361),bool)),X363),s(fun(X364,X362),X365))),s(X364,X366))),s('type.list.list'(X361),sK2(happ(s(fun(X362,fun('type.list.list'(X361),bool)),X363),s(X362,happ(s(fun(X364,X362),X365),s(X364,X366)))),X361))))) | ~p(s(bool,happ(s(fun('type.list.list'(X361),bool),happ(s(fun(X362,fun('type.list.list'(X361),bool)),X363),s(X362,happ(s(fun(X364,X362),X365),s(X364,X366))))),s('type.list.list'(X361),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f9775,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(X108,bool),happ(s(fun(X107,fun(X108,bool)),X109),s(X107,sK7(sK19(X109,X108,X107,bool),X108,X107)))),s(X108,sK8(sK19(X109,X108,X107,bool),X108,X107))))) | p(s(bool,happ(s(fun('type.pair.prod'(X107,X108),bool),sK19(X109,X108,X107,bool)),s('type.pair.prod'(X107,X108),X110))))) )),
  inference(superposition,[],[f339,f367])).
fof(f5927,plain,(
  ( ! [X321,X316,X318,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(X316,bool),happ(s(fun(X318,fun(X316,bool)),X319),s(X318,happ(s(fun(X317,X318),X320),s(X317,X321))))),s(X316,X322)))) | p(s(bool,'|dollar|exists_1'(s(fun(X316,bool),happ(s(fun(X317,fun(X316,bool)),'_dst_x0x1_2'(s(fun(X318,fun(X316,bool)),X319),s(fun(X317,X318),X320))),s(X317,X321))))))) )),
  inference(superposition,[],[f4668,f369])).
fof(f5929,plain,(
  ( ! [X333,X335,X331,X332,X334,X330,X336] : (~p(s(bool,happ(s(fun(X330,bool),happ(s(fun(X332,fun(X330,bool)),X333),s(X332,happ(s(fun(X331,X332),X334),s(X331,X335))))),s(X330,X336)))) | p(s(bool,happ(s(fun(X330,bool),happ(s(fun(X331,fun(X330,bool)),'_dst_x0x1_2'(s(fun(X332,fun(X330,bool)),X333),s(fun(X331,X332),X334))),s(X331,X335))),s(X330,'const.min.@_1'(s(fun(X330,bool),happ(s(fun(X331,fun(X330,bool)),'_dst_x0x1_2'(s(fun(X332,fun(X330,bool)),X333),s(fun(X331,X332),X334))),s(X331,X335))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f5843,plain,(
  ( ! [X327,X329,X331,X326,X332,X328,X330] : (~p(s(bool,happ(s(fun(X326,bool),happ(s(fun(X329,fun(X326,bool)),'_dst_x0x1_2'(s(fun(X327,fun(X326,bool)),X328),s(fun(X329,X327),X330))),s(X329,X331))),s(X326,X332)))) | p(s(bool,'|dollar|exists_1'(s(fun(X326,bool),happ(s(fun(X327,fun(X326,bool)),X328),s(X327,happ(s(fun(X329,X327),X330),s(X329,X331))))))))) )),
  inference(superposition,[],[f4668,f369])).
fof(f5845,plain,(
  ( ! [X344,X346,X341,X343,X345,X340,X342] : (~p(s(bool,happ(s(fun(X340,bool),happ(s(fun(X343,fun(X340,bool)),'_dst_x0x1_2'(s(fun(X341,fun(X340,bool)),X342),s(fun(X343,X341),X344))),s(X343,X345))),s(X340,X346)))) | p(s(bool,happ(s(fun(X340,bool),happ(s(fun(X341,fun(X340,bool)),X342),s(X341,happ(s(fun(X343,X341),X344),s(X343,X345))))),s(X340,'const.min.@_1'(s(fun(X340,bool),happ(s(fun(X341,fun(X340,bool)),X342),s(X341,happ(s(fun(X343,X341),X344),s(X343,X345))))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f12398,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X85,bool),happ(s(fun(X84,fun(X85,bool)),X86),s(X84,X87))),s(X85,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X84,X85),bool),sK19(X86,X85,X84,bool)),s('type.pair.prod'(X84,X85),'const.pair.,_2'(s(X84,X87),s(X85,X88)))))) ) | $spl10),
  inference(superposition,[],[f11955,f367])).
fof(f10010,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X133),'const.sum.INL_1'(s(bool,happ(s(fun(X129,bool),happ(s(fun(X128,fun(X129,bool)),X130),s(X128,X131))),s(X129,X132))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X128,X129),bool),sK19(X130,X129,X128,bool)),s('type.pair.prod'(X128,X129),'const.pair.,_2'(s(X128,X131),s(X129,X132)))))) )),
  inference(superposition,[],[f6480,f367])).
fof(f9817,plain,(
  ( ! [X350,X346,X349,X351,X347,X348] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X351,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X347,bool),happ(s(fun(X346,fun(X347,bool)),X348),s(X346,X349))),s(X347,X350))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X346,X347),bool),sK19(X348,X347,X346,bool)),s('type.pair.prod'(X346,X347),'const.pair.,_2'(s(X346,X349),s(X347,X350)))))) )),
  inference(superposition,[],[f6219,f367])).
fof(f9814,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(X328,bool),happ(s(fun(X327,fun(X328,bool)),X329),s(X327,X330))),s(X328,X331)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X332),bool),sK18(X333,happ(s(fun('type.pair.prod'(X327,X328),bool),sK19(X329,X328,X327,bool)),s('type.pair.prod'(X327,X328),'const.pair.,_2'(s(X327,X330),s(X328,X331)))),X332,bool))))) )),
  inference(superposition,[],[f4843,f367])).
fof(f9778,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X122,fun(X123,bool)),X124),s(X122,X125))),s(X123,X126)))) | p(s(bool,happ(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)),s('type.pair.prod'(X122,X123),'const.min.@_1'(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)))))))) )),
  inference(superposition,[],[f317,f367])).
fof(f9777,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X117,fun(X118,bool)),X119),s(X117,X120))),s(X118,X121)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X117,X118),bool),sK19(X119,X118,X117,bool)))))) )),
  inference(superposition,[],[f4668,f367])).
fof(f9747,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK19(X19,X18,X17,bool)),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f367,f445])).
fof(f9746,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,happ(s(fun('type.pair.prod'(X11,X12),bool),sK19(X13,X12,X11,bool)),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f367,f445])).
fof(f9776,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X113),s(X111,X114))),s(X112,X115)))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),X116),s('type.pair.prod'(X111,X112),sK6(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),sK19(X113,X112,X111,bool)),s('type.pair.prod'(X111,X112),sK5(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112))))))) )),
  inference(superposition,[],[f320,f367])).
fof(f8318,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),X72),s('type.pair.prod'(X70,X71),'const.pair.,_2'(s(X70,sK7(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70)),s(X71,sK8(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool))),s('type.pair.prod'(X70,X71),X74)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.pair.prod'(X70,X71),bool),bool),X73)))) )),
  inference(superposition,[],[f339,f836])).
fof(f8388,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f8350,f262])).
fof(f8350,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f339,f445])).
fof(f8375,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),X69)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)))))))) )),
  inference(forward_demodulation,[],[f8317,f342])).
fof(f8317,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),'const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68))))),s('type.pair.prod'(X64,X65),X69))))) )),
  inference(superposition,[],[f339,f342])).
fof(f8374,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),X63)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)))))))) )),
  inference(forward_demodulation,[],[f8316,f336])).
fof(f8316,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62))))),s('type.pair.prod'(X59,X60),X63))))) )),
  inference(superposition,[],[f339,f336])).
fof(f8373,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),X58)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)))))))) )),
  inference(forward_demodulation,[],[f8315,f341])).
fof(f8315,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57))))),s('type.pair.prod'(X53,X54),X58))))) )),
  inference(superposition,[],[f339,f341])).
fof(f8372,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),X52)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)))))))) )),
  inference(forward_demodulation,[],[f8314,f337])).
fof(f8314,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51))))),s('type.pair.prod'(X48,X49),X52))))) )),
  inference(superposition,[],[f339,f337])).
fof(f8370,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),X41)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)))))))) )),
  inference(forward_demodulation,[],[f8312,f316])).
fof(f8312,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'(X37,X38),X41))))) )),
  inference(superposition,[],[f339,f316])).
fof(f8369,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X31,X32),bool),X35),s('type.pair.prod'(X31,X32),X36)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X31,X32),bool),X35),s('type.pair.prod'(X31,X32),'const.pair.,_2'(s(X31,sK7(happ(s(fun('type.list.list'(X33),fun('type.pair.prod'(X31,X32),bool)),sK18(X34,X35,X33,fun('type.pair.prod'(X31,X32),bool))),s('type.list.list'(X33),'const.list.NIL_0')),X32,X31)),s(X32,sK8(happ(s(fun('type.list.list'(X33),fun('type.pair.prod'(X31,X32),bool)),sK18(X34,X35,X33,fun('type.pair.prod'(X31,X32),bool))),s('type.list.list'(X33),'const.list.NIL_0')),X32,X31)))))))) )),
  inference(forward_demodulation,[],[f8311,f357])).
fof(f8311,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X31,X32),bool),X35),s('type.pair.prod'(X31,X32),'const.pair.,_2'(s(X31,sK7(happ(s(fun('type.list.list'(X33),fun('type.pair.prod'(X31,X32),bool)),sK18(X34,X35,X33,fun('type.pair.prod'(X31,X32),bool))),s('type.list.list'(X33),'const.list.NIL_0')),X32,X31)),s(X32,sK8(happ(s(fun('type.list.list'(X33),fun('type.pair.prod'(X31,X32),bool)),sK18(X34,X35,X33,fun('type.pair.prod'(X31,X32),bool))),s('type.list.list'(X33),'const.list.NIL_0')),X32,X31))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X31,X32),bool),happ(s(fun('type.list.list'(X33),fun('type.pair.prod'(X31,X32),bool)),sK18(X34,X35,X33,fun('type.pair.prod'(X31,X32),bool))),s('type.list.list'(X33),'const.list.NIL_0'))),s('type.pair.prod'(X31,X32),X36))))) )),
  inference(superposition,[],[f339,f357])).
fof(f8367,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),X9)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)))))))) )),
  inference(forward_demodulation,[],[f8307,f309])).
fof(f8307,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8))),s('type.pair.prod'(X6,X7),X9))))) )),
  inference(superposition,[],[f339,f309])).
fof(f339,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9,X31] : ! [X34] : (! [X59,X60] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X59),s(X31,X60)))))) => ! [X61] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),X61)))))),
  file('thm.conSem.v_|49||49||split|4','thm.pair.pair_induction')).
fof(f8336,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,sK7(X7,bool,X8)),s(bool,X9)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK8(X7,bool,X8))) )),
  inference(superposition,[],[f339,f445])).
fof(f8335,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,sK7(X3,bool,X4)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),X5)))) | s(bool,X6) = s(bool,sK8(X3,bool,X4)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f8321,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X9),s(X8,sK8(X7,X8,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7(X7,X8,bool))) )),
  inference(superposition,[],[f339,f445])).
fof(f8320,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'T_0'),s(X4,sK8(X3,X4,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),X5)))) | s(bool,X6) = s(bool,sK7(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f6622,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74))))) )),
  inference(forward_demodulation,[],[f6596,f338])).
fof(f6596,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72))))))) )),
  inference(superposition,[],[f321,f338])).
fof(f6346,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),sK5('const.pair.FST_0',X80,'type.pair.prod'(bool,X78)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),X79)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X78),bool),X80),s('type.pair.prod'(bool,X78),sK6('const.pair.FST_0',X80,'type.pair.prod'(bool,X78))))))) )),
  inference(forward_demodulation,[],[f6315,f338])).
fof(f6315,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),X79)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X78),bool),X80),s('type.pair.prod'(bool,X78),sK6('const.pair.FST_0',X80,'type.pair.prod'(bool,X78)))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X78),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X78),sK5('const.pair.FST_0',X80,'type.pair.prod'(bool,X78))))))) )),
  inference(superposition,[],[f320,f338])).
fof(f6316,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X83)) | p(s(bool,happ(s(fun('type.list.list'(X81),bool),X84),s('type.list.list'(X81),sK6(sK18(X82,X83,X81,bool),X84,'type.list.list'(X81)))))) | p(s(bool,happ(s(fun('type.list.list'(X81),bool),sK18(X82,X83,X81,bool)),s('type.list.list'(X81),sK5(sK18(X82,X83,X81,bool),X84,'type.list.list'(X81))))))) )),
  inference(superposition,[],[f320,f357])).
fof(f4592,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f420,f836])).
fof(f3003,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f2987,f264])).
fof(f2987,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f420,f270])).
fof(f3002,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26)))))) )),
  inference(forward_demodulation,[],[f2986,f342])).
fof(f2986,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26))))) | ~p(s(bool,happ(s(fun('type.list.list'(X26),bool),'const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29))))),s('type.list.list'(X26),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f342])).
fof(f3001,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23)))))) )),
  inference(forward_demodulation,[],[f2985,f336])).
fof(f2985,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f336])).
fof(f3000,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f2984,f341])).
fof(f2984,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f341])).
fof(f2999,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16)))))) )),
  inference(forward_demodulation,[],[f2983,f337])).
fof(f2983,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f337])).
fof(f2997,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f2981,f316])).
fof(f2981,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f316])).
fof(f2995,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f2979,f357])).
fof(f2979,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f357])).
fof(f2994,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2978,f309])).
fof(f2978,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f309])).
fof(f420,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X34] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X24] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X24)))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X35),s('type.list.list'(X9),X24)))))))) => ! [X36] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X36)))))),
  file('thm.conSem.v_|49||49||split|4','thm.list.list_INDUCT')).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f6330,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(forward_demodulation,[],[f6241,f324])).
fof(f6241,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(superposition,[],[f320,f324])).
fof(f9937,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427] : (p(s(bool,happ(s(fun(X425,bool),happ(s(fun(X427,fun(X425,bool)),happ(s(fun(X426,fun(X427,fun(X425,bool))),X428),s(X426,X429))),s(X427,X430))),s(X425,sK5(happ(s(fun('type.pair.prod'(X426,X427),fun(X425,bool)),sK19(X428,X427,X426,fun(X425,bool))),s('type.pair.prod'(X426,X427),'const.pair.,_2'(s(X426,X429),s(X427,X430)))),X432,X425))))) | ~p(s(bool,happ(s(fun(X425,bool),happ(s(fun(X427,fun(X425,bool)),happ(s(fun(X426,fun(X427,fun(X425,bool))),X428),s(X426,X429))),s(X427,X430))),s(X425,X431)))) | p(s(bool,happ(s(fun(X425,bool),X432),s(X425,sK6(happ(s(fun('type.pair.prod'(X426,X427),fun(X425,bool)),sK19(X428,X427,X426,fun(X425,bool))),s('type.pair.prod'(X426,X427),'const.pair.,_2'(s(X426,X429),s(X427,X430)))),X432,X425)))))) )),
  inference(forward_demodulation,[],[f9830,f367])).
fof(f9830,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427] : (~p(s(bool,happ(s(fun(X425,bool),happ(s(fun(X427,fun(X425,bool)),happ(s(fun(X426,fun(X427,fun(X425,bool))),X428),s(X426,X429))),s(X427,X430))),s(X425,X431)))) | p(s(bool,happ(s(fun(X425,bool),X432),s(X425,sK6(happ(s(fun('type.pair.prod'(X426,X427),fun(X425,bool)),sK19(X428,X427,X426,fun(X425,bool))),s('type.pair.prod'(X426,X427),'const.pair.,_2'(s(X426,X429),s(X427,X430)))),X432,X425))))) | p(s(bool,happ(s(fun(X425,bool),happ(s(fun('type.pair.prod'(X426,X427),fun(X425,bool)),sK19(X428,X427,X426,fun(X425,bool))),s('type.pair.prod'(X426,X427),'const.pair.,_2'(s(X426,X429),s(X427,X430))))),s(X425,sK5(happ(s(fun('type.pair.prod'(X426,X427),fun(X425,bool)),sK19(X428,X427,X426,fun(X425,bool))),s('type.pair.prod'(X426,X427),'const.pair.,_2'(s(X426,X429),s(X427,X430)))),X432,X425)))))) )),
  inference(superposition,[],[f320,f367])).
fof(f6327,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK5(happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X30,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,X29)))) | p(s(bool,happ(s(fun(X26,bool),X30),s(X26,sK6(happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X30,X26)))))) )),
  inference(forward_demodulation,[],[f6238,f338])).
fof(f6238,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,X29)))) | p(s(bool,happ(s(fun(X26,bool),X30),s(X26,sK6(happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X30,X26))))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK5(happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X30,X26)))))) )),
  inference(superposition,[],[f320,f338])).
fof(f6325,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f6234,f368])).
fof(f6234,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(superposition,[],[f320,f368])).
fof(f6237,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,X24)))) | p(s(bool,happ(s(fun(X18,bool),X25),s(X18,sK6(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18))))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK5(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6236,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,X16)))) | p(s(bool,happ(s(fun(X10,bool),X17),s(X10,sK6(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10))))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK5(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6313,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70)))) | p(s(bool,happ(s(fun(X66,bool),X71),s(X66,sK6(X67,X71,X66))))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X67,X71,X66)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f5804,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun(X97,X95),X98))),s(X97,X99)))) | p(s(bool,happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f11999,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK10(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f12019,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,X40)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,sK10(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35)))))) )),
  inference(forward_demodulation,[],[f11990,f338])).
fof(f11990,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,sK10(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35))))) | p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,X40))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f12017,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f11986,f368])).
fof(f11986,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f12020,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(X41,X50)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | ~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,sK10(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42)))))) )),
  inference(forward_demodulation,[],[f11991,f367])).
fof(f11991,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | ~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,sK10(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42))))) | p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47))))),s(X41,X50))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f11989,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK10(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11988,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK10(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f12023,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f11994,f324])).
fof(f11994,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f8365,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun('type.pair.prod'(X72,X73),X74),X76),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,sK7('_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76)),X73,X72)),s(X73,sK8('_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76)),X73,X72))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X72,X73),bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76))),s('type.pair.prod'(X72,X73),X77))))) )),
  inference(superposition,[],[f339,f369])).
fof(f5889,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,happ(s(fun('type.list.list'(X96),X97),X99),s('type.list.list'(X96),sK2('_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99)),X96))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X96),bool),'_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99))),s('type.list.list'(X96),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f12396,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X78,bool),X79),s(X78,happ(s(fun(X77,X78),X80),s(X77,X81))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X77,bool),'_dst_x0x1_2'(s(fun(X78,bool),X79),s(fun(X77,X78),X80))),s(X77,X81)))) ) | $spl10),
  inference(superposition,[],[f11955,f369])).
fof(f12030,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f12002,f262])).
fof(f12002,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f345,f445])).
fof(f12027,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f11998,f342])).
fof(f11998,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f12026,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f11997,f336])).
fof(f11997,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f12025,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f11996,f341])).
fof(f11996,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f12024,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f11995,f337])).
fof(f11995,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f12022,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f11993,f316])).
fof(f11993,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f12021,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),X55),s(X51,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | ~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,sK10(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52)))))) )),
  inference(forward_demodulation,[],[f11992,f357])).
fof(f11992,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | ~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,sK10(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52))))) | p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0'))),s(X51,X58))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f12018,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f11987,f309])).
fof(f11987,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f11985,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK10(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f11984,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK10(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f11983,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK10(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f11982,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK10(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f11981,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK10(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f11980,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK10(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f11979,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK10(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f11978,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X51,X52),X55),s(fun(X52,X51),X56)))) | ~p(s(bool,happ(s(fun(X51,bool),X57),s(X51,sK10(X57,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X56,X52,X51))))) | p(s(bool,happ(s(fun(X51,bool),X57),s(X51,happ(s(fun(X52,X51),X56),s(X52,X58))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f11977,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(fun(X42,X41),X48)))) | ~p(s(bool,happ(s(fun(X41,bool),X49),s(X41,sK10(X49,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X48,X42,X41))))) | p(s(bool,happ(s(fun(X41,bool),X49),s(X41,happ(s(fun(X42,X41),X48),s(X42,X50))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f11976,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(fun(X35,X34),X38)))) | ~p(s(bool,happ(s(fun(X34,bool),X39),s(X34,sK10(X39,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X38,X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X39),s(X34,happ(s(fun(X35,X34),X38),s(X35,X40))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f11975,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK10(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11974,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK10(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11973,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK10(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f11972,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK10(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f11681,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X120),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2231,f369])).
fof(f11441,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X120,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2229,f369])).
fof(f11280,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(X71,bool),'_dst_x0x1_2'(s(fun(X69,bool),X70),s(fun(X71,X69),X72))),s(X71,X73))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X69,bool),X70),s(X69,happ(s(fun(X71,X69),X72),s(X71,X73)))))))) ) | $spl10),
  inference(superposition,[],[f2225,f369])).
fof(f10008,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6480,f369])).
fof(f9426,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X124,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6219,f369])).
fof(f6314,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | p(s(bool,happ(s(fun(X72,bool),X77),s(X72,sK6('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72))))) | p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,sK5('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f5923,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (~p(s(bool,happ(s(fun(X289,bool),X290),s(X289,happ(s(fun(X288,X289),X291),s(X288,X292)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X293),bool),sK18(X294,happ(s(fun(X288,bool),'_dst_x0x1_2'(s(fun(X289,bool),X290),s(fun(X288,X289),X291))),s(X288,X292)),X293,bool))))) )),
  inference(superposition,[],[f4843,f369])).
fof(f5905,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X184,X185),X187),s(X184,X188)))))) | s(bool,X189) = s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X184,X185),X187))),s(X184,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f494,f369])).
fof(f5903,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X173,X174),X176),s(X173,X177)))))) | s(bool,X178) != s(bool,happ(s(fun(X173,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X173,X174),X176))),s(X173,X177))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f447,f369])).
fof(f5898,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X144,bool),X145),s(X144,happ(s(fun(X143,X144),X146),s(X143,X147)))))) | s(bool,X148) != s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X144,bool),X145),s(fun(X143,X144),X146))),s(X143,X147))) | p(s(bool,X148))) )),
  inference(superposition,[],[f418,f369])).
fof(f5897,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | ~p(s(bool,X142)) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f333,f369])).
fof(f5895,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129)))))) | p(s(bool,X130)) | s(bool,X130) = s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f330,f369])).
fof(f5892,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114)))) )),
  inference(superposition,[],[f272,f369])).
fof(f5891,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f271,f369])).
fof(f5888,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun(X91,X92),X94),s(X91,X95)))))) | p(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))),s(X91,'const.min.@_1'(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f5887,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))) )),
  inference(superposition,[],[f4668,f369])).
fof(f5820,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (p(s(bool,happ(s(fun(X189,bool),'_dst_x0x1_2'(s(fun(X187,bool),X188),s(fun(X189,X187),X190))),s(X189,X191)))) | s(bool,X192) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f494,f369])).
fof(f5818,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X178,bool),'_dst_x0x1_2'(s(fun(X176,bool),X177),s(fun(X178,X176),X179))),s(X178,X180)))) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),X177),s(X176,happ(s(fun(X178,X176),X179),s(X178,X180))))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f447,f369])).
fof(f5813,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,X151) != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150))))) | p(s(bool,X151))) )),
  inference(superposition,[],[f418,f369])).
fof(f5812,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | ~p(s(bool,X145)) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) )),
  inference(superposition,[],[f333,f369])).
fof(f5810,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X130,X128),X131))),s(X130,X132)))) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f330,f369])).
fof(f5807,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117)))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f272,f369])).
fof(f5806,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X108,bool),X109),s(fun(X110,X108),X111))),s(X110,X112)))) | s(bool,'T_0') != s(bool,happ(s(fun(X108,bool),X109),s(X108,happ(s(fun(X110,X108),X111),s(X110,X112)))))) )),
  inference(superposition,[],[f271,f369])).
fof(f348,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f345,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f12658,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f12630,f262])).
fof(f12630,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f347,f445])).
fof(f12657,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f12656,f309])).
fof(f12656,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(subsumption_resolution,[],[f12628,f262])).
fof(f12628,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(superposition,[],[f347,f440])).
fof(f12655,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f12626,f342])).
fof(f12626,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f12654,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f12625,f336])).
fof(f12625,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f12653,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f12624,f341])).
fof(f12624,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f12652,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f12623,f337])).
fof(f12623,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f12651,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f12622,f324])).
fof(f12622,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73)))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f12650,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f12621,f316])).
fof(f12621,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f12649,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),X55),s(X51,sK12(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | ~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,X58))))) )),
  inference(forward_demodulation,[],[f12620,f357])).
fof(f12620,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | ~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,X58)))) | p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0'))),s(X51,sK12(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f12648,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(X41,sK12(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | ~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,X50))))) )),
  inference(forward_demodulation,[],[f12619,f367])).
fof(f12619,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | ~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,X50)))) | p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47))))),s(X41,sK12(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f12647,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,sK12(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,X40))))) )),
  inference(forward_demodulation,[],[f12618,f338])).
fof(f12618,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,X40)))) | p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,sK12(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f12646,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f12615,f309])).
fof(f12615,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f12645,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f12614,f368])).
fof(f12614,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f12627,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK12(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f12617,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33)))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK12(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f12616,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23)))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK12(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f12613,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK12(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f12612,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103)))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK12(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f12611,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95)))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK12(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f12610,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88)))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK12(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f12609,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80)))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK12(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f12608,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73)))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK12(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f12607,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65)))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK12(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f12606,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X51,X52),X55),s(fun(X52,X51),X56)))) | ~p(s(bool,happ(s(fun(X51,bool),X57),s(X51,X58)))) | p(s(bool,happ(s(fun(X51,bool),X57),s(X51,happ(s(fun(X52,X51),X56),s(X52,sK12(X57,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X56,X52,X51)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f12605,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(fun(X42,X41),X48)))) | ~p(s(bool,happ(s(fun(X41,bool),X49),s(X41,X50)))) | p(s(bool,happ(s(fun(X41,bool),X49),s(X41,happ(s(fun(X42,X41),X48),s(X42,sK12(X49,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X48,X42,X41)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f12604,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(fun(X35,X34),X38)))) | ~p(s(bool,happ(s(fun(X34,bool),X39),s(X34,X40)))) | p(s(bool,happ(s(fun(X34,bool),X39),s(X34,happ(s(fun(X35,X34),X38),s(X35,sK12(X39,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X38,X35,X34)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f12603,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33)))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK12(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f12602,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23)))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK12(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f12601,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK12(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f12600,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f12331,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f12303,f262])).
fof(f12303,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f346,f445])).
fof(f12330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f12329,f309])).
fof(f12329,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(subsumption_resolution,[],[f12301,f262])).
fof(f12301,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(superposition,[],[f346,f440])).
fof(f12328,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f12299,f342])).
fof(f12299,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f12327,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f12298,f336])).
fof(f12298,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f12326,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f12297,f341])).
fof(f12297,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f12325,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f12296,f337])).
fof(f12296,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f12324,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f12295,f324])).
fof(f12295,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f12323,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f12294,f316])).
fof(f12294,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f12322,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),X55),s(X51,sK11(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | p(s(bool,happ(s(fun(X52,bool),X57),s(X52,X58))))) )),
  inference(forward_demodulation,[],[f12293,f357])).
fof(f12293,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X52,X51),X56),s(fun(X51,X52),X55)))) | ~p(s(bool,happ(s(fun(X52,bool),X57),s(X52,happ(s(fun(X51,X52),happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0'))),s(X51,sK11(X57,X56,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X51,X52))))))) | p(s(bool,happ(s(fun(X52,bool),X57),s(X52,X58))))) )),
  inference(superposition,[],[f346,f357])).
fof(f12321,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(X41,sK11(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | p(s(bool,happ(s(fun(X42,bool),X49),s(X42,X50))))) )),
  inference(forward_demodulation,[],[f12292,f367])).
fof(f12292,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X41),X48),s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47)))))) | ~p(s(bool,happ(s(fun(X42,bool),X49),s(X42,happ(s(fun(X41,X42),happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47))))),s(X41,sK11(X49,X48,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X41,X42))))))) | p(s(bool,happ(s(fun(X42,bool),X49),s(X42,X50))))) )),
  inference(superposition,[],[f346,f367])).
fof(f12320,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,sK11(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | p(s(bool,happ(s(fun(X35,bool),X39),s(X35,X40))))) )),
  inference(forward_demodulation,[],[f12291,f338])).
fof(f12291,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X38),s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37))),s(X34,sK11(X39,X38,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X34,X35))))))) | p(s(bool,happ(s(fun(X35,bool),X39),s(X35,X40))))) )),
  inference(superposition,[],[f346,f338])).
fof(f12319,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f12288,f309])).
fof(f12288,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f12318,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f12287,f368])).
fof(f12287,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f12300,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK11(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f12290,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK11(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f12289,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK11(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f12286,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK11(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f12285,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK11(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f12284,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK11(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f12283,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK11(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f12282,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK11(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f12281,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK11(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f12280,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK11(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f12279,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X51,X52),X55),s(fun(X52,X51),X56)))) | ~p(s(bool,happ(s(fun(X51,bool),X57),s(X51,happ(s(fun(X52,X51),X56),s(X52,sK11(X57,happ(s(fun('type.list.list'(X53),fun(X51,X52)),sK18(X54,X55,X53,fun(X51,X52))),s('type.list.list'(X53),'const.list.NIL_0')),X56,X52,X51))))))) | p(s(bool,happ(s(fun(X51,bool),X57),s(X51,X58))))) )),
  inference(superposition,[],[f346,f357])).
fof(f12278,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X41,X42),happ(s(fun(X44,fun(X41,X42)),happ(s(fun(X43,fun(X44,fun(X41,X42))),X45),s(X43,X46))),s(X44,X47))),s(fun(X42,X41),X48)))) | ~p(s(bool,happ(s(fun(X41,bool),X49),s(X41,happ(s(fun(X42,X41),X48),s(X42,sK11(X49,happ(s(fun('type.pair.prod'(X43,X44),fun(X41,X42)),sK19(X45,X44,X43,fun(X41,X42))),s('type.pair.prod'(X43,X44),'const.pair.,_2'(s(X43,X46),s(X44,X47)))),X48,X42,X41))))))) | p(s(bool,happ(s(fun(X41,bool),X49),s(X41,X50))))) )),
  inference(superposition,[],[f346,f367])).
fof(f12277,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),'const.pair.FST_1'(s('type.pair.prod'(fun(X34,X35),X36),X37))),s(fun(X35,X34),X38)))) | ~p(s(bool,happ(s(fun(X34,bool),X39),s(X34,happ(s(fun(X35,X34),X38),s(X35,sK11(X39,happ(s(fun('type.pair.prod'(fun(X34,X35),X36),fun(X34,X35)),'const.pair.FST_0'),s('type.pair.prod'(fun(X34,X35),X36),X37)),X38,X35,X34))))))) | p(s(bool,happ(s(fun(X34,bool),X39),s(X34,X40))))) )),
  inference(superposition,[],[f346,f338])).
fof(f12276,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK11(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f12275,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK11(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f12274,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK11(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f12273,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK11(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f12029,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f12028,f309])).
fof(f12028,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f12000,f262])).
fof(f12000,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f345,f440])).
fof(f9938,plain,(
  ( ! [X436,X438,X434,X440,X437,X439,X433,X435] : (p(s(bool,happ(s(fun(X433,bool),happ(s(fun(X435,fun(X433,bool)),happ(s(fun(X434,fun(X435,fun(X433,bool))),X436),s(X434,X437))),s(X435,X438))),s(X433,sK6(X439,happ(s(fun('type.pair.prod'(X434,X435),fun(X433,bool)),sK19(X436,X435,X434,fun(X433,bool))),s('type.pair.prod'(X434,X435),'const.pair.,_2'(s(X434,X437),s(X435,X438)))),X433))))) | ~p(s(bool,happ(s(fun(X433,bool),happ(s(fun(X435,fun(X433,bool)),happ(s(fun(X434,fun(X435,fun(X433,bool))),X436),s(X434,X437))),s(X435,X438))),s(X433,sK5(X439,happ(s(fun('type.pair.prod'(X434,X435),fun(X433,bool)),sK19(X436,X435,X434,fun(X433,bool))),s('type.pair.prod'(X434,X435),'const.pair.,_2'(s(X434,X437),s(X435,X438)))),X433))))) | ~p(s(bool,happ(s(fun(X433,bool),X439),s(X433,X440))))) )),
  inference(forward_demodulation,[],[f9831,f367])).
fof(f9831,plain,(
  ( ! [X436,X438,X434,X440,X437,X439,X433,X435] : (~p(s(bool,happ(s(fun(X433,bool),happ(s(fun(X435,fun(X433,bool)),happ(s(fun(X434,fun(X435,fun(X433,bool))),X436),s(X434,X437))),s(X435,X438))),s(X433,sK5(X439,happ(s(fun('type.pair.prod'(X434,X435),fun(X433,bool)),sK19(X436,X435,X434,fun(X433,bool))),s('type.pair.prod'(X434,X435),'const.pair.,_2'(s(X434,X437),s(X435,X438)))),X433))))) | ~p(s(bool,happ(s(fun(X433,bool),X439),s(X433,X440)))) | p(s(bool,happ(s(fun(X433,bool),happ(s(fun('type.pair.prod'(X434,X435),fun(X433,bool)),sK19(X436,X435,X434,fun(X433,bool))),s('type.pair.prod'(X434,X435),'const.pair.,_2'(s(X434,X437),s(X435,X438))))),s(X433,sK6(X439,happ(s(fun('type.pair.prod'(X434,X435),fun(X433,bool)),sK19(X436,X435,X434,fun(X433,bool))),s('type.pair.prod'(X434,X435),'const.pair.,_2'(s(X434,X437),s(X435,X438)))),X433)))))) )),
  inference(superposition,[],[f321,f367])).
fof(f6613,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6582,f262])).
fof(f6582,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f321,f445])).
fof(f6606,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69))))) )),
  inference(forward_demodulation,[],[f6565,f342])).
fof(f6565,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69)))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64)))))) )),
  inference(superposition,[],[f321,f342])).
fof(f6605,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63))))) )),
  inference(forward_demodulation,[],[f6564,f336])).
fof(f6564,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63)))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59)))))) )),
  inference(superposition,[],[f321,f336])).
fof(f6604,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58))))) )),
  inference(forward_demodulation,[],[f6563,f341])).
fof(f6563,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58)))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53)))))) )),
  inference(superposition,[],[f321,f341])).
fof(f6603,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52))))) )),
  inference(forward_demodulation,[],[f6562,f337])).
fof(f6562,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48)))))) )),
  inference(superposition,[],[f321,f337])).
fof(f6602,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47))))) )),
  inference(forward_demodulation,[],[f6561,f324])).
fof(f6561,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47)))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)))))) )),
  inference(superposition,[],[f321,f324])).
fof(f6601,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41))))) )),
  inference(forward_demodulation,[],[f6560,f316])).
fof(f6560,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37)))))) )),
  inference(superposition,[],[f321,f316])).
fof(f6600,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X35,happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X31))))) | ~p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK5(X35,happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X31))))) | ~p(s(bool,happ(s(fun(X31,bool),X35),s(X31,X36))))) )),
  inference(forward_demodulation,[],[f6559,f357])).
fof(f6559,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK5(X35,happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X31))))) | ~p(s(bool,happ(s(fun(X31,bool),X35),s(X31,X36)))) | p(s(bool,happ(s(fun(X31,bool),happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0'))),s(X31,sK6(X35,happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X31)))))) )),
  inference(superposition,[],[f321,f357])).
fof(f6599,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK6(X29,happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK5(X29,happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30))))) )),
  inference(forward_demodulation,[],[f6558,f338])).
fof(f6558,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK5(X29,happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28))),s(X26,sK6(X29,happ(s(fun('type.pair.prod'(fun(X26,bool),X27),fun(X26,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,bool),X27),X28)),X26)))))) )),
  inference(superposition,[],[f321,f338])).
fof(f6598,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f6555,f309])).
fof(f6555,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f6597,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f6554,f368])).
fof(f6554,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0)))))) )),
  inference(superposition,[],[f321,f368])).
fof(f6595,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,happ(s(fun(X66,X67),X69),s(X66,sK5(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66))))))) | ~p(s(bool,happ(s(fun(X66,bool),X70),s(X66,X71)))) | p(s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69))),s(X66,sK6(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6566,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | ~p(s(bool,happ(s(fun(X70,bool),X73),s(X70,X74)))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK6(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f321,f836])).
fof(f6557,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,sK5(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18))))) | ~p(s(bool,happ(s(fun(X18,bool),X24),s(X18,X25)))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK6(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6556,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,sK5(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10))))) | ~p(s(bool,happ(s(fun(X10,bool),X16),s(X10,X17)))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK6(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6339,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6300,f262])).
fof(f6300,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f320,f445])).
fof(f6334,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(forward_demodulation,[],[f6245,f342])).
fof(f6245,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(superposition,[],[f320,f342])).
fof(f6333,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(forward_demodulation,[],[f6244,f336])).
fof(f6244,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(superposition,[],[f320,f336])).
fof(f6332,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(forward_demodulation,[],[f6243,f341])).
fof(f6243,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(superposition,[],[f320,f341])).
fof(f6331,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(forward_demodulation,[],[f6242,f337])).
fof(f6242,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(superposition,[],[f320,f337])).
fof(f6329,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(forward_demodulation,[],[f6240,f316])).
fof(f6240,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(superposition,[],[f320,f316])).
fof(f6328,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK5(happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X36,X31))))) | ~p(s(bool,happ(s(fun(X31,bool),X34),s(X31,X35)))) | p(s(bool,happ(s(fun(X31,bool),X36),s(X31,sK6(happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X36,X31)))))) )),
  inference(forward_demodulation,[],[f6239,f357])).
fof(f6239,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X34),s(X31,X35)))) | p(s(bool,happ(s(fun(X31,bool),X36),s(X31,sK6(happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X36,X31))))) | p(s(bool,happ(s(fun(X31,bool),happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0'))),s(X31,sK5(happ(s(fun('type.list.list'(X32),fun(X31,bool)),sK18(X33,X34,X32,fun(X31,bool))),s('type.list.list'(X32),'const.list.NIL_0')),X36,X31)))))) )),
  inference(superposition,[],[f320,f357])).
fof(f6326,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(forward_demodulation,[],[f6235,f309])).
fof(f6235,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f6246,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,X73)))) | p(s(bool,happ(s(fun(X70,bool),X74),s(X70,sK6(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK5(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f320,f836])).
fof(f4763,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4730,f262])).
fof(f4730,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4668,f445])).
fof(f4684,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4668,f836])).
fof(f4668,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4539,f271])).
fof(f4539,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f319,f836])).
fof(f4589,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f317,f836])).
fof(f1904,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1883,f262])).
fof(f1883,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f317,f445])).
fof(f347,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f346,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f321,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X34] : (! [X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))) => (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => ? [X38] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.MONO_EXISTS')).
fof(f320,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9] : ! [X34,X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,'const.min.@_1'(s(fun(X9,bool),X34)))))))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.SELECT_AX')).
fof(f12401,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X95),s('type.num.num',X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X94),s(fun('type.num.num',bool),X95),s('type.num.num','const.num.SUC_1'(s('type.num.num',X96)))))) ) | $spl10),
  inference(superposition,[],[f11955,f324])).
fof(f10013,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X142),s('type.num.num',X143))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X141),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.SUC_1'(s('type.num.num',X143)))))) )),
  inference(superposition,[],[f6480,f324])).
fof(f9430,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X138,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X136),s('type.num.num',X137))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X135),s(fun('type.num.num',bool),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))))) )),
  inference(superposition,[],[f6219,f324])).
fof(f5372,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun('type.num.num',bool),X124),s('type.num.num',X125)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X126),bool),sK18(X127,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.SUC_1'(s('type.num.num',X125)))),X126,bool))))) )),
  inference(superposition,[],[f4843,f324])).
fof(f2663,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f324,f445])).
fof(f2661,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f324,f445])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f319,f445])).
fof(f6621,plain,(
  ( ! [X64,X65] : (p(s(bool,sK6(X64,'_i_0_0',bool))) | ~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))))) )),
  inference(forward_demodulation,[],[f6594,f309])).
fof(f6594,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X64,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f6569,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X7,X8,bool))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool))) )),
  inference(superposition,[],[f321,f445])).
fof(f6568,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,sK6(X3,X4,bool))))) | s(bool,X6) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f321,f445])).
fof(f1376,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f319,f445])).
fof(f3329,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK14(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK14(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f3330,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK14(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK14(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f351,f445])).
fof(f12661,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f12660,f262])).
fof(f12660,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f12635,f2070])).
fof(f12635,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f347,f2531])).
fof(f12443,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(subsumption_resolution,[],[f12364,f432])).
fof(f12364,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(superposition,[],[f11955,f2803])).
fof(f12334,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f12333,f262])).
fof(f12333,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f12308,f2070])).
fof(f12308,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f346,f2531])).
fof(f12145,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f12054,f432])).
fof(f12054,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5393,f2803])).
fof(f12033,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f12032,f262])).
fof(f12032,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f12007,f2070])).
fof(f12007,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f345,f2531])).
fof(f11965,plain,(
  ( ! [X54,X53] : (s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(subsumption_resolution,[],[f11883,f264])).
fof(f11883,plain,(
  ( ! [X54,X53] : (p(s(bool,'F_0')) | s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(superposition,[],[f2426,f2913])).
fof(f11905,plain,(
  ( ! [X98] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X98),s(bool,sK15('_i_0_0',X98,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X98) | s(bool,'T_0') = s(bool,sK15('_i_0_0',X98,bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2913])).
fof(f11904,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X97),s(bool,sK15(X97,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X97) | s(bool,'T_0') = s(bool,sK15(X97,'_i_0_0',bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2913])).
fof(f11834,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103))) | s(bool,'T_0') = s(bool,sK4(X103,X104,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X104)))) ) | $spl10),
  inference(superposition,[],[f2913,f836])).
fof(f11829,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) ) | $spl10),
  inference(superposition,[],[f2913,f324])).
fof(f11826,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X81,bool),happ(s(fun(X80,fun(X81,bool)),X82),s(X80,X83))),s(X81,X84))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X80,X81),bool),sK19(X82,X81,X80,bool)),s('type.pair.prod'(X80,X81),'const.pair.,_2'(s(X80,X83),s(X81,X84)))))) ) | $spl10),
  inference(superposition,[],[f2913,f367])).
fof(f11825,plain,(
  ( ! [X78,X79] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),X79))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X78),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X78),X79)))) ) | $spl10),
  inference(superposition,[],[f2913,f338])).
fof(f11824,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X73,X74),X76),s(X73,X77))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X73,X74),X76))),s(X73,X77)))) ) | $spl10),
  inference(superposition,[],[f2913,f369])).
fof(f11823,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X70,bool),'_dst_x0x1_2'(s(fun(X68,bool),X69),s(fun(X70,X68),X71))),s(X70,X72))))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,happ(s(fun(X70,X68),X71),s(X70,X72)))))) ) | $spl10),
  inference(superposition,[],[f2913,f369])).
fof(f10054,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9979,f432])).
fof(f9979,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6480,f2803])).
fof(f9903,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(forward_demodulation,[],[f9751,f2070])).
fof(f9751,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(superposition,[],[f367,f2531])).
fof(f9813,plain,(
  ( ! [X325,X321,X323,X324,X326,X322] : (s(bool,'T_0') != s(bool,happ(s(fun(X322,bool),happ(s(fun(X321,fun(X322,bool)),X323),s(X321,X324))),s(X322,X325))) | s(bool,X326) = s(bool,happ(s(fun('type.pair.prod'(X321,X322),bool),sK19(X323,X322,X321,bool)),s('type.pair.prod'(X321,X322),'const.pair.,_2'(s(X321,X324),s(X322,X325))))) | s(bool,'F_0') = s(bool,X326)) ) | $spl10),
  inference(superposition,[],[f3633,f367])).
fof(f9801,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X256),'const.sum.INL_1'(s(bool,happ(s(fun(X252,bool),happ(s(fun(X251,fun(X252,bool)),X253),s(X251,X254))),s(X252,X255))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X251,X252),bool),sK19(X253,X252,X251,bool)),s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X254),s(X252,X255)))))) )),
  inference(superposition,[],[f2421,f367])).
fof(f9800,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X250,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X246,bool),happ(s(fun(X245,fun(X246,bool)),X247),s(X245,X248))),s(X246,X249))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X245,X246),bool),sK19(X247,X246,X245,bool)),s('type.pair.prod'(X245,X246),'const.pair.,_2'(s(X245,X248),s(X246,X249)))))) )),
  inference(superposition,[],[f2419,f367])).
fof(f9796,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X228,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X224,bool),happ(s(fun(X223,fun(X224,bool)),X225),s(X223,X226))),s(X224,X227))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X223,X224),bool),sK19(X225,X224,X223,bool)),s('type.pair.prod'(X223,X224),'const.pair.,_2'(s(X223,X226),s(X224,X227)))))) )),
  inference(superposition,[],[f540,f367])).
fof(f9795,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X222),'const.sum.INL_1'(s(bool,happ(s(fun(X218,bool),happ(s(fun(X217,fun(X218,bool)),X219),s(X217,X220))),s(X218,X221))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X217,X218),bool),sK19(X219,X218,X217,bool)),s('type.pair.prod'(X217,X218),'const.pair.,_2'(s(X217,X220),s(X218,X221)))))) )),
  inference(superposition,[],[f513,f367])).
fof(f9793,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X207,bool),happ(s(fun(X206,fun(X207,bool)),X208),s(X206,X209))),s(X207,X210))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X206,X207),bool),sK19(X208,X207,X206,bool)),s('type.pair.prod'(X206,X207),'const.pair.,_2'(s(X206,X209),s(X207,X210)))))) )),
  inference(superposition,[],[f472,f367])).
fof(f9754,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X58,bool),happ(s(fun(X57,fun(X58,bool)),X59),s(X57,X60))),s(X58,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X57,X58),bool),sK19(X59,X58,X57,bool)),s('type.pair.prod'(X57,X58),'const.pair.,_2'(s(X57,X60),s(X58,X61)))))) ) | $spl10),
  inference(superposition,[],[f367,f2803])).
fof(f9753,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,X54))),s(X52,X55))) | s(bool,X56) = s(bool,happ(s(fun('type.pair.prod'(X51,X52),bool),sK19(X53,X52,X51,bool)),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl10),
  inference(superposition,[],[f367,f2803])).
fof(f9745,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),sK19(X8,X7,X6,bool)),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f367,f270])).
fof(f9470,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9397,f432])).
fof(f9397,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6219,f2803])).
fof(f8798,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8797,f262])).
fof(f8797,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f8766,f2070])).
fof(f8766,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f350,f2531])).
fof(f8524,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8523,f262])).
fof(f8523,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f8491,f2070])).
fof(f8491,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f349,f2531])).
fof(f13954,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f8414,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2524,f2661,f2663,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3321,f3322,f3329,f3330,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4507,f4508,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4559,f4561,f4562,f4567,f4569,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4742,f4743,f4779,f4789,f4823,f4829,f4843,f4873,f4872,f4866,f4928,f4938,f4972,f5054,f5079,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5389,f5394,f5400,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5688,f5689,f5700,f5803,f5804,f5806,f5807,f5810,f5812,f5813,f5818,f5820,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5895,f5897,f5898,f5903,f5905,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6219,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6345,f6313,f6314,f6346,f6316,f6480,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6568,f6569,f6613,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6789,f6790,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6819,f6820,f6831,f6832,f6833,f6836,f6837,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8320,f8321,f8335,f8336,f8388,f8365,f8454,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8609,f8610,f8620,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8841,f8845,f8846,f8848,f8849,f8850,f8851,f8852,f9225,f9229,f9230,f9232,f9233,f9234,f9235,f9236,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9466,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9561,f9745,f9746,f9747,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10049,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10148,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12714,f12719,f12720,f12722,f12723,f12724,f12725,f12726,f8395,f8411,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f264,f275,f276,f472,f513,f540,f633,f686,f743,f909,f1378,f1488,f2371,f2417,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2529,f2530,f2689,f3320,f3328,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4506,f4557,f4558,f4568,f4570,f4571,f4575,f4576,f5388,f5390,f5402,f5687,f5808,f5809,f5819,f5821,f5822,f5826,f5827,f5893,f5894,f5904,f5906,f5907,f5911,f5912,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6229,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6490,f6567,f6612,f6787,f6788,f6818,f6834,f6835,f8319,f8334])).
fof(f8414,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(subsumption_resolution,[],[f8413,f262])).
fof(f8413,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f8402,f2070])).
fof(f8402,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(superposition,[],[f8395,f2531])).
fof(f8391,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f8390,f262])).
fof(f8390,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(forward_demodulation,[],[f8355,f2070])).
fof(f8355,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f8383,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(forward_demodulation,[],[f8340,f2070])).
fof(f8340,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f8377,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'T_0'),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f8325,f2070])).
fof(f8325,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f8343,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,sK7(X33,bool,X34)),s(bool,X35)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK8(X33,bool,X34))) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f8342,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,sK7(X29,bool,X30)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),X31)))) | s(bool,X32) = s(bool,sK8(X29,bool,X30)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f8334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK7(X0,bool,X1)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK8(X0,bool,X1))) )),
  inference(superposition,[],[f339,f270])).
fof(f8328,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,sK8(X33,X34,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK7(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f8327,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,sK8(X29,X30,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),X31)))) | s(bool,X32) = s(bool,sK7(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f8319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK8(X0,X1,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool))) )),
  inference(superposition,[],[f339,f270])).
fof(f6882,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6841,f2070])).
fof(f6841,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6878,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'T_0'))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(forward_demodulation,[],[f6824,f2070])).
fof(f6824,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6864,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6794,f2070])).
fof(f6794,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6844,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X46,X47,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X46,X47,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6843,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X42,X43,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X42,X43,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6835,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X13,X14,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,X14,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f6834,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15('_i_0_0',X11,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15('_i_0_0',X11,bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f6827,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,happ(s(fun(bool,X35),X33),s(bool,X36))) != s(X35,happ(s(fun(bool,X35),X34),s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK15(X33,X34,bool,X35))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6826,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,happ(s(fun(bool,X31),X29),s(bool,'F_0'))) != s(X31,happ(s(fun(bool,X31),X30),s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK15(X29,X30,bool,X31)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6818,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,X2))) )),
  inference(superposition,[],[f353,f270])).
fof(f6797,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X47,X46,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X47,X46,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6796,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X43,X42,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X43,X42,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6788,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X14,X13,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X14,X13,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f6787,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15(X11,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15(X11,'_i_0_0',bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f6616,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6615,f262])).
fof(f6615,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(forward_demodulation,[],[f6587,f2070])).
fof(f6587,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(superposition,[],[f321,f2531])).
fof(f6612,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6(X0,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(forward_demodulation,[],[f6611,f309])).
fof(f6611,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f6580,f262])).
fof(f6580,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f321,f472])).
fof(f6608,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f6573,f2070])).
fof(f6573,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2531])).
fof(f6576,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,sK6(X33,X34,bool))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK5(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2803])).
fof(f6575,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X30),s(bool,sK6(X29,X30,bool))))) | s(bool,X32) = s(bool,sK5(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f321,f2803])).
fof(f6567,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f321,f270])).
fof(f6490,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6430,f264])).
fof(f6430,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'F_0')) | s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2426,f2421])).
fof(f6394,plain,(
  ( ! [X130,X131,X129] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f2421,f836])).
fof(f6393,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X125,bool),'const.pair.,_2'(s(X125,X126),s(bool,X127)))))) )),
  inference(superposition,[],[f2421,f342])).
fof(f6391,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X121),'const.sum.INL_1'(s(bool,X119))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120)))))) )),
  inference(superposition,[],[f2421,f341])).
fof(f6389,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X112),s('type.num.num',X113))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X111),s(fun('type.num.num',bool),X112),s('type.num.num','const.num.SUC_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f2421,f324])).
fof(f6388,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X110),'const.sum.INL_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2421,f316])).
fof(f6387,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,X106))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X104),bool),sK18(X105,X106,X104,bool)),s('type.list.list'(X104),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2421,f357])).
fof(f6386,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X101),X102))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X101),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X101),X102)))) )),
  inference(superposition,[],[f2421,f338])).
fof(f6385,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f2421,f369])).
fof(f6384,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f2421,f369])).
fof(f6342,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(subsumption_resolution,[],[f6341,f262])).
fof(f6341,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(forward_demodulation,[],[f6305,f2070])).
fof(f6305,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(superposition,[],[f320,f2531])).
fof(f6229,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f6170,f264])).
fof(f6170,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f2426,f2419])).
fof(f6135,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X126))))) | s(bool,'T_0') = s(bool,sK4(X126,X127,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X127)))) )),
  inference(superposition,[],[f2419,f836])).
fof(f6134,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X125,bool),'const.sum.INR_1'(s(bool,X124))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X122,bool),'const.pair.,_2'(s(X122,X123),s(bool,X124)))))) )),
  inference(superposition,[],[f2419,f342])).
fof(f6132,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,X116))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),'const.pair.,_2'(s(bool,X116),s(X115,X117)))))) )),
  inference(superposition,[],[f2419,f341])).
fof(f6130,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X109),s('type.num.num',X110))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.SUC_1'(s('type.num.num',X110)))))) )),
  inference(superposition,[],[f2419,f324])).
fof(f6129,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,X105))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X105),s(fun('type.num.num',bool),X106),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2419,f316])).
fof(f6128,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK18(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2419,f357])).
fof(f6127,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X98),X99))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X98),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X98),X99)))) )),
  inference(superposition,[],[f2419,f338])).
fof(f6126,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,happ(s(fun(X92,X93),X95),s(X92,X96))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X93,bool),X94),s(fun(X92,X93),X95))),s(X92,X96)))) )),
  inference(superposition,[],[f2419,f369])).
fof(f6125,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X88,bool),'_dst_x0x1_2'(s(fun(X86,bool),X87),s(fun(X88,X86),X89))),s(X88,X90))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90)))))) )),
  inference(superposition,[],[f2419,f369])).
fof(f5968,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(forward_demodulation,[],[f5693,f2070])).
fof(f5693,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(superposition,[],[f369,f2531])).
fof(f5922,plain,(
  ( ! [X282,X285,X287,X283,X284,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X283,bool),X284),s(X283,happ(s(fun(X282,X283),X285),s(X282,X286))))) | s(bool,X287) = s(bool,happ(s(fun(X282,bool),'_dst_x0x1_2'(s(fun(X283,bool),X284),s(fun(X282,X283),X285))),s(X282,X286))) | s(bool,'F_0') = s(bool,X287)) ) | $spl10),
  inference(superposition,[],[f3633,f369])).
fof(f5912,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (p(s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X224,X225),X227),s(X224,X228)))))) | s(bool,X229) != s(bool,happ(s(fun(X224,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X224,X225),X227))),s(X224,X228))) | s(bool,'F_0') = s(bool,X229)) )),
  inference(superposition,[],[f2426,f369])).
fof(f5911,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (~p(s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X218,X219),X221),s(X218,X222)))))) | s(bool,X223) = s(bool,happ(s(fun(X218,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X218,X219),X221))),s(X218,X222))) | s(bool,'F_0') = s(bool,X223)) )),
  inference(superposition,[],[f2424,f369])).
fof(f5907,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X197,bool),X198),s(X197,happ(s(fun(X196,X197),X199),s(X196,X200))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X197,bool),X198),s(fun(X196,X197),X199))),s(X196,X200)))) )),
  inference(superposition,[],[f540,f369])).
fof(f5906,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X195),'const.sum.INL_1'(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f513,f369])).
fof(f5904,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183)))) )),
  inference(superposition,[],[f472,f369])).
fof(f5894,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f276,f369])).
fof(f5893,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) )),
  inference(superposition,[],[f275,f369])).
fof(f5837,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X287,bool),'_dst_x0x1_2'(s(fun(X285,bool),X286),s(fun(X287,X285),X288))),s(X287,X289))) | s(bool,X290) = s(bool,happ(s(fun(X285,bool),X286),s(X285,happ(s(fun(X287,X285),X288),s(X287,X289))))) | s(bool,'F_0') = s(bool,X290)) ) | $spl10),
  inference(superposition,[],[f3633,f369])).
fof(f5827,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (p(s(bool,happ(s(fun(X229,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X229,X227),X230))),s(X229,X231)))) | s(bool,X232) != s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X229,X227),X230),s(X229,X231))))) | s(bool,'F_0') = s(bool,X232)) )),
  inference(superposition,[],[f2426,f369])).
fof(f5826,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (~p(s(bool,happ(s(fun(X223,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X223,X221),X224))),s(X223,X225)))) | s(bool,X226) = s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X223,X221),X224),s(X223,X225))))) | s(bool,'F_0') = s(bool,X226)) )),
  inference(superposition,[],[f2424,f369])).
fof(f5822,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),X200),s(fun(X201,X199),X202))),s(X201,X203))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X199,bool),X200),s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))))) )),
  inference(superposition,[],[f540,f369])).
fof(f5821,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X198),'const.sum.INL_1'(s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X193,bool),X194),s(fun(X195,X193),X196))),s(X195,X197))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,happ(s(fun(X195,X193),X196),s(X195,X197)))))) )),
  inference(superposition,[],[f513,f369])).
fof(f5819,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X184,X182),X185))),s(X184,X186))))) | s(bool,'F_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X184,X182),X185),s(X184,X186)))))) )),
  inference(superposition,[],[f472,f369])).
fof(f5809,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X123,bool),X124),s(fun(X125,X123),X126))),s(X125,X127)))) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,happ(s(fun(X125,X123),X126),s(X125,X127)))))) )),
  inference(superposition,[],[f276,f369])).
fof(f5808,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'F_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) )),
  inference(superposition,[],[f275,f369])).
fof(f5696,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,X82) = s(bool,happ(s(fun(X79,bool),'_dst_x0x1_2'(s(fun(X77,bool),X78),s(fun(X79,X77),X80))),s(X79,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,happ(s(fun(X77,bool),X78),s(X77,happ(s(fun(X79,X77),X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f369,f2803])).
fof(f5695,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X71,bool),X72),s(fun(X73,X71),X74))),s(X73,X75))) | s(bool,X76) = s(bool,happ(s(fun(X71,bool),X72),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f369,f2803])).
fof(f5687,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X28,X26),X29))),s(X28,X30))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30)))))) )),
  inference(superposition,[],[f369,f270])).
fof(f5402,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5365,f262])).
fof(f5365,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4843,f513])).
fof(f5390,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f5357,f262])).
fof(f5357,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f4843,f540])).
fof(f5388,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5355,f262])).
fof(f5355,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4843,f472])).
fof(f5381,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f5380,f262])).
fof(f5380,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f5344,f2070])).
fof(f5344,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f4843,f2531])).
fof(f4975,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4974,f262])).
fof(f4974,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4946,f2070])).
fof(f4946,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f355,f2531])).
fof(f4918,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4883,f264])).
fof(f4883,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4829,f2803])).
fof(f4848,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4839,f264])).
fof(f4839,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4742,f2803])).
fof(f4826,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4825,f262])).
fof(f4825,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4797,f2070])).
fof(f4797,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f354,f2531])).
fof(f4766,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4765,f262])).
fof(f4765,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f4735,f2070])).
fof(f4735,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f4668,f2531])).
fof(f4667,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4512,f2070])).
fof(f4512,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2531])).
fof(f4586,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3633,f836])).
fof(f4576,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | s(bool,X113) != s(bool,sK4(X111,X112,bool)) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) )),
  inference(superposition,[],[f2426,f836])).
fof(f4575,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | s(bool,X110) = s(bool,sK4(X108,X109,bool)) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X109)))) )),
  inference(superposition,[],[f2424,f836])).
fof(f4571,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f540,f836])).
fof(f4570,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f513,f836])).
fof(f4568,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f472,f836])).
fof(f4558,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f276,f836])).
fof(f4557,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f275,f836])).
fof(f4515,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2803])).
fof(f4514,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f836,f2803])).
fof(f4506,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f836,f270])).
fof(f4463,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4406,f262])).
fof(f4406,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f540])).
fof(f4383,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94)))))) )),
  inference(superposition,[],[f540,f342])).
fof(f4382,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90)))))) )),
  inference(superposition,[],[f540,f336])).
fof(f4381,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X88,bool),'const.sum.INR_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f540,f341])).
fof(f4380,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83)))))) )),
  inference(superposition,[],[f540,f337])).
fof(f4379,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f540,f324])).
fof(f4378,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f540,f316])).
fof(f4377,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f540,f338])).
fof(f4376,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK18(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f540,f357])).
fof(f4344,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4288,f262])).
fof(f4288,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f513])).
fof(f4265,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f513,f342])).
fof(f4264,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X86),'const.sum.INL_1'(s(bool,X87)))))) )),
  inference(superposition,[],[f513,f336])).
fof(f4263,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f513,f341])).
fof(f4262,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,X80)))))) )),
  inference(superposition,[],[f513,f337])).
fof(f4261,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f513,f324])).
fof(f4260,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f513,f316])).
fof(f4259,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f513,f338])).
fof(f4258,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK18(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f513,f357])).
fof(f4148,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3633,f342])).
fof(f4147,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,X87)) ) | $spl10),
  inference(superposition,[],[f3633,f336])).
fof(f4146,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3633,f341])).
fof(f4145,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X80) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3633,f337])).
fof(f4144,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X75),s('type.num.num',X76))) | s(bool,X77) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl10),
  inference(superposition,[],[f3633,f324])).
fof(f4143,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3633,f316])).
fof(f4142,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3633,f338])).
fof(f4141,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK18(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3633,f357])).
fof(f4140,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3633,f309])).
fof(f4081,plain,(
  ( ! [X101,X99,X100,X98] : (s(X98,X100) = s(X98,sK4(X100,X99,X98)) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(bool,'F_0') = s(bool,X101)) ) | $spl10),
  inference(subsumption_resolution,[],[f3704,f264])).
fof(f3704,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,'F_0')) | s(X98,X100) = s(X98,sK4(X100,X99,X98)) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(bool,'F_0') = s(bool,X101)) ) | $spl10),
  inference(superposition,[],[f318,f2803])).
fof(f4078,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(subsumption_resolution,[],[f3696,f264])).
fof(f3696,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(superposition,[],[f420,f2803])).
fof(f4073,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3668,f264])).
fof(f3668,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f330,f2803])).
fof(f3729,plain,(
  ( ! [X187,X185,X188,X186] : (s(X187,f(s(bool,'F_0'))) != s(X187,g(s(bool,'F_0'))) | s(fun(bool,X187),X185) = s(fun(bool,X187),X186) | s(bool,X188) = s(bool,sK14(X185,X186,X187,bool)) | s(bool,'F_0') = s(bool,X188)) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3728,plain,(
  ( ! [X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,X184) = s(bool,sK4(X182,X183,bool)) | s(bool,'F_0') = s(bool,X184)) ) | $spl10),
  inference(superposition,[],[f319,f2803])).
fof(f3718,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'F_0') != s(bool,f(s(X148,sK14(X149,X150,bool,X148)))) | s(fun(X148,bool),X149) = s(fun(X148,bool),X150) | s(bool,X151) = s(bool,g(s(X148,sK14(X149,X150,bool,X148)))) | s(bool,'F_0') = s(bool,X151)) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3632,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f275,f2803])).
fof(f3559,plain,(
  ( ! [X198,X200,X199,X201] : (s(X200,f(s(bool,X201))) != s(X200,g(s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,'F_0') = s(bool,X201) | s(bool,'F_0') = s(bool,sK14(X198,X199,X200,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3558,plain,(
  ( ! [X196,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X196),s(bool,X197)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X196)))) | s(bool,'F_0') = s(bool,X197) | s(bool,'F_0') = s(bool,sK4(X195,X196,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2803])).
fof(f3548,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X164) != s(bool,f(s(X161,sK14(X162,X163,bool,X161)))) | s(fun(X161,bool),X162) = s(fun(X161,bool),X163) | s(bool,'F_0') = s(bool,X164) | s(bool,'F_0') = s(bool,g(s(X161,sK14(X162,X163,bool,X161))))) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3451,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2803,f342])).
fof(f3450,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') = s(bool,X71) | s(bool,X72) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X70),'const.sum.INL_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f2803,f336])).
fof(f3449,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2803,f341])).
fof(f3448,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X63,bool),'const.sum.INR_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2803,f337])).
fof(f3447,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61))) | s(bool,X62) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl10),
  inference(superposition,[],[f2803,f324])).
fof(f3446,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2803,f316])).
fof(f3445,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2803,f338])).
fof(f3444,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK18(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2803,f357])).
fof(f3380,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2803,f342])).
fof(f3379,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X83),'const.sum.INL_1'(s(bool,X84)))))) ) | $spl10),
  inference(superposition,[],[f2803,f336])).
fof(f3378,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2803,f341])).
fof(f3377,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X76,bool),'const.sum.INR_1'(s(bool,X77)))))) ) | $spl10),
  inference(superposition,[],[f2803,f337])).
fof(f3376,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl10),
  inference(superposition,[],[f2803,f324])).
fof(f3375,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2803,f316])).
fof(f3374,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2803,f338])).
fof(f3373,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK18(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2803,f357])).
fof(f3372,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2803,f309])).
fof(f3340,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f3334,f2070])).
fof(f3334,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f351,f2531])).
fof(f3337,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f3326,f2070])).
fof(f3326,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2531])).
fof(f3328,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK14(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK14(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f351,f270])).
fof(f3320,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,X2,bool))) )),
  inference(superposition,[],[f351,f270])).
fof(f2965,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2866,f2070])).
fof(f2866,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2531])).
fof(f2951,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2950,f262])).
fof(f2950,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2835,f2070])).
fof(f2835,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f317,f2531])).
fof(f2949,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2948,f262])).
fof(f2948,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2834,f2070])).
fof(f2834,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f319,f2531])).
fof(f2917,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2798,f2070])).
fof(f2798,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(superposition,[],[f2531,f324])).
fof(f2915,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2796,f2070])).
fof(f2796,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2531,f338])).
fof(f2913,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2794,f2070])).
fof(f2794,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2531,f309])).
fof(f2689,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f472,f324])).
fof(f2531,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2509,f262])).
fof(f2509,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f2424,f2070])).
fof(f2530,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2508,f262])).
fof(f2508,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2424,f440])).
fof(f2529,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2507,f262])).
fof(f2507,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2424,f472])).
fof(f2427,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2394,f262])).
fof(f2394,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f447,f472])).
fof(f2426,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2425,f309])).
fof(f2425,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2391,f262])).
fof(f2391,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f418,f472])).
fof(f2424,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2423,f309])).
fof(f2423,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2390,f262])).
fof(f2390,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f333,f472])).
fof(f2422,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(forward_demodulation,[],[f2376,f309])).
fof(f2376,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f472,f342])).
fof(f2421,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(forward_demodulation,[],[f2375,f309])).
fof(f2375,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X35))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35)))))) )),
  inference(superposition,[],[f472,f336])).
fof(f2420,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2374,f309])).
fof(f2374,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f472,f341])).
fof(f2419,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f2373,f309])).
fof(f2373,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30)))))) )),
  inference(superposition,[],[f472,f337])).
fof(f2418,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2372,f309])).
fof(f2372,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f472,f316])).
fof(f2417,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2370,f309])).
fof(f2370,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f472,f357])).
fof(f2371,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f472,f338])).
fof(f1488,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X12),X13)))) )),
  inference(superposition,[],[f338,f270])).
fof(f1378,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f319,f270])).
fof(f909,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,X2,X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f357,f270])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f342,f270])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f341,f270])).
fof(f633,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f270])).
fof(f540,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f337,f270])).
fof(f513,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f336,f270])).
fof(f472,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f309,f270])).
fof(f276,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|3')).
fof(f275,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f264,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.conSem.v_|49||49||split|4',aHL_FALSITY)).
fof(f13136,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22)))) ) | $spl10),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4593,f4598,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5305,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6783,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6817,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10695,f10709,f11226,f11301,f11280,f11281,f11291,f11387,f11462,f11441,f11442,f11452,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11543,f11627,f11702,f11681,f11682,f11692,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450])).
fof(f8450,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f8437,f264])).
fof(f8437,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f8414,f2803])).
fof(f8445,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8429,f264])).
fof(f8429,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(superposition,[],[f8414,f270])).
fof(f8411,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8397,f262])).
fof(f8397,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f8395,f445])).
fof(f8395,plain,(
  ( ! [X70,X71] : (~p(s(bool,sK7('const.pair.FST_0',X70,bool))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),X71))))) )),
  inference(forward_demodulation,[],[f8394,f338])).
fof(f8394,plain,(
  ( ! [X70,X71] : (~p(s(bool,sK7('const.pair.FST_0',X70,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X70),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X70),X71))))) )),
  inference(forward_demodulation,[],[f8364,f341])).
fof(f8364,plain,(
  ( ! [X70,X71] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,sK7('const.pair.FST_0',X70,bool)),s(X70,sK8('const.pair.FST_0',X70,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X70),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X70),X71))))) )),
  inference(superposition,[],[f339,f338])).
fof(f12727,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f9561,f2803])).
fof(f12726,plain,(
  ( ! [X175,X173,X174] : (s(bool,X173) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X175,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X173,X174,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X174)))) ) | $spl26),
  inference(superposition,[],[f9561,f836])).
fof(f12725,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,X171) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X172,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X169,bool),'const.pair.,_2'(s(X169,X170),s(bool,X171)))))) ) | $spl26),
  inference(superposition,[],[f9561,f342])).
fof(f12724,plain,(
  ( ! [X167,X166,X168] : (s(bool,X167) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X168,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X166),'const.sum.INL_1'(s(bool,X167)))))) ) | $spl26),
  inference(superposition,[],[f9561,f336])).
fof(f12723,plain,(
  ( ! [X165,X163,X164,X162] : (s(bool,X163) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X165,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) ) | $spl26),
  inference(superposition,[],[f9561,f341])).
fof(f12722,plain,(
  ( ! [X161,X159,X160] : (s(bool,X160) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X161,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X159,bool),'const.sum.INR_1'(s(bool,X160)))))) ) | $spl26),
  inference(superposition,[],[f9561,f337])).
fof(f12720,plain,(
  ( ! [X154,X152,X153] : (s(bool,X152) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X154,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X152),s(fun('type.num.num',bool),X153),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f9561,f316])).
fof(f12719,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,X150) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X151,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X148),bool),sK18(X149,X150,X148,bool)),s('type.list.list'(X148),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f9561,f357])).
fof(f12714,plain,(
  ( ! [X125,X126] : (s(bool,X125) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X125)))) ) | $spl26),
  inference(superposition,[],[f9561,f309])).
fof(f12547,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f12530])).
fof(f12530,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f12437,f445])).
fof(f12406,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107))) | s(bool,'T_0') = s(bool,sK4(X107,X108,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X108)))) ) | $spl10),
  inference(superposition,[],[f11955,f836])).
fof(f12437,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f12344,f8454])).
fof(f12344,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('const.pair.FST_0',X0,bool))))) ) | ($spl10 | $spl26)),
  inference(unit_resulting_resolution,[],[f8610,f11955])).
fof(f12155,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f12101,f262])).
fof(f12101,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(superposition,[],[f4843,f5393])).
fof(f12150,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f12075,f262])).
fof(f12075,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f447,f5393])).
fof(f11955,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f11857,f264])).
fof(f11857,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f271,f2913])).
fof(f11627,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2231,f445])).
fof(f11542,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.conLang.exp',X92))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X93),s('type.list.list'('type.string.char'),X94),s('type.conLang.exp',X95))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X93) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.pair.SND_1'(s('type.pair.prod'(X88,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.,_2'(s(X88,X89),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X90)))))) )),
  inference(superposition,[],[f399,f342])).
fof(f11541,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X82),s('type.list.list'('type.string.char'),X83),s('type.conLang.exp',X84))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X85),s('type.list.list'('type.string.char'),X86),s('type.conLang.exp',X87))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X85) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X81),'const.sum.INL_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X82)))))) )),
  inference(superposition,[],[f399,f336])).
fof(f11540,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75,X73] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X74),s('type.list.list'('type.string.char'),X76),s('type.conLang.exp',X77))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X78),s('type.list.list'('type.string.char'),X79),s('type.conLang.exp',X80))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X78) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X73),'const.pair.,_2'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X74),s(X73,X75)))))) )),
  inference(superposition,[],[f399,f341])).
fof(f11539,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X67),s('type.list.list'('type.string.char'),X68),s('type.conLang.exp',X69))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X70),s('type.list.list'('type.string.char'),X71),s('type.conLang.exp',X72))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X70) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.sum.OUTR_1'(s('type.sum.sum'(X66,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.sum.INR_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X67)))))) )),
  inference(superposition,[],[f399,f337])).
fof(f11538,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X63),s('type.list.list'('type.string.char'),X64),s('type.conLang.exp',X65))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X59),s('type.num.num',X60))),s('type.list.list'('type.string.char'),X61),s('type.conLang.exp',X62))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X63) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.ind_type.FCONS_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X58),s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60)))))) )),
  inference(superposition,[],[f399,f324])).
fof(f11537,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X51),s('type.list.list'('type.string.char'),X53),s('type.conLang.exp',X54))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X55),s('type.list.list'('type.string.char'),X56),s('type.conLang.exp',X57))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X55) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.ind_type.FCONS_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X51),s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X52),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f399,f316])).
fof(f11536,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X49] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X45),s('type.list.list'('type.string.char'),X46),s('type.conLang.exp',X47))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X48),s('type.list.list'('type.string.char'),X49),s('type.conLang.exp',X50))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X48) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.list.list'(X43),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),sK18(X44,X45,X43,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),s('type.list.list'(X43),'const.list.NIL_0')))) )),
  inference(superposition,[],[f399,f357])).
fof(f11535,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X42,X40] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X40),s('type.list.list'('type.string.char'),X41),s('type.conLang.exp',X42))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun(X34,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X33,fun(X34,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X35),s(X33,X36))),s(X34,X37))),s('type.list.list'('type.string.char'),X38),s('type.conLang.exp',X39))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X40) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.pair.prod'(X33,X34),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),sK19(X35,X34,X33,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),s('type.pair.prod'(X33,X34),'const.pair.,_2'(s(X33,X36),s(X34,X37)))))) )),
  inference(superposition,[],[f399,f367])).
fof(f11534,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X30),s('type.list.list'('type.string.char'),X31),s('type.conLang.exp',X32))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X26),X27))),s('type.list.list'('type.string.char'),X28),s('type.conLang.exp',X29))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X30) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X26),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X26),X27)))) )),
  inference(superposition,[],[f399,f338])).
fof(f11533,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.conLang.exp',X25))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun(X17,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.list.list'('type.string.char'),X21),s('type.conLang.exp',X22))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X23) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun(X16,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'_dst_x0x1_2'(s(fun(X17,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f399,f369])).
fof(f11532,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.conLang.exp',X15))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun(X8,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X7),s(fun(X8,X6),X9))),s(X8,X10))),s('type.list.list'('type.string.char'),X11),s('type.conLang.exp',X12))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X13) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun(X6,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f399,f369])).
fof(f11531,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'_i_0_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0)))) )),
  inference(superposition,[],[f399,f309])).
fof(f11387,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2229,f445])).
fof(f11226,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f2225,f445])).
fof(f10709,plain,(
  ( ! [X2,X0,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X1))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2)) )),
  inference(equality_resolution,[],[f396])).
fof(f10685,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.list.list'('type.conSem.v'),X17) != s('type.list.list'('type.conSem.v'),happ(s(fun(X13,'type.list.list'('type.conSem.v')),X14),s(X13,happ(s(fun(X12,X13),X15),s(X12,X16))))) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X19),s('type.list.list'('type.conSem.v'),X17))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.conSem.v'),happ(s(fun(X12,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X13,'type.list.list'('type.conSem.v')),X14),s(fun(X12,X13),X15))),s(X12,X16))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X19)) )),
  inference(superposition,[],[f396,f369])).
fof(f10684,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.list.list'('type.conSem.v'),X9) != s('type.list.list'('type.conSem.v'),happ(s(fun(X6,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X4,'type.list.list'('type.conSem.v')),X5),s(fun(X6,X4),X7))),s(X6,X8))) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11),s('type.list.list'('type.conSem.v'),X9))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X10),s('type.list.list'('type.conSem.v'),happ(s(fun(X4,'type.list.list'('type.conSem.v')),X5),s(X4,happ(s(fun(X6,X4),X7),s(X6,X8))))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X10) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11)) )),
  inference(superposition,[],[f396,f369])).
fof(f10797,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),happ(s(fun(X1,'type.list.list'('type.conSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),happ(s(fun(X3,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.conSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(forward_demodulation,[],[f10593,f342])).
fof(f10593,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),happ(s(fun(X1,'type.list.list'('type.conSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.SND_1'(s('type.pair.prod'(X6,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'const.pair.,_2'(s(X6,X7),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X3,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.conSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f369,f342,f396])).
fof(f10775,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X4,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f10615,f338])).
fof(f10615,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'const.pair.FST_0'),s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X4,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f369,f338,f396])).
fof(f10753,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X4,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f10637,f324])).
fof(f10637,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.conSem.v'),happ(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X7),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X4,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.conSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f369,f324,f396])).
fof(f10648,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.conSem.v')),X7),s(fun(X5,X6),X8))),s(X5,X9))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X6,'type.list.list'('type.conSem.v')),X7),s(X6,happ(s(fun(X5,X6),X8),s(X5,X9)))))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f10735,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),X5))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),X5)))) )),
  inference(forward_demodulation,[],[f10656,f342])).
fof(f10656,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),X5))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X6,'type.list.list'('type.conSem.v')),'const.pair.,_2'(s(X6,X7),s('type.list.list'('type.conSem.v'),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f342,f369,f396])).
fof(f10734,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.conSem.v'),X5),X6))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.conSem.v'),X5),X6)))))) )),
  inference(forward_demodulation,[],[f10657,f338])).
fof(f10657,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.conSem.v'),X5),X6))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.conSem.v'),X5),'type.list.list'('type.conSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.conSem.v'),X5),X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f369,f396])).
fof(f10733,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun('type.num.num','type.list.list'('type.conSem.v')),X5),s('type.num.num',X6))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun('type.num.num','type.list.list'('type.conSem.v')),X5),s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f10658,f324])).
fof(f10658,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun('type.num.num','type.list.list'('type.conSem.v')),X5),s('type.num.num',X6))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.conSem.v'),X7),s(fun('type.num.num','type.list.list'('type.conSem.v')),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(unit_resulting_resolution,[],[f324,f369,f396])).
fof(f10659,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X7,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f10732,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.conSem.v'))),X7),s(X6,X8))),s(X5,X9))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.conSem.v'))),X7),s(X6,X8))),s(X5,X9)))))) )),
  inference(forward_demodulation,[],[f10660,f367])).
fof(f10660,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.conSem.v'))),X7),s(X6,X8))),s(X5,X9))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun('type.pair.prod'(X6,X5),'type.list.list'('type.conSem.v')),sK19(X7,X5,X6,'type.list.list'('type.conSem.v'))),s('type.pair.prod'(X6,X5),'const.pair.,_2'(s(X6,X8),s(X5,X9)))))))) )),
  inference(unit_resulting_resolution,[],[f367,f369,f396])).
fof(f10711,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X7,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(forward_demodulation,[],[f10681,f367])).
fof(f10681,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.conSem.v'),happ(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'(X1,X0),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),sK19(X2,X0,X1,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X4))))),s('type.list.list'('type.conSem.v'),happ(s(fun(X7,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.conSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f367,f396])).
fof(f10148,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f10131])).
fof(f10131,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f10049,f445])).
fof(f10018,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X161),'const.sum.INL_1'(s(bool,X159))))) | s(bool,'T_0') = s(bool,sK4(X159,X160,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X160)))) )),
  inference(superposition,[],[f6480,f836])).
fof(f10017,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X158),'const.sum.INL_1'(s(bool,X157))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))))) )),
  inference(superposition,[],[f6480,f342])).
fof(f10016,plain,(
  ( ! [X154,X152,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X154),'const.sum.INL_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X153)))))) )),
  inference(superposition,[],[f6480,f336])).
fof(f10015,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X151),'const.sum.INL_1'(s(bool,X149))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X149),s(X148,X150)))))) )),
  inference(superposition,[],[f6480,f341])).
fof(f10014,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X147),'const.sum.INL_1'(s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X146)))))) )),
  inference(superposition,[],[f6480,f337])).
fof(f10012,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X140),'const.sum.INL_1'(s(bool,X138))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6480,f316])).
fof(f10011,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X136))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X134),bool),sK18(X135,X136,X134,bool)),s('type.list.list'(X134),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6480,f357])).
fof(f10049,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9959,f8454])).
fof(f9959,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8610,f6480])).
fof(f9884,plain,(
  ( ! [X872,X866,X864,X870,X868,X873,X867,X865,X871,X869] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X869),s('type.list.list'('type.string.char'),X870),s('type.conLang.exp',X871))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X872),s('type.list.list'('type.string.char'),X873),s('type.conLang.exp',happ(s(fun(X865,'type.conLang.exp'),happ(s(fun(X864,fun(X865,'type.conLang.exp')),X866),s(X864,X867))),s(X865,X868))))) | s('type.conLang.exp',X871) = s('type.conLang.exp',happ(s(fun('type.pair.prod'(X864,X865),'type.conLang.exp'),sK19(X866,X865,X864,'type.conLang.exp')),s('type.pair.prod'(X864,X865),'const.pair.,_2'(s(X864,X867),s(X865,X868)))))) )),
  inference(superposition,[],[f401,f367])).
fof(f9878,plain,(
  ( ! [X810,X808,X812,X806,X804,X811,X809,X813,X807,X805] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X809),s('type.list.list'('type.string.char'),X810),s('type.conLang.exp',X811))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X812),s('type.list.list'('type.string.char'),happ(s(fun(X805,'type.list.list'('type.string.char')),happ(s(fun(X804,fun(X805,'type.list.list'('type.string.char'))),X806),s(X804,X807))),s(X805,X808))),s('type.conLang.exp',X813))) | s('type.list.list'('type.string.char'),X810) = s('type.list.list'('type.string.char'),happ(s(fun('type.pair.prod'(X804,X805),'type.list.list'('type.string.char')),sK19(X806,X805,X804,'type.list.list'('type.string.char'))),s('type.pair.prod'(X804,X805),'const.pair.,_2'(s(X804,X807),s(X805,X808)))))) )),
  inference(superposition,[],[f400,f367])).
fof(f9872,plain,(
  ( ! [X753,X755,X749,X751,X752,X754,X748,X750] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X754),s('type.list.list'('type.conSem.v'),X755))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X753),s('type.list.list'('type.conSem.v'),happ(s(fun(X749,'type.list.list'('type.conSem.v')),happ(s(fun(X748,fun(X749,'type.list.list'('type.conSem.v'))),X750),s(X748,X751))),s(X749,X752))))) | s('type.list.list'('type.conSem.v'),X755) = s('type.list.list'('type.conSem.v'),happ(s(fun('type.pair.prod'(X748,X749),'type.list.list'('type.conSem.v')),sK19(X750,X749,X748,'type.list.list'('type.conSem.v'))),s('type.pair.prod'(X748,X749),'const.pair.,_2'(s(X748,X751),s(X749,X752)))))) )),
  inference(superposition,[],[f395,f367])).
fof(f9863,plain,(
  ( ! [X689,X685,X687,X683,X688,X690,X684,X686] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X688),s('type.list.list'('type.conSem.v'),X689))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X684,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X683,fun(X684,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),X685),s(X683,X686))),s(X684,X687))),s('type.list.list'('type.conSem.v'),X690))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X688) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'(X683,X684),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),sK19(X685,X684,X683,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),s('type.pair.prod'(X683,X684),'const.pair.,_2'(s(X683,X686),s(X684,X687)))))) )),
  inference(superposition,[],[f394,f367])).
fof(f9858,plain,(
  ( ! [X656,X658,X654,X657,X653,X655] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X658))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X654,'type.ast.lit'),happ(s(fun(X653,fun(X654,'type.ast.lit')),X655),s(X653,X656))),s(X654,X657))))) | s('type.ast.lit',X658) = s('type.ast.lit',happ(s(fun('type.pair.prod'(X653,X654),'type.ast.lit'),sK19(X655,X654,X653,'type.ast.lit')),s('type.pair.prod'(X653,X654),'const.pair.,_2'(s(X653,X656),s(X654,X657)))))) )),
  inference(superposition,[],[f334,f367])).
fof(f9948,plain,(
  ( ! [X532,X534,X528,X530,X533,X529,X531,X527] : (s(fun(X527,X528),X534) = s(fun(X527,X528),happ(s(fun(X530,fun(X527,X528)),happ(s(fun(X529,fun(X530,fun(X527,X528))),X531),s(X529,X532))),s(X530,X533))) | s(X528,happ(s(fun(X527,X528),X534),s(X527,sK15(happ(s(fun('type.pair.prod'(X529,X530),fun(X527,X528)),sK19(X531,X530,X529,fun(X527,X528))),s('type.pair.prod'(X529,X530),'const.pair.,_2'(s(X529,X532),s(X530,X533)))),X534,X527,X528)))) != s(X528,happ(s(fun(X527,X528),happ(s(fun(X530,fun(X527,X528)),happ(s(fun(X529,fun(X530,fun(X527,X528))),X531),s(X529,X532))),s(X530,X533))),s(X527,sK15(happ(s(fun('type.pair.prod'(X529,X530),fun(X527,X528)),sK19(X531,X530,X529,fun(X527,X528))),s('type.pair.prod'(X529,X530),'const.pair.,_2'(s(X529,X532),s(X530,X533)))),X534,X527,X528))))) )),
  inference(forward_demodulation,[],[f9842,f367])).
fof(f9842,plain,(
  ( ! [X532,X534,X528,X530,X533,X529,X531,X527] : (s(X528,happ(s(fun(X527,X528),X534),s(X527,sK15(happ(s(fun('type.pair.prod'(X529,X530),fun(X527,X528)),sK19(X531,X530,X529,fun(X527,X528))),s('type.pair.prod'(X529,X530),'const.pair.,_2'(s(X529,X532),s(X530,X533)))),X534,X527,X528)))) != s(X528,happ(s(fun(X527,X528),happ(s(fun(X530,fun(X527,X528)),happ(s(fun(X529,fun(X530,fun(X527,X528))),X531),s(X529,X532))),s(X530,X533))),s(X527,sK15(happ(s(fun('type.pair.prod'(X529,X530),fun(X527,X528)),sK19(X531,X530,X529,fun(X527,X528))),s('type.pair.prod'(X529,X530),'const.pair.,_2'(s(X529,X532),s(X530,X533)))),X534,X527,X528)))) | s(fun(X527,X528),X534) = s(fun(X527,X528),happ(s(fun('type.pair.prod'(X529,X530),fun(X527,X528)),sK19(X531,X530,X529,fun(X527,X528))),s('type.pair.prod'(X529,X530),'const.pair.,_2'(s(X529,X532),s(X530,X533)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f9947,plain,(
  ( ! [X524,X526,X520,X522,X525,X521,X523,X519] : (s(fun(X519,X520),X526) = s(fun(X519,X520),happ(s(fun(X522,fun(X519,X520)),happ(s(fun(X521,fun(X522,fun(X519,X520))),X523),s(X521,X524))),s(X522,X525))) | s(X520,happ(s(fun(X519,X520),X526),s(X519,sK15(X526,happ(s(fun('type.pair.prod'(X521,X522),fun(X519,X520)),sK19(X523,X522,X521,fun(X519,X520))),s('type.pair.prod'(X521,X522),'const.pair.,_2'(s(X521,X524),s(X522,X525)))),X519,X520)))) != s(X520,happ(s(fun(X519,X520),happ(s(fun(X522,fun(X519,X520)),happ(s(fun(X521,fun(X522,fun(X519,X520))),X523),s(X521,X524))),s(X522,X525))),s(X519,sK15(X526,happ(s(fun('type.pair.prod'(X521,X522),fun(X519,X520)),sK19(X523,X522,X521,fun(X519,X520))),s('type.pair.prod'(X521,X522),'const.pair.,_2'(s(X521,X524),s(X522,X525)))),X519,X520))))) )),
  inference(forward_demodulation,[],[f9841,f367])).
fof(f9841,plain,(
  ( ! [X524,X526,X520,X522,X525,X521,X523,X519] : (s(X520,happ(s(fun(X519,X520),X526),s(X519,sK15(X526,happ(s(fun('type.pair.prod'(X521,X522),fun(X519,X520)),sK19(X523,X522,X521,fun(X519,X520))),s('type.pair.prod'(X521,X522),'const.pair.,_2'(s(X521,X524),s(X522,X525)))),X519,X520)))) != s(X520,happ(s(fun(X519,X520),happ(s(fun(X522,fun(X519,X520)),happ(s(fun(X521,fun(X522,fun(X519,X520))),X523),s(X521,X524))),s(X522,X525))),s(X519,sK15(X526,happ(s(fun('type.pair.prod'(X521,X522),fun(X519,X520)),sK19(X523,X522,X521,fun(X519,X520))),s('type.pair.prod'(X521,X522),'const.pair.,_2'(s(X521,X524),s(X522,X525)))),X519,X520)))) | s(fun(X519,X520),X526) = s(fun(X519,X520),happ(s(fun('type.pair.prod'(X521,X522),fun(X519,X520)),sK19(X523,X522,X521,fun(X519,X520))),s('type.pair.prod'(X521,X522),'const.pair.,_2'(s(X521,X524),s(X522,X525)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f9761,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))) = s(X6,happ(s(fun('type.pair.prod'(X1,X2),X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(superposition,[],[f369,f367])).
fof(f9561,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f9544])).
fof(f9544,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f9466,f445])).
fof(f9435,plain,(
  ( ! [X154,X155,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X155,bool),'const.sum.INR_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,sK4(X153,X154,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154)))) )),
  inference(superposition,[],[f6219,f836])).
fof(f9434,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X152,bool),'const.sum.INR_1'(s(bool,X151))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X149,bool),'const.pair.,_2'(s(X149,X150),s(bool,X151)))))) )),
  inference(superposition,[],[f6219,f342])).
fof(f9433,plain,(
  ( ! [X146,X147,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X148,bool),'const.sum.INR_1'(s(bool,X147))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X146),'const.sum.INL_1'(s(bool,X147)))))) )),
  inference(superposition,[],[f6219,f336])).
fof(f9432,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X143))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,X143),s(X142,X144)))))) )),
  inference(superposition,[],[f6219,f341])).
fof(f9431,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X141,bool),'const.sum.INR_1'(s(bool,X140))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X139,bool),'const.sum.INR_1'(s(bool,X140)))))) )),
  inference(superposition,[],[f6219,f337])).
fof(f9429,plain,(
  ( ! [X134,X132,X133] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X134,bool),'const.sum.INR_1'(s(bool,X132))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X132),s(fun('type.num.num',bool),X133),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6219,f316])).
fof(f9428,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,X130))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X128),bool),sK18(X129,X130,X128,bool)),s('type.list.list'(X128),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6219,f357])).
fof(f9466,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9377,f8454])).
fof(f9377,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8610,f6219])).
fof(f9323,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X88),s('type.list.list'('type.string.char'),X86),s('type.conLang.exp',X89))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.conLang.exp',X92))) | s('type.list.list'('type.string.char'),X91) = s('type.list.list'('type.string.char'),sK4(X86,X87,'type.list.list'('type.string.char'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.string.char'),bool),X87)))) )),
  inference(superposition,[],[f400,f836])).
fof(f9322,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X81),s('type.list.list'('type.string.char'),X80),s('type.conLang.exp',X82))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X83),s('type.list.list'('type.string.char'),X84),s('type.conLang.exp',X85))) | s('type.list.list'('type.string.char'),X84) = s('type.list.list'('type.string.char'),'const.pair.SND_1'(s('type.pair.prod'(X78,'type.list.list'('type.string.char')),'const.pair.,_2'(s(X78,X79),s('type.list.list'('type.string.char'),X80)))))) )),
  inference(superposition,[],[f400,f342])).
fof(f9321,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X73),s('type.list.list'('type.string.char'),X72),s('type.conLang.exp',X74))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X75),s('type.list.list'('type.string.char'),X76),s('type.conLang.exp',X77))) | s('type.list.list'('type.string.char'),X76) = s('type.list.list'('type.string.char'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.string.char'),X71),'const.sum.INL_1'(s('type.list.list'('type.string.char'),X72)))))) )),
  inference(superposition,[],[f400,f336])).
fof(f9320,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X66),s('type.list.list'('type.string.char'),X64),s('type.conLang.exp',X67))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X68),s('type.list.list'('type.string.char'),X69),s('type.conLang.exp',X70))) | s('type.list.list'('type.string.char'),X69) = s('type.list.list'('type.string.char'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.string.char'),X63),'const.pair.,_2'(s('type.list.list'('type.string.char'),X64),s(X63,X65)))))) )),
  inference(superposition,[],[f400,f341])).
fof(f9319,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X58),s('type.list.list'('type.string.char'),X57),s('type.conLang.exp',X59))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X60),s('type.list.list'('type.string.char'),X61),s('type.conLang.exp',X62))) | s('type.list.list'('type.string.char'),X61) = s('type.list.list'('type.string.char'),'const.sum.OUTR_1'(s('type.sum.sum'(X56,'type.list.list'('type.string.char')),'const.sum.INR_1'(s('type.list.list'('type.string.char'),X57)))))) )),
  inference(superposition,[],[f400,f337])).
fof(f9318,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X53),s('type.list.list'('type.string.char'),X54),s('type.conLang.exp',X55))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X51),s('type.list.list'('type.string.char'),happ(s(fun('type.num.num','type.list.list'('type.string.char')),X49),s('type.num.num',X50))),s('type.conLang.exp',X52))) | s('type.list.list'('type.string.char'),X54) = s('type.list.list'('type.string.char'),'const.ind_type.FCONS_3'(s('type.list.list'('type.string.char'),X48),s(fun('type.num.num','type.list.list'('type.string.char')),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f400,f324])).
fof(f9317,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X43),s('type.list.list'('type.string.char'),X41),s('type.conLang.exp',X44))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X45),s('type.list.list'('type.string.char'),X46),s('type.conLang.exp',X47))) | s('type.list.list'('type.string.char'),X46) = s('type.list.list'('type.string.char'),'const.ind_type.FCONS_3'(s('type.list.list'('type.string.char'),X41),s(fun('type.num.num','type.list.list'('type.string.char')),X42),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f400,f316])).
fof(f9316,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X36),s('type.list.list'('type.string.char'),X35),s('type.conLang.exp',X37))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X38),s('type.list.list'('type.string.char'),X39),s('type.conLang.exp',X40))) | s('type.list.list'('type.string.char'),X39) = s('type.list.list'('type.string.char'),happ(s(fun('type.list.list'(X33),'type.list.list'('type.string.char')),sK18(X34,X35,X33,'type.list.list'('type.string.char'))),s('type.list.list'(X33),'const.list.NIL_0')))) )),
  inference(superposition,[],[f400,f357])).
fof(f9315,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X30),s('type.list.list'('type.string.char'),X31),s('type.conLang.exp',X32))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.string.char'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.string.char'),X26),X27))),s('type.conLang.exp',X29))) | s('type.list.list'('type.string.char'),X31) = s('type.list.list'('type.string.char'),happ(s(fun('type.pair.prod'('type.list.list'('type.string.char'),X26),'type.list.list'('type.string.char')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.string.char'),X26),X27)))) )),
  inference(superposition,[],[f400,f338])).
fof(f9314,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.conLang.exp',X25))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X21),s('type.list.list'('type.string.char'),happ(s(fun(X17,'type.list.list'('type.string.char')),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.conLang.exp',X22))) | s('type.list.list'('type.string.char'),X24) = s('type.list.list'('type.string.char'),happ(s(fun(X16,'type.list.list'('type.string.char')),'_dst_x0x1_2'(s(fun(X17,'type.list.list'('type.string.char')),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f400,f369])).
fof(f9313,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.conLang.exp',X15))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),happ(s(fun(X8,'type.list.list'('type.string.char')),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.string.char')),X7),s(fun(X8,X6),X9))),s(X8,X10))),s('type.conLang.exp',X12))) | s('type.list.list'('type.string.char'),X14) = s('type.list.list'('type.string.char'),happ(s(fun(X6,'type.list.list'('type.string.char')),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f400,f369])).
fof(f9236,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) != s(bool,sK7('const.pair.FST_0',X153,bool)) | s(bool,'F_0') = s(bool,sK4(X151,X152,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X152)))) ) | $spl26),
  inference(superposition,[],[f8620,f836])).
fof(f9235,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X147,bool),'const.pair.,_2'(s(X147,X148),s(bool,X149)))))) ) | $spl26),
  inference(superposition,[],[f8620,f342])).
fof(f9234,plain,(
  ( ! [X146,X144,X145] : (s(bool,X145) != s(bool,sK7('const.pair.FST_0',X146,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,X145)))))) ) | $spl26),
  inference(superposition,[],[f8620,f336])).
fof(f9233,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X141) != s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,X141),s(X140,X142)))))) ) | $spl26),
  inference(superposition,[],[f8620,f341])).
fof(f9232,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) != s(bool,sK7('const.pair.FST_0',X139,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,X138)))))) ) | $spl26),
  inference(superposition,[],[f8620,f337])).
fof(f9230,plain,(
  ( ! [X132,X130,X131] : (s(bool,X130) != s(bool,sK7('const.pair.FST_0',X132,bool)) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X130),s(fun('type.num.num',bool),X131),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8620,f316])).
fof(f9229,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,X128) != s(bool,sK7('const.pair.FST_0',X129,bool)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X126),bool),sK18(X127,X128,X126,bool)),s('type.list.list'(X126),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8620,f357])).
fof(f9225,plain,(
  ( ! [X109,X110] : (s(bool,X109) != s(bool,sK7('const.pair.FST_0',X110,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X109)))) ) | $spl26),
  inference(superposition,[],[f8620,f309])).
fof(f8852,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) = s(bool,sK7('const.pair.FST_0',X153,bool)) | s(bool,'T_0') = s(bool,sK4(X151,X152,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X152)))) ) | $spl26),
  inference(superposition,[],[f8609,f836])).
fof(f8851,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) = s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X147,bool),'const.pair.,_2'(s(X147,X148),s(bool,X149)))))) ) | $spl26),
  inference(superposition,[],[f8609,f342])).
fof(f8850,plain,(
  ( ! [X146,X144,X145] : (s(bool,X145) = s(bool,sK7('const.pair.FST_0',X146,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,X145)))))) ) | $spl26),
  inference(superposition,[],[f8609,f336])).
fof(f8849,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X141) = s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,X141),s(X140,X142)))))) ) | $spl26),
  inference(superposition,[],[f8609,f341])).
fof(f8848,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) = s(bool,sK7('const.pair.FST_0',X139,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,X138)))))) ) | $spl26),
  inference(superposition,[],[f8609,f337])).
fof(f8846,plain,(
  ( ! [X132,X130,X131] : (s(bool,X130) = s(bool,sK7('const.pair.FST_0',X132,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X130),s(fun('type.num.num',bool),X131),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8609,f316])).
fof(f8845,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,X128) = s(bool,sK7('const.pair.FST_0',X129,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X126),bool),sK18(X127,X128,X126,bool)),s('type.list.list'(X126),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8609,f357])).
fof(f8841,plain,(
  ( ! [X109,X110] : (s(bool,X109) = s(bool,sK7('const.pair.FST_0',X110,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X109)))) ) | $spl26),
  inference(superposition,[],[f8609,f309])).
fof(f8795,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f8761,f262])).
fof(f8761,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f350,f445])).
fof(f8620,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(subsumption_resolution,[],[f8567,f264])).
fof(f8567,plain,(
  ( ! [X41,X40] : (p(s(bool,'F_0')) | s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(superposition,[],[f2426,f8454])).
fof(f8610,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(subsumption_resolution,[],[f8544,f264])).
fof(f8544,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(superposition,[],[f271,f8454])).
fof(f8609,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(subsumption_resolution,[],[f8532,f432])).
fof(f8532,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f8454,f445])).
fof(f8521,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f8486,f262])).
fof(f8486,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f349,f445])).
fof(f8520,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(subsumption_resolution,[],[f8484,f262])).
fof(f8484,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f349,f440])).
fof(f8104,plain,(
  ( ! [X14,X13] : (s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X14),s('type.list.list'('type.conSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13)),X14,'type.list.list'('type.conSem.v'),'type.conSem.v')))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.conSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13)),X14,'type.list.list'('type.conSem.v'),'type.conSem.v')))) | s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X14) = s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13)))) )),
  inference(superposition,[],[f353,f329])).
fof(f8103,plain,(
  ( ! [X12,X11] : (s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X12),s('type.list.list'('type.conSem.v'),sK15(X12,happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11)),'type.list.list'('type.conSem.v'),'type.conSem.v')))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11),s('type.list.list'('type.conSem.v'),sK15(X12,happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11)),'type.list.list'('type.conSem.v'),'type.conSem.v')))) | s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X12) = s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X11)))) )),
  inference(superposition,[],[f353,f329])).
fof(f7594,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X2),s('type.list.list'('type.conSem.v'),X3))))) = s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v',X2))),s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),X3)))))))))) )),
  inference(superposition,[],[f328,f269])).
fof(f7065,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X88),s('type.list.list'('type.string.char'),X89),s('type.conLang.exp',X86))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.conLang.exp',X92))) | s('type.conLang.exp',X92) = s('type.conLang.exp',sK4(X86,X87,'type.conLang.exp')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.conLang.exp',bool),X87)))) )),
  inference(superposition,[],[f401,f836])).
fof(f7064,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X81),s('type.list.list'('type.string.char'),X82),s('type.conLang.exp',X80))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X83),s('type.list.list'('type.string.char'),X84),s('type.conLang.exp',X85))) | s('type.conLang.exp',X85) = s('type.conLang.exp','const.pair.SND_1'(s('type.pair.prod'(X78,'type.conLang.exp'),'const.pair.,_2'(s(X78,X79),s('type.conLang.exp',X80)))))) )),
  inference(superposition,[],[f401,f342])).
fof(f7063,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X73),s('type.list.list'('type.string.char'),X74),s('type.conLang.exp',X72))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X75),s('type.list.list'('type.string.char'),X76),s('type.conLang.exp',X77))) | s('type.conLang.exp',X77) = s('type.conLang.exp','const.sum.OUTL_1'(s('type.sum.sum'('type.conLang.exp',X71),'const.sum.INL_1'(s('type.conLang.exp',X72)))))) )),
  inference(superposition,[],[f401,f336])).
fof(f7062,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X66),s('type.list.list'('type.string.char'),X67),s('type.conLang.exp',X64))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X68),s('type.list.list'('type.string.char'),X69),s('type.conLang.exp',X70))) | s('type.conLang.exp',X70) = s('type.conLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.conLang.exp',X63),'const.pair.,_2'(s('type.conLang.exp',X64),s(X63,X65)))))) )),
  inference(superposition,[],[f401,f341])).
fof(f7061,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X58),s('type.list.list'('type.string.char'),X59),s('type.conLang.exp',X57))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X60),s('type.list.list'('type.string.char'),X61),s('type.conLang.exp',X62))) | s('type.conLang.exp',X62) = s('type.conLang.exp','const.sum.OUTR_1'(s('type.sum.sum'(X56,'type.conLang.exp'),'const.sum.INR_1'(s('type.conLang.exp',X57)))))) )),
  inference(superposition,[],[f401,f337])).
fof(f7060,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X53),s('type.list.list'('type.string.char'),X54),s('type.conLang.exp',X55))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X51),s('type.list.list'('type.string.char'),X52),s('type.conLang.exp',happ(s(fun('type.num.num','type.conLang.exp'),X49),s('type.num.num',X50))))) | s('type.conLang.exp',X55) = s('type.conLang.exp','const.ind_type.FCONS_3'(s('type.conLang.exp',X48),s(fun('type.num.num','type.conLang.exp'),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f401,f324])).
fof(f7059,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X43),s('type.list.list'('type.string.char'),X44),s('type.conLang.exp',X41))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X45),s('type.list.list'('type.string.char'),X46),s('type.conLang.exp',X47))) | s('type.conLang.exp',X47) = s('type.conLang.exp','const.ind_type.FCONS_3'(s('type.conLang.exp',X41),s(fun('type.num.num','type.conLang.exp'),X42),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f401,f316])).
fof(f7058,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X36),s('type.list.list'('type.string.char'),X37),s('type.conLang.exp',X35))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X38),s('type.list.list'('type.string.char'),X39),s('type.conLang.exp',X40))) | s('type.conLang.exp',X40) = s('type.conLang.exp',happ(s(fun('type.list.list'(X33),'type.conLang.exp'),sK18(X34,X35,X33,'type.conLang.exp')),s('type.list.list'(X33),'const.list.NIL_0')))) )),
  inference(superposition,[],[f401,f357])).
fof(f7057,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X30),s('type.list.list'('type.string.char'),X31),s('type.conLang.exp',X32))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.string.char'),X29),s('type.conLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.conLang.exp',X26),X27))))) | s('type.conLang.exp',X32) = s('type.conLang.exp',happ(s(fun('type.pair.prod'('type.conLang.exp',X26),'type.conLang.exp'),'const.pair.FST_0'),s('type.pair.prod'('type.conLang.exp',X26),X27)))) )),
  inference(superposition,[],[f401,f338])).
fof(f7056,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.conLang.exp',X25))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X21),s('type.list.list'('type.string.char'),X22),s('type.conLang.exp',happ(s(fun(X17,'type.conLang.exp'),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))))) | s('type.conLang.exp',X25) = s('type.conLang.exp',happ(s(fun(X16,'type.conLang.exp'),'_dst_x0x1_2'(s(fun(X17,'type.conLang.exp'),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f401,f369])).
fof(f7055,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.conLang.exp',X15))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',happ(s(fun(X8,'type.conLang.exp'),'_dst_x0x1_2'(s(fun(X6,'type.conLang.exp'),X7),s(fun(X8,X6),X9))),s(X8,X10))))) | s('type.conLang.exp',X15) = s('type.conLang.exp',happ(s(fun(X6,'type.conLang.exp'),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f401,f369])).
fof(f6851,plain,(
  ( ! [X68] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK15('const.conSem.Vectorv_0',X68,'type.list.list'('type.conSem.v'),'type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X68),s('type.list.list'('type.conSem.v'),sK15('const.conSem.Vectorv_0',X68,'type.list.list'('type.conSem.v'),'type.conSem.v')))) | s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.conSem.Vectorv_0') = s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f6850,plain,(
  ( ! [X67] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK15('const.conSem.Litv_0',X67,'type.ast.lit','type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.ast.lit','type.conSem.v'),X67),s('type.ast.lit',sK15('const.conSem.Litv_0',X67,'type.ast.lit','type.conSem.v')))) | s(fun('type.ast.lit','type.conSem.v'),'const.conSem.Litv_0') = s(fun('type.ast.lit','type.conSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f6849,plain,(
  ( ! [X66] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK15('const.conSem.Loc_0',X66,'type.num.num','type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.num.num','type.conSem.v'),X66),s('type.num.num',sK15('const.conSem.Loc_0',X66,'type.num.num','type.conSem.v')))) | s(fun('type.num.num','type.conSem.v'),'const.conSem.Loc_0') = s(fun('type.num.num','type.conSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f6848,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6837,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X20,X21,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,X21,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f6836,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f6833,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f6832,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6831,plain,(
  ( ! [X0,X1] : (s(X0,sK15('_i_0_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15('_i_0_0',X1,X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f6820,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,X9))) )),
  inference(superposition,[],[f353,f445])).
fof(f6819,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK15(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f353,f445])).
fof(f6876,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65))))) )),
  inference(forward_demodulation,[],[f6816,f342])).
fof(f6816,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f6875,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60))))) )),
  inference(forward_demodulation,[],[f6815,f336])).
fof(f6815,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f6874,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54))))) )),
  inference(forward_demodulation,[],[f6814,f341])).
fof(f6814,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f6873,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49))))) )),
  inference(forward_demodulation,[],[f6813,f337])).
fof(f6813,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f6872,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43))))) )),
  inference(forward_demodulation,[],[f6812,f324])).
fof(f6812,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f6871,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38))))) )),
  inference(forward_demodulation,[],[f6811,f316])).
fof(f6811,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f6870,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(fun(X31,X32),X35) = s(fun(X31,X32),X36) | s(X32,happ(s(fun(X31,X32),X35),s(X31,sK15(happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X36,X31,X32)))) != s(X32,happ(s(fun(X31,X32),X36),s(X31,sK15(happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X36,X31,X32))))) )),
  inference(forward_demodulation,[],[f6810,f357])).
fof(f6810,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X32,happ(s(fun(X31,X32),X35),s(X31,sK15(happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X36,X31,X32)))) != s(X32,happ(s(fun(X31,X32),X36),s(X31,sK15(happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X36,X31,X32)))) | s(fun(X31,X32),X36) = s(fun(X31,X32),happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f6869,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X30,X26,X27)))) != s(X27,happ(s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))),s(X26,sK15(happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X30,X26,X27))))) )),
  inference(forward_demodulation,[],[f6809,f338])).
fof(f6809,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X30,X26,X27)))) != s(X27,happ(s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))),s(X26,sK15(happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X30,X26,X27)))) | s(fun(X26,X27),X30) = s(fun(X26,X27),happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)))) )),
  inference(superposition,[],[f353,f338])).
fof(f6808,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6807,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f6868,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7))))) )),
  inference(forward_demodulation,[],[f6806,f309])).
fof(f6806,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f6867,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1))))) )),
  inference(forward_demodulation,[],[f6805,f368])).
fof(f6805,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6804,plain,(
  ( ! [X68] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK15(X68,'const.conSem.Vectorv_0','type.list.list'('type.conSem.v'),'type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X68),s('type.list.list'('type.conSem.v'),sK15(X68,'const.conSem.Vectorv_0','type.list.list'('type.conSem.v'),'type.conSem.v')))) | s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.conSem.Vectorv_0') = s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f6803,plain,(
  ( ! [X67] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK15(X67,'const.conSem.Litv_0','type.ast.lit','type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.ast.lit','type.conSem.v'),X67),s('type.ast.lit',sK15(X67,'const.conSem.Litv_0','type.ast.lit','type.conSem.v')))) | s(fun('type.ast.lit','type.conSem.v'),'const.conSem.Litv_0') = s(fun('type.ast.lit','type.conSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f6802,plain,(
  ( ! [X66] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK15(X66,'const.conSem.Loc_0','type.num.num','type.conSem.v')))) != s('type.conSem.v',happ(s(fun('type.num.num','type.conSem.v'),X66),s('type.num.num',sK15(X66,'const.conSem.Loc_0','type.num.num','type.conSem.v')))) | s(fun('type.num.num','type.conSem.v'),'const.conSem.Loc_0') = s(fun('type.num.num','type.conSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f6801,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6790,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X21,X20,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X21,X20,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f6789,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f6786,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f6785,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6784,plain,(
  ( ! [X0,X1] : (s(X0,sK15(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f6862,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65))))) )),
  inference(forward_demodulation,[],[f6782,f342])).
fof(f6782,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f6861,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60))))) )),
  inference(forward_demodulation,[],[f6781,f336])).
fof(f6781,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f6860,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54))))) )),
  inference(forward_demodulation,[],[f6780,f341])).
fof(f6780,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f6859,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49))))) )),
  inference(forward_demodulation,[],[f6779,f337])).
fof(f6779,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f6858,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43))))) )),
  inference(forward_demodulation,[],[f6778,f324])).
fof(f6778,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f6857,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38))))) )),
  inference(forward_demodulation,[],[f6777,f316])).
fof(f6777,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f6856,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(fun(X31,X32),X35) = s(fun(X31,X32),X36) | s(X32,happ(s(fun(X31,X32),X35),s(X31,sK15(X36,happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X31,X32)))) != s(X32,happ(s(fun(X31,X32),X36),s(X31,sK15(X36,happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X31,X32))))) )),
  inference(forward_demodulation,[],[f6776,f357])).
fof(f6776,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X32,happ(s(fun(X31,X32),X35),s(X31,sK15(X36,happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X31,X32)))) != s(X32,happ(s(fun(X31,X32),X36),s(X31,sK15(X36,happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')),X31,X32)))) | s(fun(X31,X32),X36) = s(fun(X31,X32),happ(s(fun('type.list.list'(X33),fun(X31,X32)),sK18(X34,X35,X33,fun(X31,X32))),s('type.list.list'(X33),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f6855,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X30,happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X26,X27)))) != s(X27,happ(s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))),s(X26,sK15(X30,happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X26,X27))))) )),
  inference(forward_demodulation,[],[f6775,f338])).
fof(f6775,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X30,happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X26,X27)))) != s(X27,happ(s(fun(X26,X27),'const.pair.FST_1'(s('type.pair.prod'(fun(X26,X27),X28),X29))),s(X26,sK15(X30,happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)),X26,X27)))) | s(fun(X26,X27),X30) = s(fun(X26,X27),happ(s(fun('type.pair.prod'(fun(X26,X27),X28),fun(X26,X27)),'const.pair.FST_0'),s('type.pair.prod'(fun(X26,X27),X28),X29)))) )),
  inference(superposition,[],[f353,f338])).
fof(f6774,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6773,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f6854,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7))))) )),
  inference(forward_demodulation,[],[f6772,f309])).
fof(f6772,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f6853,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) )),
  inference(forward_demodulation,[],[f6771,f368])).
fof(f6771,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6480,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6408,f264])).
fof(f6408,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2421])).
fof(f6219,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6149,f264])).
fof(f6149,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2419])).
fof(f5956,plain,(
  ( ! [X536,X538,X532,X534,X537,X533,X535,X531] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X536),s('type.list.list'('type.conSem.v'),X537))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X538),s('type.list.list'('type.conSem.v'),happ(s(fun(X532,'type.list.list'('type.conSem.v')),X533),s(X532,happ(s(fun(X531,X532),X534),s(X531,X535))))))) | s('type.list.list'('type.conSem.v'),X537) = s('type.list.list'('type.conSem.v'),happ(s(fun(X531,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X532,'type.list.list'('type.conSem.v')),X533),s(fun(X531,X532),X534))),s(X531,X535)))) )),
  inference(superposition,[],[f395,f369])).
fof(f5950,plain,(
  ( ! [X484,X486,X488,X490,X485,X487,X483,X489] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X488),s('type.list.list'('type.conSem.v'),X489))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X484,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X485),s(X484,happ(s(fun(X483,X484),X486),s(X483,X487))))),s('type.list.list'('type.conSem.v'),X490))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X488) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X483,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X484,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X485),s(fun(X483,X484),X486))),s(X483,X487)))) )),
  inference(superposition,[],[f394,f369])).
fof(f5946,plain,(
  ( ! [X461,X463,X459,X460,X462,X464] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X464))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X460,'type.ast.lit'),X461),s(X460,happ(s(fun(X459,X460),X462),s(X459,X463))))))) | s('type.ast.lit',X464) = s('type.ast.lit',happ(s(fun(X459,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X460,'type.ast.lit'),X461),s(fun(X459,X460),X462))),s(X459,X463)))) )),
  inference(superposition,[],[f334,f369])).
fof(f5944,plain,(
  ( ! [X447,X449,X451,X452,X448,X450] : (s('type.ast.lit',X452) != s('type.ast.lit',happ(s(fun(X448,'type.ast.lit'),X449),s(X448,happ(s(fun(X447,X448),X450),s(X447,X451))))) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X452))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X447,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X448,'type.ast.lit'),X449),s(fun(X447,X448),X450))),s(X447,X451)))))) )),
  inference(superposition,[],[f335,f369])).
fof(f5933,plain,(
  ( ! [X358,X364,X360,X362,X357,X359,X365,X361,X363] : (s(fun(X357,X358),X364) != s(fun(X357,X358),happ(s(fun(X360,fun(X357,X358)),X361),s(X360,happ(s(fun(X359,X360),X362),s(X359,X363))))) | s(X358,happ(s(fun(X357,X358),X364),s(X357,X365))) = s(X358,happ(s(fun(X357,X358),happ(s(fun(X359,fun(X357,X358)),'_dst_x0x1_2'(s(fun(X360,fun(X357,X358)),X361),s(fun(X359,X360),X362))),s(X359,X363))),s(X357,X365)))) )),
  inference(superposition,[],[f352,f369])).
fof(f5928,plain,(
  ( ! [X325,X327,X323,X329,X324,X326,X328] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X323,bool),happ(s(fun(X325,fun(X323,bool)),X326),s(X325,happ(s(fun(X324,X325),X327),s(X324,X328))))))) | s(X323,X329) = s(X323,sK4(X329,happ(s(fun(X324,fun(X323,bool)),'_dst_x0x1_2'(s(fun(X325,fun(X323,bool)),X326),s(fun(X324,X325),X327))),s(X324,X328)),X323))) )),
  inference(superposition,[],[f836,f369])).
fof(f5886,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK4(X85,'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84)),X81))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f319,f369])).
fof(f5872,plain,(
  ( ! [X557,X553,X555,X551,X556,X552,X554,X550] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X555),s('type.list.list'('type.conSem.v'),X556))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X557),s('type.list.list'('type.conSem.v'),happ(s(fun(X552,'type.list.list'('type.conSem.v')),'_dst_x0x1_2'(s(fun(X550,'type.list.list'('type.conSem.v')),X551),s(fun(X552,X550),X553))),s(X552,X554))))) | s('type.list.list'('type.conSem.v'),X556) = s('type.list.list'('type.conSem.v'),happ(s(fun(X550,'type.list.list'('type.conSem.v')),X551),s(X550,happ(s(fun(X552,X550),X553),s(X552,X554)))))) )),
  inference(superposition,[],[f395,f369])).
fof(f5866,plain,(
  ( ! [X509,X505,X507,X502,X508,X504,X506,X503] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X507),s('type.list.list'('type.conSem.v'),X508))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X504,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X502,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X503),s(fun(X504,X502),X505))),s(X504,X506))),s('type.list.list'('type.conSem.v'),X509))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X507) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun(X502,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X503),s(X502,happ(s(fun(X504,X502),X505),s(X504,X506)))))) )),
  inference(superposition,[],[f394,f369])).
fof(f5862,plain,(
  ( ! [X478,X480,X482,X479,X481,X483] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X483))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X480,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X478,'type.ast.lit'),X479),s(fun(X480,X478),X481))),s(X480,X482))))) | s('type.ast.lit',X483) = s('type.ast.lit',happ(s(fun(X478,'type.ast.lit'),X479),s(X478,happ(s(fun(X480,X478),X481),s(X480,X482)))))) )),
  inference(superposition,[],[f334,f369])).
fof(f5860,plain,(
  ( ! [X466,X469,X471,X467,X468,X470] : (s('type.ast.lit',X471) != s('type.ast.lit',happ(s(fun(X468,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X466,'type.ast.lit'),X467),s(fun(X468,X466),X469))),s(X468,X470))) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X471))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X466,'type.ast.lit'),X467),s(X466,happ(s(fun(X468,X466),X469),s(X468,X470)))))))) )),
  inference(superposition,[],[f335,f369])).
fof(f5849,plain,(
  ( ! [X372,X374,X368,X370,X367,X373,X375,X369,X371] : (s(fun(X367,X368),X374) != s(fun(X367,X368),happ(s(fun(X371,fun(X367,X368)),'_dst_x0x1_2'(s(fun(X369,fun(X367,X368)),X370),s(fun(X371,X369),X372))),s(X371,X373))) | s(X368,happ(s(fun(X367,X368),X374),s(X367,X375))) = s(X368,happ(s(fun(X367,X368),happ(s(fun(X369,fun(X367,X368)),X370),s(X369,happ(s(fun(X371,X369),X372),s(X371,X373))))),s(X367,X375)))) )),
  inference(superposition,[],[f352,f369])).
fof(f5844,plain,(
  ( ! [X333,X335,X337,X339,X334,X336,X338] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X333,bool),happ(s(fun(X336,fun(X333,bool)),'_dst_x0x1_2'(s(fun(X334,fun(X333,bool)),X335),s(fun(X336,X334),X337))),s(X336,X338))))) | s(X333,X339) = s(X333,sK4(X339,happ(s(fun(X334,fun(X333,bool)),X335),s(X334,happ(s(fun(X336,X334),X337),s(X336,X338)))),X333))) )),
  inference(superposition,[],[f836,f369])).
fof(f5700,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (s(fun(X101,X102),'_dst_x0x1_2'(s(fun(X103,X102),happ(s(fun(X105,fun(X103,X102)),X106),s(X105,X107))),s(fun(X101,X103),X104))) = s(fun(X101,X102),happ(s(fun(X105,fun(X101,X102)),'_dst_x0x1_2'(s(fun(fun(X103,X102),fun(X101,X102)),'_dstsel_0_x01_1'(s(fun(X101,X103),X104))),s(fun(X105,fun(X103,X102)),X106))),s(X105,X107)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5689,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X39,X37),X40))),s(X39,X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X39,X37),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f369,f445])).
fof(f5688,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X33,X31),X34))),s(X33,X35))) | s(bool,X36) = s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X33,X31),X34),s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f369,f445])).
fof(f5686,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f369,f338])).
fof(f5685,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5684,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(X110,happ(s(fun(fun(X105,X106),X110),X111),s(fun(X105,X106),'_dst_x0x1_2'(s(fun(X107,X106),X109),s(fun(X105,X107),X108))))) = s(X110,happ(s(fun(fun(X107,X106),X110),'_dst_x0x1_2'(s(fun(fun(X105,X106),X110),X111),s(fun(fun(X107,X106),fun(X105,X106)),'_dstsel_0_x01_1'(s(fun(X105,X107),X108))))),s(fun(X107,X106),X109)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5683,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(X29,happ(s(fun(X25,X29),X30),s(X25,X28))) = s(X29,happ(s(fun('type.list.list'(X26),X29),'_dst_x0x1_2'(s(fun(X25,X29),X30),s(fun('type.list.list'(X26),X25),sK18(X27,X28,X26,X25)))),s('type.list.list'(X26),'const.list.NIL_0')))) )),
  inference(superposition,[],[f369,f357])).
fof(f5682,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5681,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f369,f369])).
fof(f5679,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f369,f352])).
fof(f5680,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X0,'type.ast.lit'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun(X2,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X0,'type.ast.lit'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f369,f335])).
fof(f5400,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X94),bool),sK18(X95,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK18(X93,'T_0',X92,bool))),X94,bool))))) )),
  inference(subsumption_resolution,[],[f5364,f262])).
fof(f5364,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X94),bool),sK18(X95,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK18(X93,'T_0',X92,bool))),X94,bool))))) )),
  inference(superposition,[],[f4843,f4872])).
fof(f5394,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f5360,f262])).
fof(f5360,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4843,f836])).
fof(f5393,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5359,f262])).
fof(f5359,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4843,f2070])).
fof(f5389,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(subsumption_resolution,[],[f5356,f262])).
fof(f5356,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(superposition,[],[f4843,f440])).
fof(f5378,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5339,f262])).
fof(f5339,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4843,f445])).
fof(f5304,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X44),s('type.list.list'('type.conSem.v'),X45))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.conSem.v'),X47))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X46) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.SND_1'(s('type.pair.prod'(X42,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'const.pair.,_2'(s(X42,X43),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X44)))))) )),
  inference(superposition,[],[f394,f342])).
fof(f5303,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X38),s('type.list.list'('type.conSem.v'),X39))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.conSem.v'),X41))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X40) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.sum.OUTL_1'(s('type.sum.sum'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37),'const.sum.INL_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X38)))))) )),
  inference(superposition,[],[f394,f336])).
fof(f5302,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X32),s('type.list.list'('type.conSem.v'),X34))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X35),s('type.list.list'('type.conSem.v'),X36))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X35) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X31),'const.pair.,_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f394,f341])).
fof(f5301,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X27),s('type.list.list'('type.conSem.v'),X28))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.conSem.v'),X30))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X29) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'const.sum.INR_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X27)))))) )),
  inference(superposition,[],[f394,f337])).
fof(f5300,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.conSem.v'),X25))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X21),s('type.num.num',X22))),s('type.list.list'('type.conSem.v'),X23))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X24) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X20),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f394,f324])).
fof(f5299,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X15),s('type.list.list'('type.conSem.v'),X17))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.conSem.v'),X19))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X15),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f394,f316])).
fof(f5298,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.conSem.v'),X14))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X10),X11))),s('type.list.list'('type.conSem.v'),X12))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X10),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'const.pair.FST_0'),s('type.pair.prod'('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X10),X11)))) )),
  inference(superposition,[],[f394,f338])).
fof(f5297,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6),s('type.list.list'('type.conSem.v'),X7))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.conSem.v'),X9))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X8) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.list.list'(X4),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),sK18(X5,X6,X4,'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f394,f357])).
fof(f5296,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),'type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))),'_i_0_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0)))) )),
  inference(superposition,[],[f394,f309])).
fof(f5082,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5081,f262])).
fof(f5081,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f5072,f2070])).
fof(f5072,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f4743,f2531])).
fof(f5079,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5067,f262])).
fof(f5067,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4743,f445])).
fof(f5054,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5007,f262])).
fof(f5007,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f447,f4872])).
fof(f4866,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f357,f4829,f447])).
fof(f4872,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f432,f4829,f2424])).
fof(f4873,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4829,f4742])).
fof(f4829,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f262,f4742])).
fof(f4743,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4668,f338])).
fof(f4673,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4623,f262])).
fof(f4623,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f447,f836])).
fof(f4606,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X217),s('type.list.list'('type.conSem.v'),X218))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X219),s('type.list.list'('type.conSem.v'),X215))) | s('type.list.list'('type.conSem.v'),X218) = s('type.list.list'('type.conSem.v'),sK4(X215,X216,'type.list.list'('type.conSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.conSem.v'),bool),X216)))) )),
  inference(superposition,[],[f395,f836])).
fof(f4600,plain,(
  ( ! [X194,X195,X193] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X193))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X195))) | s('type.ast.lit',X195) = s('type.ast.lit',sK4(X193,X194,'type.ast.lit')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X194)))) )),
  inference(superposition,[],[f334,f836])).
fof(f4590,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f319,f836])).
fof(f4669,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(forward_demodulation,[],[f4574,f2070])).
fof(f4574,plain,(
  ( ! [X107,X106] : (s(bool,X106) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(superposition,[],[f2168,f836])).
fof(f4508,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f836,f445])).
fof(f4507,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f836,f445])).
fof(f4483,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X45),s('type.list.list'('type.conSem.v'),X44))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.conSem.v'),X47))) | s('type.list.list'('type.conSem.v'),X47) = s('type.list.list'('type.conSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X42,'type.list.list'('type.conSem.v')),'const.pair.,_2'(s(X42,X43),s('type.list.list'('type.conSem.v'),X44)))))) )),
  inference(superposition,[],[f395,f342])).
fof(f4482,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X39),s('type.list.list'('type.conSem.v'),X38))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.conSem.v'),X41))) | s('type.list.list'('type.conSem.v'),X41) = s('type.list.list'('type.conSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.conSem.v'),X37),'const.sum.INL_1'(s('type.list.list'('type.conSem.v'),X38)))))) )),
  inference(superposition,[],[f395,f336])).
fof(f4481,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X34),s('type.list.list'('type.conSem.v'),X32))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X35),s('type.list.list'('type.conSem.v'),X36))) | s('type.list.list'('type.conSem.v'),X36) = s('type.list.list'('type.conSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.conSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.conSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f395,f341])).
fof(f4480,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X28),s('type.list.list'('type.conSem.v'),X27))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.conSem.v'),X30))) | s('type.list.list'('type.conSem.v'),X30) = s('type.list.list'('type.conSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.list.list'('type.conSem.v')),'const.sum.INR_1'(s('type.list.list'('type.conSem.v'),X27)))))) )),
  inference(superposition,[],[f395,f337])).
fof(f4479,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.conSem.v'),X25))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X23),s('type.list.list'('type.conSem.v'),happ(s(fun('type.num.num','type.list.list'('type.conSem.v')),X21),s('type.num.num',X22))))) | s('type.list.list'('type.conSem.v'),X25) = s('type.list.list'('type.conSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.conSem.v'),X20),s(fun('type.num.num','type.list.list'('type.conSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f395,f324])).
fof(f4478,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.conSem.v'),X15))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.conSem.v'),X19))) | s('type.list.list'('type.conSem.v'),X19) = s('type.list.list'('type.conSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.conSem.v'),X15),s(fun('type.num.num','type.list.list'('type.conSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f395,f316])).
fof(f4477,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.conSem.v'),X14))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X12),s('type.list.list'('type.conSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.conSem.v'),X10),X11))))) | s('type.list.list'('type.conSem.v'),X14) = s('type.list.list'('type.conSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.conSem.v'),X10),'type.list.list'('type.conSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.conSem.v'),X10),X11)))) )),
  inference(superposition,[],[f395,f338])).
fof(f4476,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X7),s('type.list.list'('type.conSem.v'),X6))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.conSem.v'),X9))) | s('type.list.list'('type.conSem.v'),X9) = s('type.list.list'('type.conSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.conSem.v')),sK18(X5,X6,X4,'type.list.list'('type.conSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f395,f357])).
fof(f4475,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | s('type.list.list'('type.conSem.v'),X3) = s('type.list.list'('type.conSem.v'),happ(s(fun('type.list.list'('type.conSem.v'),'type.list.list'('type.conSem.v')),'_i_0_0'),s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(superposition,[],[f395,f309])).
fof(f4107,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3633,f445])).
fof(f3633,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f432,f2803])).
fof(f3415,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2803,f2803])).
fof(f3322,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK14(X7,X8,X9,bool))) )),
  inference(superposition,[],[f351,f445])).
fof(f3321,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK14(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f2803,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2531,f2531])).
fof(f2704,plain,(
  ( ! [X158,X161,X159,X160] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X161))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',happ(s(fun('type.num.num','type.ast.lit'),X159),s('type.num.num',X160))))) | s('type.ast.lit',X161) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X158),s(fun('type.num.num','type.ast.lit'),X159),s('type.num.num','const.num.SUC_1'(s('type.num.num',X160)))))) )),
  inference(superposition,[],[f334,f324])).
fof(f2730,plain,(
  ( ! [X134,X138,X136,X135,X137] : (p(s(bool,'|dollar|exists_1'(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137)))))) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137))),s(X134,sK4(X138,'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))),X134)))))) )),
  inference(forward_demodulation,[],[f2698,f324])).
fof(f2698,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,happ(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137))),s(X134,sK4(X138,'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))),X134))))) | p(s(bool,'|dollar|exists_1'(s(fun(X134,bool),'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137))))))))) )),
  inference(superposition,[],[f319,f324])).
fof(f2524,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2499,f262])).
fof(f2499,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2424,f445])).
fof(f2232,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2197,f2070])).
fof(f2197,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(superposition,[],[f2168,f342])).
fof(f2231,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2196,f2070])).
fof(f2196,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(superposition,[],[f2168,f336])).
fof(f2230,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2195,f2070])).
fof(f2195,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(superposition,[],[f2168,f341])).
fof(f2229,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2194,f2070])).
fof(f2194,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(superposition,[],[f2168,f337])).
fof(f2228,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X24) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2193,f2070])).
fof(f2193,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2168,f316])).
fof(f2226,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2191,f2070])).
fof(f2191,plain,(
  ( ! [X21,X19,X20] : (s(bool,X21) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2168,f357])).
fof(f2225,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f2190,f2070])).
fof(f2190,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(superposition,[],[f2168,f309])).
fof(f2168,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2143,f262])).
fof(f2143,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f447,f2070])).
fof(f2051,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f357,f1908,f447])).
fof(f2070,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1908,f272])).
fof(f1392,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1379,f262])).
fof(f1379,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f319,f445])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1367,f309])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f319,f309])).
fof(f1386,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1368,f357])).
fof(f1368,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f319,f357])).
fof(f1533,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1514,f338])).
fof(f1514,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f319,f338])).
fof(f1387,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1369,f316])).
fof(f1369,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f319,f316])).
fof(f1388,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12)))))) )),
  inference(forward_demodulation,[],[f1370,f337])).
fof(f1370,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14))))))))) )),
  inference(superposition,[],[f319,f337])).
fof(f1389,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1371,f341])).
fof(f1371,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f319,f341])).
fof(f1390,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | ~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21)))))) )),
  inference(forward_demodulation,[],[f1372,f336])).
fof(f1372,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23))))))))) )),
  inference(superposition,[],[f319,f336])).
fof(f1391,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X28)))) | ~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25)))))) )),
  inference(forward_demodulation,[],[f1373,f342])).
fof(f1373,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25))))) | p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28))))))))) )),
  inference(superposition,[],[f319,f342])).
fof(f1497,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f319,f338])).
fof(f1519,plain,(
  ( ! [X99,X97,X98] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X99))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X97),X98))))) | s('type.ast.lit',X99) = s('type.ast.lit',happ(s(fun('type.pair.prod'('type.ast.lit',X97),'type.ast.lit'),'const.pair.FST_0'),s('type.pair.prod'('type.ast.lit',X97),X98)))) )),
  inference(superposition,[],[f334,f338])).
fof(f1360,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1117,f262])).
fof(f1117,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f447,f445])).
fof(f1359,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1116,f262])).
fof(f1116,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f418,f445])).
fof(f1358,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1115,f262])).
fof(f1115,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f333,f445])).
fof(f1006,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f445,f342])).
fof(f1005,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f445,f336])).
fof(f1004,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f445,f341])).
fof(f1003,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f445,f337])).
fof(f1002,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f445,f316])).
fof(f1001,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK18(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f445,f357])).
fof(f1000,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f445,f309])).
fof(f990,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f445,f445])).
fof(f973,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) )),
  inference(superposition,[],[f445,f342])).
fof(f972,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X39)))))) )),
  inference(superposition,[],[f445,f336])).
fof(f971,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f445,f341])).
fof(f970,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,X32)))))) )),
  inference(superposition,[],[f445,f337])).
fof(f969,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f445,f316])).
fof(f968,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f445,f357])).
fof(f967,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f445,f309])).
fof(f934,plain,(
  ( ! [X103,X105,X106,X104] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X105))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X106))) | s('type.ast.lit',X106) = s('type.ast.lit',happ(s(fun('type.list.list'(X103),'type.ast.lit'),sK18(X104,X105,X103,'type.ast.lit')),s('type.list.list'(X103),'const.list.NIL_0')))) )),
  inference(superposition,[],[f334,f357])).
fof(f842,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X17))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X18))) | s('type.ast.lit',X18) = s('type.ast.lit','const.pair.SND_1'(s('type.pair.prod'(X15,'type.ast.lit'),'const.pair.,_2'(s(X15,X16),s('type.ast.lit',X17)))))) )),
  inference(superposition,[],[f334,f342])).
fof(f841,plain,(
  ( ! [X14,X12,X13] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X13))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X14))) | s('type.ast.lit',X14) = s('type.ast.lit','const.sum.OUTL_1'(s('type.sum.sum'('type.ast.lit',X12),'const.sum.INL_1'(s('type.ast.lit',X13)))))) )),
  inference(superposition,[],[f334,f336])).
fof(f840,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X9))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X11))) | s('type.ast.lit',X11) = s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X8),'const.pair.,_2'(s('type.ast.lit',X9),s(X8,X10)))))) )),
  inference(superposition,[],[f334,f341])).
fof(f839,plain,(
  ( ! [X6,X7,X5] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X6))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.ast.lit',X7) = s('type.ast.lit','const.sum.OUTR_1'(s('type.sum.sum'(X5,'type.ast.lit'),'const.sum.INR_1'(s('type.ast.lit',X6)))))) )),
  inference(superposition,[],[f334,f337])).
fof(f838,plain,(
  ( ! [X4,X2,X3] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.ast.lit',X4) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X2),s(fun('type.num.num','type.ast.lit'),X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f334,f316])).
fof(f837,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X1) = s('type.ast.lit',happ(s(fun('type.ast.lit','type.ast.lit'),'_i_0_0'),s('type.ast.lit',X0)))) )),
  inference(superposition,[],[f334,f309])).
fof(f836,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f835,f264])).
fof(f835,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f318,f270])).
fof(f803,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f790,f262])).
fof(f790,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f447,f440])).
fof(f494,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f489,f264])).
fof(f489,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f330,f270])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f275,f270])).
fof(f440,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f265,f272])).
fof(f432,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f264,f271])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & ((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5)) | s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16,X17,X22,X18,X19,X23] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X16),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X17),s('type.list.list'('type.string.char'),X22))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X19),s('type.list.list'('type.string.char'),X23))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X16) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X18) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X17) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X19) & s('type.list.list'('type.string.char'),X22) = s('type.list.list'('type.string.char'),X23)))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_|49||49||split|3')).
fof(f404,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4)) )),
  inference(cnf_transformation,[],[f259])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)) )),
  inference(cnf_transformation,[],[f259])).
fof(f406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X4) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3)) )),
  inference(cnf_transformation,[],[f259])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.conLang.exp',X2) = s('type.conLang.exp',X5))) & (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.conLang.exp',X2) != s('type.conLang.exp',X5) | s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5)))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.conLang.exp',X2) = s('type.conLang.exp',X5))) & ((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.conLang.exp',X2) != s('type.conLang.exp',X5)) | s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5)))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.conLang.exp',X2) = s('type.conLang.exp',X5)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X16,X17,X22,X18,X19,X23] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X16),s('type.list.list'('type.string.char'),X17),s('type.conLang.exp',X22))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X18),s('type.list.list'('type.string.char'),X19),s('type.conLang.exp',X23))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X16) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X18) & s('type.list.list'('type.string.char'),X17) = s('type.list.list'('type.string.char'),X19) & s('type.conLang.exp',X22) = s('type.conLang.exp',X23)))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_|49||49||split|2')).
fof(f400,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4)) )),
  inference(cnf_transformation,[],[f257])).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.conLang.exp',X2) = s('type.conLang.exp',X5)) )),
  inference(cnf_transformation,[],[f257])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.conLang.exp',X2) != s('type.conLang.exp',X5) | s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f397,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X25,X26,X52,X27] : ! [X28,X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X52,X25),X28),s(fun(X26,X52),happ(s(fun(X27,fun(X26,X52)),X29),s(X27,X30))))) = s(fun(X26,X25),happ(s(fun(X27,fun(X26,X25)),'_dst_x00x11_2'(s(fun(X52,X25),X28),s(fun(X27,fun(X26,X52)),X29))),s(X27,X30)))),
  file('thm.conSem.v_|49||49||split|4','_dst_x00x11_ax')).
fof(f394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : ((s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),X3))) & (s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),X3) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : ((s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),X3))) & ((s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),X3)) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3)))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) <=> (s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.conSem.v'),X19))) <=> (s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16) = s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X18) & s('type.list.list'('type.conSem.v'),X17) = s('type.list.list'('type.conSem.v'),X19)))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_|49||49||split|1')).
fof(f395,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f396,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),X3) | s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.conSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f418,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X38,X55,X43,X56] : ((s(bool,X38) = s(bool,X55) & (p(s(bool,X55)) => s(bool,X43) = s(bool,X56))) => ((p(s(bool,X38)) => p(s(bool,X43))) <=> (p(s(bool,X55)) => p(s(bool,X56)))))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.IMP_CONG')).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.conSem.v_|49||49||split|4','_dst_x0x1_ax')).
fof(f368,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X25,X26,X27] : ! [X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X30),s(fun(X26,X27),X29))) = s(fun(X26,X25),happ(s(fun(fun(X27,X25),fun(X26,X25)),'_dstsel_0_x01_1'(s(fun(X26,X27),X29))),s(fun(X27,X25),X30)))),
  file('thm.conSem.v_|49||49||split|4','_dstsel_0_x01_ax')).
fof(f367,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ! [X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X45,X9,X31] : ! [X10] : ? [X33] : ! [X38,X43] : s(X45,happ(s(fun(X31,X45),happ(s(fun(X9,fun(X31,X45)),X10),s(X9,X38))),s(X31,X43))) = s(X45,happ(s(fun('type.pair.prod'(X9,X31),X45),X33),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.conSem.v_|49||49||split|4','thm.pair.pair_Axiom')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.conSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v'))),'const.conSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.conSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v'))),'const.conSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X20,X21] : s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X20),s('type.list.list'('type.string.char'),X21))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.conSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),fun('type.list.list'('type.string.char'),'type.conSem.v'))),'const.conSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X20))),s('type.list.list'('type.string.char'),X21)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_3_const.conSem.Recclosure')).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v',happ(s(fun('type.conLang.exp','type.conSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v'))),'const.conSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0))),s('type.list.list'('type.string.char'),X1))),s('type.conLang.exp',X2)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X2))) = s('type.conSem.v',happ(s(fun('type.conLang.exp','type.conSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v'))),'const.conSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X0))),s('type.list.list'('type.string.char'),X1))),s('type.conLang.exp',X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8,X20,X21] : s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.string.char'),X20),s('type.conLang.exp',X21))) = s('type.conSem.v',happ(s(fun('type.conLang.exp','type.conSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.string.char'),fun('type.conLang.exp','type.conSem.v'))),'const.conSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8))),s('type.list.list'('type.string.char'),X20))),s('type.conLang.exp',X21)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_3_const.conSem.Closure')).
fof(f357,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X31,X9] : ! [X32,X11] : ? [X33] : (s(X31,X32) = s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X31,happ(s(fun(X31,X31),happ(s(fun('type.list.list'(X9),fun(X31,X31)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X31,X31))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),X17))))))),
  file('thm.conSem.v_|49||49||split|4','thm.list.list_Axiom')).
fof(f358,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f352,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X31,X9] : ! [X10,X37] : (s(fun(X9,X31),X10) = s(fun(X9,X31),X37) <=> ! [X38] : s(X31,happ(s(fun(X9,X31),X10),s(X9,X38))) = s(X31,happ(s(fun(X9,X31),X37),s(X9,X38))))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.FUN_EQ_THM')).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f351,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.conSem.v_|49||49||split|4',aHL_EXT)).
fof(f342,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X31] : ! [X38,X43] : s(X31,X43) = s(X31,'const.pair.SND_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.conSem.v_|49||49||split|4','thm.pair.SND')).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31,X9] : ! [X38,X43] : s(X9,X38) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.conSem.v_|49||49||split|4','thm.pair.FST')).
fof(f340,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X31,X9] : ! [X48] : ? [X10] : ! [X49,X50,X51] : s(X31,happ(s(fun('type.ind_type.recspace'(X9),X31),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X49),s(X9,X50),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))))) = s(X31,happ(s(fun(fun('type.num.num',X31),X31),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)))),X48),s('type.num.num',X49))),s(X9,X50))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))),s(fun('type.num.num',X31),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X31),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51)))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.CONSTR_REC')).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9,X31] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X31),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X31),X8)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_1_const.pair.FST')).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X31] : ! [X38] : s(X31,X38) = s(X31,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X31),'const.sum.INR_1'(s(X31,X38)))))),
  file('thm.conSem.v_|49||49||split|4','thm.sum.OUTR')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X31,X9] : ! [X38] : s(X9,X38) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X31),'const.sum.INL_1'(s(X9,X38)))))),
  file('thm.conSem.v_|49||49||split|4','thm.sum.OUTL')).
fof(f334,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X6))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) <=> s('type.ast.lit',X6) = s('type.ast.lit',X7))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_|49||49||split|0')).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f233])).
fof(f330,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X40,X41] : ((p(s(bool,X40)) => p(s(bool,X41))) => ((p(s(bool,X41)) => p(s(bool,X40))) => s(bool,X40) = s(bool,X41)))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f329,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.conSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.conSem.v'),X1))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.conSem.v'),X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X20] : s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.conSem.v'),X20))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.conSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X8))),s('type.list.list'('type.conSem.v'),X20)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_2_const.conSem.Conv')).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X0),s('type.list.list'('type.conSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v',X0))),s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X0),s('type.list.list'('type.conSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v',X0))),s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),X1)))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X16,X17] : s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X16),s('type.list.list'('type.conSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v',X16))),s('type.num.num','const.conSem.v3_size_1'(s('type.list.list'('type.conSem.v'),X17)))))))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_size_def|split|10')).
fof(f326,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X8,X20,X21] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X20),s('type.num.num',X21))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X20))),s('type.num.num',X21)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X20,X21] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X20),s(fun('type.num.num','type.ind_type.recspace'(X9)),X21))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X20))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X21)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_1_3_const.ind_type.CONSTR')).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X10,X42] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X42))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.FCONS|split|1')).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.conSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f272,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f269,plain,(
  ( ! [X0] : (s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.conSem.v_size_1'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_size_def|split|4')).
fof(f268,plain,(
  ( ! [X0] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.conSem.Vectorv_0'),s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.conSem.Vectorv_0'),s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X8))) = s('type.conSem.v',happ(s(fun('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.conSem.Vectorv_0'),s('type.list.list'('type.conSem.v'),X8)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_1_const.conSem.Vectorv')).
fof(f267,plain,(
  ( ! [X0] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v',happ(s(fun('type.ast.lit','type.conSem.v'),'const.conSem.Litv_0'),s('type.ast.lit',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v',happ(s(fun('type.ast.lit','type.conSem.v'),'const.conSem.Litv_0'),s('type.ast.lit',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X8))) = s('type.conSem.v',happ(s(fun('type.ast.lit','type.conSem.v'),'const.conSem.Litv_0'),s('type.ast.lit',X8)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_1_const.conSem.Litv')).
fof(f266,plain,(
  ( ! [X0] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v',happ(s(fun('type.num.num','type.conSem.v'),'const.conSem.Loc_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v',happ(s(fun('type.num.num','type.conSem.v'),'const.conSem.Loc_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X8))) = s('type.conSem.v',happ(s(fun('type.num.num','type.conSem.v'),'const.conSem.Loc_0'),s('type.num.num',X8)))),
  file('thm.conSem.v_|49||49||split|4','ARITY_0_1_const.conSem.Loc')).
fof(f265,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.conSem.v_|49||49||split|4','thm.ind_type.ISO_REFL')).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.conSem.v_|49||49||split|4',aHL_TRUTH)).
fof(f19892,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(forward_demodulation,[],[f19795,f336])).
fof(f19795,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(duplicate_literal_removal,[],[f19633])).
fof(f19633,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(superposition,[],[f11955,f14486])).
fof(f14486,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(trivial_inequality_removal,[],[f14475])).
fof(f14475,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(superposition,[],[f13131,f445])).
fof(f19951,plain,(
  ( ! [X74,X72] : (s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(subsumption_resolution,[],[f19950,f264])).
fof(f19950,plain,(
  ( ! [X74,X72] : (p(s(bool,'F_0')) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(forward_demodulation,[],[f19612,f336])).
fof(f19612,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X73),'const.sum.INL_1'(s(bool,'F_0')))))) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(superposition,[],[f2426,f14486])).
fof(f1394,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1395,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2265,plain,(
  ( ! [X23] : (~p(s(bool,sK4(X23,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f20140,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl42),
  inference(cnf_transformation,[],[f20140_D])).
fof(f20140_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f20139,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2382,f262])).
fof(f2382,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f319,f472])).
fof(f6337,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f6337_D])).
fof(f6337_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4761,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4728,f262])).
fof(f4728,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4668,f472])).
fof(f1906,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) | $spl8),
  inference(cnf_transformation,[],[f1906_D])).
fof(f1906_D,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4741,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4668,f309])).
fof(f17447,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4668,f309])).
fof(f18184,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f18080,f262])).
fof(f18080,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f4668,f472])).
fof(f19875,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl40),
  inference(cnf_transformation,[],[f19875_D])).
fof(f19875_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f19802,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(duplicate_literal_removal,[],[f19615])).
fof(f19615,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(superposition,[],[f2594,f14486])).
fof(f2594,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2593,f264])).
fof(f2593,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2560,f1409])).
fof(f1409,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f276])).
fof(f2560,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2426,f1433])).
fof(f1433,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1411,f262])).
fof(f1411,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1395,f445])).
fof(f15618,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(cnf_transformation,[],[f15618_D])).
fof(f15618_D,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f15568,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f15431])).
fof(f15431,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(superposition,[],[f2594,f14235])).
fof(f13953,plain,(
  $false | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f13131,f13809])).
fof(f13809,plain,(
  ( ! [X89] : (s(bool,'T_0') = s(bool,X89)) ) | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4575,f4576,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4742,f4743,f4779,f4789,f4823,f4829,f4843,f4873,f4872,f4866,f4928,f4938,f4972,f5054,f5079,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5388,f5389,f5390,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6612,f6613,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8334,f8335,f8336,f8388,f8365,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f8395,f8411,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f13808])).
fof(f13808,plain,(
  ( ! [X88,X87,X89] : (s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4567,f8455])).
fof(f13804,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4843,f8455])).
fof(f13803,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5368,f8455])).
fof(f13802,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK18(X114,happ(s(fun('type.list.list'(X110),bool),sK18(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5369,f8455])).
fof(f13801,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK18(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5371,f8455])).
fof(f13800,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK18(X131,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X129)))),X130,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5373,f8455])).
fof(f13799,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5374,f8455])).
fof(f13798,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X139),bool),sK18(X140,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X138)))),X139,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5375,f8455])).
fof(f13797,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))),X144,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5376,f8455])).
fof(f13952,plain,(
  $false | ($spl28 | $spl36)),
  inference(subsumption_resolution,[],[f13133,f13809])).
fof(f13951,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f264,f8455])).
fof(f13950,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4575,f4576,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4742,f4743,f4779,f4789,f4823,f4829,f4843,f4873,f4872,f4866,f4928,f4938,f4972,f5054,f5079,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5388,f5389,f5390,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6612,f6613,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8334,f8335,f8336,f8388,f8365,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f8395,f8411,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13946,f13949])).
fof(f13949,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl28),
  inference(subsumption_resolution,[],[f275,f8455])).
fof(f13946,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f13945])).
fof(f13945,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) ) | $spl28),
  inference(subsumption_resolution,[],[f2424,f8455])).
fof(f13824,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f13823,f8455])).
fof(f13823,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5812,f8455])).
fof(f13822,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f13821,f8455])).
fof(f13821,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f5897,f8455])).
fof(f13820,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl28),
  inference(subsumption_resolution,[],[f350,f8455])).
fof(f13819,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f355,f8455])).
fof(f13818,plain,(
  ( ! [X389,X391,X385,X387,X393,X388,X390,X386,X392] : (s(X386,X393) = s(X386,happ(s(fun(X385,X386),happ(s(fun(X387,fun(X385,X386)),X388),s(X387,happ(s(fun(X389,X387),X390),s(X389,X391))))),s(X385,happ(s(fun(X386,X385),X392),s(X386,X393)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5851,f8455])).
fof(f13817,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X411,X404,X406] : (s(X403,X411) = s(X403,happ(s(fun(X404,X403),X410),s(X404,happ(s(fun(X403,X404),happ(s(fun(X405,fun(X403,X404)),X406),s(X405,happ(s(fun(X407,X405),X408),s(X407,X409))))),s(X403,X411)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5853,f8455])).
fof(f13816,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8460,f8455])).
fof(f13815,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8735,f8455])).
fof(f13814,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (s(X376,X383) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),'_dst_x0x1_2'(s(fun(X378,fun(X375,X376)),X379),s(fun(X377,X378),X380))),s(X377,X381))),s(X375,happ(s(fun(X376,X375),X382),s(X376,X383)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5935,f8455])).
fof(f13813,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (s(X393,X401) = s(X393,happ(s(fun(X394,X393),X400),s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X396,fun(X393,X394)),X397),s(fun(X395,X396),X398))),s(X395,X399))),s(X393,X401)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5937,f8455])).
fof(f13812,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8461,f8455])).
fof(f13811,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8736,f8455])).
fof(f13810,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) ) | $spl28),
  inference(subsumption_resolution,[],[f4671,f8455])).
fof(f13947,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4575,f4576,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4742,f4743,f4779,f4789,f4823,f4829,f4843,f4873,f4872,f4866,f4928,f4938,f4972,f5054,f5079,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5388,f5389,f5390,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6612,f6613,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8334,f8335,f8336,f8388,f8365,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f8395,f8411,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13946])).
fof(f13904,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452,f13806,f11291,f13807,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13830,f13835,f13843,f13903])).
fof(f13903,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f3632,f8455])).
fof(f13843,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f13842])).
fof(f13842,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f9470,f13809])).
fof(f13835,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f13834])).
fof(f13834,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f10054,f13809])).
fof(f13830,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f13829])).
fof(f13829,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f12443,f13809])).
fof(f13807,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452,f13806,f11291])).
fof(f13806,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool))))))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11291,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452])).
fof(f13805,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11291,f11387,f11462,f11441,f11442,f11452,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692])).
fof(f13868,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f13867,f8455])).
fof(f13867,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5808,f13822])).
fof(f13866,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f13865,f8455])).
fof(f13865,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5893,f13824])).
fof(f13844,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452,f13806,f11291,f13807,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13830,f13835,f13843])).
fof(f13836,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452,f13806,f11291,f13807,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13830,f13835])).
fof(f13831,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4506,f4507,f4508,f4667,f4514,f4515,f4518,f4519,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4668,f4555,f4556,f4557,f4558,f4559,f4561,f4562,f4567,f4568,f4569,f4570,f4571,f4669,f4575,f4576,f4586,f4589,f4590,f4591,f4592,f4600,f4606,f4671,f4672,f4673,f4684,f4763,f4766,f4742,f4743,f4779,f4789,f4823,f4826,f4829,f4843,f4848,f4873,f4872,f4866,f4918,f4928,f4938,f4972,f4975,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5296,f5297,f5298,f5299,f5300,f5301,f5302,f5303,f5304,f5378,f5381,f5388,f5389,f5390,f5391,f5393,f5394,f5400,f5402,f5368,f5369,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5680,f5679,f5681,f5682,f5683,f5684,f5685,f5686,f5687,f5688,f5689,f5968,f5695,f5696,f5700,f5803,f5804,f5806,f5807,f5808,f5809,f5810,f5812,f5813,f5818,f5819,f5820,f5821,f5822,f5826,f5827,f5837,f5838,f5843,f5844,f5845,f5847,f5848,f5849,f5851,f5852,f5853,f5854,f5860,f5862,f5866,f5872,f5886,f5887,f5888,f5889,f5891,f5892,f5893,f5894,f5895,f5897,f5898,f5903,f5904,f5905,f5906,f5907,f5911,f5912,f5922,f5923,f5927,f5928,f5929,f5931,f5932,f5933,f5935,f5936,f5937,f5938,f5944,f5946,f5950,f5956,f6125,f6126,f6127,f6128,f6129,f6130,f6132,f6134,f6135,f6219,f6229,f6325,f6326,f6236,f6237,f6327,f6328,f6329,f6330,f6331,f6332,f6333,f6334,f6246,f6336,f6339,f6342,f6345,f6313,f6314,f6346,f6316,f6384,f6385,f6386,f6387,f6388,f6389,f6391,f6393,f6394,f6480,f6490,f6597,f6598,f6556,f6557,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6566,f6567,f6568,f6569,f6608,f6575,f6576,f6612,f6613,f6616,f6621,f6595,f6622,f6853,f6854,f6773,f6774,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6864,f6796,f6797,f6801,f6802,f6803,f6804,f6867,f6868,f6807,f6808,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6818,f6819,f6820,f6878,f6826,f6827,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6882,f6843,f6844,f6848,f6849,f6850,f6851,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7594,f8103,f8104,f8366,f8367,f8308,f8309,f8368,f8369,f8370,f8371,f8372,f8373,f8374,f8375,f8318,f8319,f8320,f8321,f8377,f8327,f8328,f8334,f8335,f8336,f8383,f8342,f8343,f8388,f8391,f8365,f8444,f8460,f8461,f8470,f8473,f8474,f8483,f8520,f8521,f8524,f8735,f8736,f8745,f8748,f8749,f8758,f8795,f8798,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9470,f9425,f9426,f9427,f9428,f9429,f9430,f9431,f9432,f9433,f9434,f9435,f9745,f9746,f9747,f9903,f9753,f9754,f9761,f9775,f9776,f9777,f9778,f9793,f9795,f9796,f9800,f9801,f9813,f9814,f9817,f9828,f9937,f9938,f9832,f9940,f9947,f9948,f9953,f9858,f9863,f9872,f9878,f9884,f10054,f10007,f10008,f10009,f10010,f10011,f10012,f10013,f10014,f10015,f10016,f10017,f10018,f10711,f10732,f10659,f10733,f10734,f10735,f10648,f10753,f10775,f10797,f10684,f10685,f10709,f11226,f11301,f11280,f11281,f11387,f11462,f11441,f11442,f11531,f11532,f11533,f11534,f11535,f11536,f11537,f11538,f11539,f11540,f11541,f11542,f11627,f11702,f11681,f11682,f11823,f11824,f11825,f11826,f11829,f11834,f11955,f11965,f11904,f11905,f11972,f11973,f11974,f11975,f11976,f11977,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f12017,f12018,f11988,f11989,f12019,f12020,f12021,f12022,f12023,f12024,f12025,f12026,f12027,f11999,f12029,f12030,f12033,f12145,f12150,f12155,f12273,f12274,f12275,f12276,f12277,f12278,f12279,f12280,f12281,f12282,f12283,f12284,f12285,f12286,f12318,f12319,f12289,f12290,f12320,f12321,f12322,f12323,f12324,f12325,f12326,f12327,f12328,f12300,f12330,f12331,f12334,f12443,f12395,f12396,f12397,f12398,f12401,f12406,f12600,f12601,f12602,f12603,f12604,f12605,f12606,f12607,f12608,f12609,f12610,f12611,f12612,f12613,f12645,f12646,f12616,f12617,f12647,f12648,f12649,f12650,f12651,f12652,f12653,f12654,f12655,f12627,f12657,f12658,f12661,f13130,f13132,f8395,f8411,f8414,f8445,f8450,f13136,f8455,f13797,f13798,f13799,f13800,f13801,f13802,f13803,f13804,f11543,f10695,f5305,f4598,f6817,f6783,f4593,f11692,f13805,f11452,f13806,f11291,f13807,f13808,f13809,f13810,f13811,f13812,f13813,f13814,f13815,f13816,f13817,f13818,f13819,f13820,f13822,f13824,f13830])).
fof(f13828,plain,(
  $false | ($spl6 | $spl10 | $spl28)),
  inference(subsumption_resolution,[],[f13129,f13809])).
fof(f13129,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f12387,f1410])).
fof(f1410,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f271])).
fof(f12387,plain,(
  ( ! [X65] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X65,'_i_0_0',bool))) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f11955,f1409])).
fof(f13827,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f13134,f13809])).
fof(f13134,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9999,f1410])).
fof(f9999,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6480,f1409])).
fof(f13826,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f13135,f13809])).
fof(f13135,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9417,f1410])).
fof(f9417,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6219,f1409])).
fof(f13792,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f13414,f13053])).
fof(f13414,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f13053,f13135])).
fof(f13791,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f13413,f13053])).
fof(f13413,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f13053,f13134])).
fof(f13790,plain,(
  $false | ($spl6 | $spl10 | $spl30)),
  inference(subsumption_resolution,[],[f13412,f13053])).
fof(f13412,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl6 | $spl10 | $spl30)),
  inference(backward_demodulation,[],[f13053,f13129])).
fof(f13782,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f13781,f8455])).
fof(f13781,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13306,f13053])).
fof(f13306,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(backward_demodulation,[],[f13053,f5893])).
fof(f13780,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f13779,f8455])).
fof(f13779,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13296,f13053])).
fof(f13296,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(backward_demodulation,[],[f13053,f5808])).
fof(f13772,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f13771,f8455])).
fof(f13771,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f13138,f13053])).
fof(f13138,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl30),
  inference(backward_demodulation,[],[f13053,f275])).
fof(f13770,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13137,f262])).
fof(f13137,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(backward_demodulation,[],[f13053,f264])).
fof(f13133,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl36),
  inference(cnf_transformation,[],[f13133_D])).
fof(f13133_D,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13132,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(subsumption_resolution,[],[f11895,f11955])).
fof(f11895,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X78))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(superposition,[],[f6219,f2913])).
fof(f13131,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl34),
  inference(cnf_transformation,[],[f13131_D])).
fof(f13131_D,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13130,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f11896,f11955])).
fof(f11896,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f6480,f2913])).
fof(f13054,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl32),
  inference(cnf_transformation,[],[f13054_D])).
fof(f13054_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f12980,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(duplicate_literal_removal,[],[f12790])).
fof(f12790,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(superposition,[],[f2594,f9561])).
fof(f8454,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl26),
  inference(cnf_transformation,[],[f8454_D])).
fof(f8454_D,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8455,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl28),
  inference(cnf_transformation,[],[f8455_D])).
fof(f8455_D,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8444,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl10),
  inference(superposition,[],[f8414,f341])).
fof(f6551,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(cnf_transformation,[],[f6551_D])).
fof(f6551_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6553,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl24),
  inference(cnf_transformation,[],[f6553_D])).
fof(f6553_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6535,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f6338,f309])).
fof(f6338,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl20),
  inference(cnf_transformation,[],[f6338_D])).
fof(f6338_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6345,plain,(
  ( ! [X64,X65] : (p(s(bool,sK5('_i_0_0',X65,bool))) | ~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool)))))) )),
  inference(forward_demodulation,[],[f6312,f309])).
fof(f6312,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X65,bool)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f6336,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f6335,f309])).
fof(f6335,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6298,f262])).
fof(f6298,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f320,f472])).
fof(f5396,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5396_D])).
fof(f5396_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5395,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5361,f275])).
fof(f5361,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4843,f2532])).
fof(f2532,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2510,f262])).
fof(f2510,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f2424,f1610])).
fof(f1610,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1557,f272])).
fof(f1557,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1410,f318])).
fof(f5392,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl14),
  inference(cnf_transformation,[],[f5392_D])).
fof(f5392_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5391,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5358,f275])).
fof(f5358,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4843,f2531])).
fof(f5085,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl12),
  inference(cnf_transformation,[],[f5085_D])).
fof(f5085_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5078,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4743,f341])).
fof(f2047,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1957,f264])).
fof(f1957,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f330])).
fof(f2046,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1959,f1906])).
fof(f1959,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f330])).
fof(f2045,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1967,f264])).
fof(f1967,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f330])).
fof(f2044,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1969,f1906])).
fof(f1969,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f330])).
fof(f1977,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f318])).
fof(f1974,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f440,f1906,f271])).
fof(f1972,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f276])).
fof(f1970,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1968,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f1906,f330])).
fof(f1966,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1964,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1906,f432,f1906,f330])).
fof(f1960,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1958,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f1906,f330])).
fof(f1956,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1954,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1906,f432,f1906,f330])).
fof(f1938,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f342,f1906,f418])).
fof(f1936,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f336,f1906,f418])).
fof(f1934,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f341,f1906,f418])).
fof(f1932,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f337,f1906,f418])).
fof(f1930,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f316,f1906,f418])).
fof(f1928,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f357,f1906,f418])).
fof(f1926,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f309,f1906,f418])).
fof(f1923,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f262,f1610,f1906,f418])).
fof(f1921,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f440,f1906,f418])).
fof(f1918,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1410,f1906,f494])).
fof(f1916,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f432,f1906,f494])).
fof(f1910,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f1906])).
fof(f1908,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f1908_D])).
fof(f1908_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1905,plain,(
  ( ! [X23] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f1888,f309])).
fof(f1888,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f317,f309])).
fof(f1384,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f431,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f430])).
fof(f430,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f427,f428])).
fof(f426,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f428,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f198])).
fof(f198,plain,(
  ? [X0,X1] : ((s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ? [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.conSem.v_|49||49||split|4','thm.conSem.v_|49||49||split|4')).
fof(f427,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f429,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f261,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
% SZS output end Proof for thm.conSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 213
% Passive clauses: 2531
% Generated clauses: 113959
% Final active clauses: 142
% Final passive clauses: 1324
% Input formulas: 69
% Initial clauses: 147
% Discarded non-redundant clauses: 179
% 
% Duplicate literals: 991
% Trivial inequalities: 231
% Fw subsumption resolutions: 1503
% Fw demodulations: 60859
% Bw demodulations: 660
% Condensations: 1340
% Global subsumptions: 14
% 
% Simple tautologies: 353
% Equational tautologies: 195
% Forward subsumptions: 22146
% Fw demodulations to eq. taut.: 23722
% Bw demodulations to eq. taut.: 66
% 
% Unit resulting resolution: 30520
% Factoring: 1
% Forward superposition: 8542
% Backward superposition: 7247
% Self superposition: 166
% Equality factoring: 90
% Equality resolution: 40
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 2644
% SAT solver unit clauses: 231
% SAT solver binary clauses: 885
% 
% Sat splits: 23
% Sat splitting refutations: 51
% 
% Memory used [KB]: 55009
% Time elapsed: 4.427 s
% ------------------------------
% ------------------------------
% Success in time 5.78 s
