Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relation
% SZS output start Proof for thm.clos_relation
fof(f9923,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9886,f75,f110,f108,f9156,f107,f80,f113,f140,f170,f172,f171,f195,f124,f128,f126])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f128,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl13),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f124,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relation.val_rel_closure',aHL_BOOL_CASES)).
fof(f195,plain,(
  $false | ($spl13 | $spl14)),
  inference(trivial_inequality_removal,[],[f192])).
fof(f192,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl13 | $spl14)),
  inference(superposition,[],[f128,f171])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f172,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl7),
  inference(subsumption_resolution,[],[f164,f118])).
fof(f118,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl7),
  inference(superposition,[],[f113,f81])).
fof(f164,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f88,f81])).
fof(f88,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X30,X31] : ((p(s(bool,X30)) => p(s(bool,X31))) => ((p(s(bool,X31)) => p(s(bool,X30))) => s(bool,X30) = s(bool,X31)))),
  file('thm.clos_relation.val_rel_closure','thm.bool.IMP_ANTISYM_AX')).
fof(f140,plain,(
  $false | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f129,f113])).
fof(f129,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f126,f74])).
fof(f74,plain,(
  p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK1),s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))) & (~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1),s('type.closSem.v',X3),s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X3)))) & (~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) | ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X3),s('type.closSem.v',X2))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1),s('type.closSem.v',X3),s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X3))))) & (~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) | ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X3),s('type.closSem.v',X2))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1),s('type.closSem.v',X3),s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X3))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X3),s('type.closSem.v',X2))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1),s('type.closSem.v',X3),s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X3))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X3),s('type.closSem.v',X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X7),s('type.closSem.v',X9),s('type.closSem.v',X8)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X9))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X8)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X9),s('type.closSem.v',X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X7),s('type.closSem.v',X9),s('type.closSem.v',X8)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X9))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X8)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X9),s('type.closSem.v',X8))))))),
  file('thm.clos_relation.val_rel_closure','thm.clos_relation.val_rel_closure')).
fof(f113,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f80,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relation.val_rel_closure',aHL_FALSITY)).
fof(f107,plain,(
  p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | $spl0),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f9156,plain,(
  ( ! [X0] : (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl16)),
  inference(subsumption_resolution,[],[f9155,f76])).
fof(f76,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relation.val_rel_closure',aHL_TRUTH)).
fof(f9155,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl16)),
  inference(forward_demodulation,[],[f9154,f125])).
fof(f125,plain,(
  s(bool,'T_0') = s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3))) | $spl7),
  inference(subsumption_resolution,[],[f120,f113])).
fof(f120,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))),
  inference(superposition,[],[f74,f81])).
fof(f9154,plain,(
  ( ! [X0] : (~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))) | p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(subsumption_resolution,[],[f9127,f172])).
fof(f9127,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))) | p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f86,f367])).
fof(f367,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK1),s('type.closSem.v',sK3),s('type.closSem.v',sK2))) | s(bool,'F_0') = s(bool,X12)) ) | $spl16),
  inference(resolution,[],[f242,f73])).
fof(f73,plain,(
  p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK1),s('type.closSem.v',sK3),s('type.closSem.v',sK2))))),
  inference(cnf_transformation,[],[f62])).
fof(f242,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl16),
  inference(resolution,[],[f172,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f59])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) | p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X1))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) | (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X2),s('type.closSem.v',X1))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2))))) | (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X2),s('type.closSem.v',X1))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X2),s('type.closSem.v',X1))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X1)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X2),s('type.closSem.v',X1))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X10,X11,X12] : ((p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X12),s('type.closSem.v',X11),s('type.closSem.v',X10)))) & p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X11))))) => (p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X10)))) & p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X11),s('type.closSem.v',X10))))))),
  file('thm.clos_relation.val_rel_closure','thm.clos_relation.val_rel_is_closure')).
fof(f108,plain,(
  ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) | $spl1),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f110,plain,(
  ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | $spl3),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f75,plain,(
  ~p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK2))))),
  inference(cnf_transformation,[],[f62])).
fof(f9886,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl3 | $spl7 | $spl16)),
  inference(subsumption_resolution,[],[f9885,f113])).
fof(f9885,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X0)) ) | ($spl3 | $spl7 | $spl16)),
  inference(forward_demodulation,[],[f9884,f245])).
fof(f245,plain,(
  s(bool,'F_0') = s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2))) | ($spl3 | $spl16)),
  inference(resolution,[],[f172,f110])).
fof(f9884,plain,(
  ( ! [X0] : (p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl16)),
  inference(subsumption_resolution,[],[f9883,f76])).
fof(f9883,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl16)),
  inference(forward_demodulation,[],[f9882,f125])).
fof(f9882,plain,(
  ( ! [X0] : (~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))) | p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(subsumption_resolution,[],[f9854,f172])).
fof(f9854,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',sK3)))) | p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',sK3),s('type.closSem.v',sK2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f87,f367])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3),s('type.closSem.v',X2),s('type.closSem.v',X1)))) | ~p(s(bool,'const.closProps.is_closure_1'(s('type.closSem.v',X2)))) | p(s(bool,'const.closProps.check_closures_2'(s('type.closSem.v',X2),s('type.closSem.v',X1))))) )),
  inference(cnf_transformation,[],[f57])).
% SZS output end Proof for thm.clos_relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 91
% Passive clauses: 1162
% Generated clauses: 9647
% Final active clauses: 61
% Final passive clauses: 975
% Input formulas: 21
% Initial clauses: 34
% 
% Duplicate literals: 451
% Trivial inequalities: 95
% Fw subsumption resolutions: 1415
% Fw demodulations: 1550
% Bw demodulations: 26
% 
% Simple tautologies: 162
% Equational tautologies: 56
% Forward subsumptions: 4219
% 
% Binary resolution: 217
% Factoring: 1
% Forward superposition: 3243
% Backward superposition: 2269
% Self superposition: 132
% Equality factoring: 124
% Equality resolution: 11
% 
% Split clauses: 270
% Split components: 540
% SAT solver clauses: 617
% SAT solver unit clauses: 4
% SAT solver binary clauses: 54
% SAT solver learnt clauses: 1
% 
% Sat splits: 270
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 606
% TWLsolver calls for satisfiability: 88
% 
% Memory used [KB]: 4093
% Time elapsed: 0.358 s
% ------------------------------
% ------------------------------
% Success in time 0.37 s
