Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (144538)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 2495
% Generated clauses: 17373
% Final active clauses: 78
% Final passive clauses: 2284
% Input formulas: 21
% Initial clauses: 47
% 
% Duplicate literals: 1099
% Trivial inequalities: 235
% Fw subsumption resolutions: 2409
% Fw demodulations: 1142
% Bw demodulations: 6
% 
% Simple tautologies: 400
% Equational tautologies: 90
% Forward subsumptions: 7096
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 46
% Forward superposition: 7353
% Backward superposition: 4290
% Self superposition: 315
% Equality factoring: 241
% Equality resolution: 14
% 
% Split clauses: 884
% Split components: 1809
% SAT solver clauses: 2180
% SAT solver unit clauses: 72
% SAT solver binary clauses: 477
% 
% Sat splits: 884
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2168
% TWLsolver calls for satisfiability: 111
% 
% Memory used [KB]: 6652
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f2480,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2467,f85,f139,f129,f153,f165,f167,f83,f137,f87,f141,f88,f142,f2288,f2292,f2290,f84,f135,f2262,f2261,f127,f155,f2323,f128,f154])).
fof(f154,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) ) | $spl14),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f128,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | (s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ? [X6,X7] : (s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X7))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) => ? [X6,X7] : (s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X7))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) => ? [X6,X7] : (s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X7))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X8,X9,X10] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9))),s('type.list.list'(X7),X10)))) => ? [X12,X13] : (s('type.list.list'(X7),X10) = s('type.list.list'(X7),'const.list.APPEND_2'(s('type.list.list'(X7),X12),s('type.list.list'(X7),X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X8),s('type.list.list'(X7),X12)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X9),s('type.list.list'(X7),X13))))))),
  file('thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND','thm.bvp_to_wordProps.EVERY2_APPEND_IMP')).
fof(f2323,plain,(
  $false | ($spl4 | $spl14 | $spl22 | $spl133)),
  inference(subsumption_resolution,[],[f2322,f139])).
fof(f2322,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl22 | $spl133)),
  inference(forward_demodulation,[],[f2304,f167])).
fof(f2304,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4)))) | ($spl14 | $spl133)),
  inference(resolution,[],[f2292,f154])).
fof(f155,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0))))) ) | $spl16),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f127,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X1),X5) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),sK7(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f2261,plain,(
  s('type.list.list'(sK1),sK4) = s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0)),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0)))) | $spl124),
  inference(cnf_transformation,[],[f2261_D])).
fof(f2261_D,plain,(
  s('type.list.list'(sK1),sK4) = s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0)),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0)))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2262,plain,(
  s('type.list.list'(sK1),sK4) = s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0)),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0)))) | ($spl4 | $spl16 | $spl22)),
  inference(subsumption_resolution,[],[f2237,f139])).
fof(f2237,plain,(
  ~p(s(bool,'T_0')) | s('type.list.list'(sK1),sK4) = s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0)),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0)))) | ($spl16 | $spl22)),
  inference(superposition,[],[f155,f167])).
fof(f135,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK4) != s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X7))))) ) | $spl0),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK4) != s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X7))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f84,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),X6)))) | s('type.list.list'(sK1),sK4) != s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),X6),s('type.list.list'(sK1),X7)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4)))) & ! [X6,X7] : (s('type.list.list'(sK1),sK4) != s('type.list.list'(sK1),'const.list.APPEND_2'(s('type.list.list'(sK1),X6),s('type.list.list'(sK1),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f46])).
fof(f46,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X4)))) & ! [X6,X7] : (s('type.list.list'(X1),X4) != s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X4)))) => ? [X6,X7] : (s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X1),X4)))) => ? [X6,X7] : (s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.APPEND_2'(s('type.list.list'(X1),X6),s('type.list.list'(X1),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X7))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9))),s('type.list.list'(X7),X10)))) => ? [X12,X13] : (s('type.list.list'(X7),X10) = s('type.list.list'(X7),'const.list.APPEND_2'(s('type.list.list'(X7),X12),s('type.list.list'(X7),X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X8),s('type.list.list'(X7),X12)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X9),s('type.list.list'(X7),X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9))),s('type.list.list'(X7),X10)))) => ? [X12,X13] : (s('type.list.list'(X7),X10) = s('type.list.list'(X7),'const.list.APPEND_2'(s('type.list.list'(X7),X12),s('type.list.list'(X7),X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X8),s('type.list.list'(X7),X12)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X9),s('type.list.list'(X7),X13))))))),
  file('thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND','thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND')).
fof(f2290,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0))))) | $spl131),
  inference(cnf_transformation,[],[f2290_D])).
fof(f2290_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f2292,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0))))) | $spl133),
  inference(cnf_transformation,[],[f2292_D])).
fof(f2292_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0))))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f2288,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0))))) | ($spl0 | $spl124)),
  inference(trivial_inequality_removal,[],[f2279])).
fof(f2279,plain,(
  s('type.list.list'(sK1),sK4) != s('type.list.list'(sK1),sK4) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK2),s('type.list.list'(sK1),sK7(sK4,sK3,sK2,sK5,sK1,sK0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK4,sK3,sK2,sK5,sK1,sK0))))) | ($spl0 | $spl124)),
  inference(superposition,[],[f135,f2261])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f88,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND',aHL_BOOL_CASES)).
fof(f141,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f87,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND',aHL_FALSITY)).
fof(f137,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4)))) | $spl2),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f83,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4))))),
  inference(cnf_transformation,[],[f55])).
fof(f167,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4))) | $spl22),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f165,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4))) | ($spl2 | $spl7 | $spl8)),
  inference(subsumption_resolution,[],[f160,f141])).
fof(f160,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4))) | ($spl2 | $spl8)),
  inference(superposition,[],[f137,f142])).
fof(f153,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) ) | $spl12),
  inference(cnf_transformation,[],[f153_D])).
fof(f153_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f129,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),sK8(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f139,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f85,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.EVERY2_APPEND_IMP_APPEND',aHL_TRUTH)).
fof(f2467,plain,(
  $false | ($spl4 | $spl12 | $spl22 | $spl131)),
  inference(subsumption_resolution,[],[f2466,f139])).
fof(f2466,plain,(
  ~p(s(bool,'T_0')) | ($spl12 | $spl22 | $spl131)),
  inference(forward_demodulation,[],[f2448,f167])).
fof(f2448,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK1),sK4)))) | ($spl12 | $spl131)),
  inference(resolution,[],[f2290,f153])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 21
% Passive clauses: 87
% Generated clauses: 2323
% Final active clauses: 19
% Final passive clauses: 60
% Input formulas: 21
% Initial clauses: 47
% 
% Duplicate literals: 501
% Trivial inequalities: 20
% Fw subsumption resolutions: 60
% Fw demodulations: 31
% Bw demodulations: 6
% Condensations: 724
% Global subsumptions: 9
% 
% Simple tautologies: 67
% Equational tautologies: 39
% Forward subsumptions: 564
% 
% Binary resolution: 12
% Forward superposition: 278
% Backward superposition: 336
% Self superposition: 130
% Equality factoring: 126
% 
% Split clauses: 30
% Split components: 61
% SAT solver clauses: 319
% SAT solver unit clauses: 58
% SAT solver binary clauses: 69
% SAT solver learnt clauses: 19
% 
% Sat splits: 30
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 282
% TWLsolver calls for satisfiability: 116
% 
% Memory used [KB]: 2302
% Time elapsed: 0.246 s
% ------------------------------
% ------------------------------
% Success in time 0.86 s
