Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.conSem
% (75590)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 174
% Passive clauses: 3424
% Generated clauses: 23203
% Final active clauses: 128
% Final passive clauses: 2957
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 923
% Trivial inequalities: 116
% Fw subsumption resolutions: 1723
% Fw demodulations: 3663
% Bw demodulations: 36
% 
% Simple tautologies: 137
% Equational tautologies: 52
% Forward subsumptions: 12648
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 454
% Factoring: 3
% Forward superposition: 8749
% Backward superposition: 7086
% Self superposition: 230
% Equality factoring: 106
% Equality resolution: 9
% 
% Split clauses: 247
% Split components: 494
% SAT solver clauses: 705
% SAT solver unit clauses: 41
% SAT solver binary clauses: 122
% 
% Sat splits: 247
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 705
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.conSem
% (75613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 1117
% Generated clauses: 14279
% Final active clauses: 99
% Final passive clauses: 833
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 392
% Trivial inequalities: 21
% Fw subsumption resolutions: 512
% Fw demodulations: 1491
% Bw demodulations: 179
% Condensations: 852
% 
% Simple tautologies: 47
% Equational tautologies: 35
% Forward subsumptions: 8058
% 
% Binary resolution: 247
% Factoring: 1
% Forward superposition: 5066
% Backward superposition: 4451
% Self superposition: 170
% Equality factoring: 58
% Equality resolution: 1
% 
% Split clauses: 31
% Split components: 62
% SAT solver clauses: 1795
% SAT solver unit clauses: 130
% SAT solver binary clauses: 37
% 
% Sat splits: 31
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1795
% TWLsolver calls for satisfiability: 245
% 
% Memory used [KB]: 23027
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.conSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.conSem
% SZS output start Proof for thm.conSem
fof(f27123,plain,(
  $false),
  inference(subsumption_resolution,[],[f27122,f214])).
fof(f214,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.conSem.state_fupdfupds_comp|split|5',aHL_TRUTH)).
fof(f27122,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f27120,f2837])).
fof(f2837,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2488,f296])).
fof(f296,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f216,f217])).
fof(f217,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.conSem.state_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f216,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.conSem.state_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f2488,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2404,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X14,X6] : ! [X17,X22] : s(X6,X17) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X17),s(X14,X22)))))),
  file('thm.conSem.state_fupdfupds_comp|split|5','thm.pair.FST')).
fof(f2404,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f259,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.conSem.state_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f27120,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),fun(fun(sK0,'type.conSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f27056,f7267])).
fof(f7267,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),fun(fun(sK0,'type.conSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f213,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f213,plain,(
  s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))) != s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))) != s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) != s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) = s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) = s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9))))),s(fun(X6,'type.conSem.state'(X7)),X8))) = s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10))),s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9))),s(fun(X6,'type.conSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9))))),s(fun(X6,'type.conSem.state'(X7)),X8))) = s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10))),s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9))),s(fun(X6,'type.conSem.state'(X7)),X8)))))),
  file('thm.conSem.state_fupdfupds_comp|split|5','thm.conSem.state_fupdfupds_comp|split|5')).
fof(f27056,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2))))),s(fun(X3,'type.conSem.state'(X0)),X4))) = s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1))),s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2))),s(fun(X3,'type.conSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f273,f228])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9))))) = s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X10))),s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.conSem.v')),'type.list.list'('type.option.option'('type.conSem.v'))),X9)))))),
  file('thm.conSem.state_fupdfupds_comp|split|5','thm.conSem.state_fupdfupds_comp|split|4')).
fof(f273,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14,X6,X18,X15] : ! [X10,X9,X8] : s(fun(X18,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X18,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X18,X15),X8))))) = s(fun(X18,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X18,X15),X8)))),
  file('thm.conSem.state_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.conSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 182
% Passive clauses: 2391
% Generated clauses: 27664
% Final active clauses: 110
% Final passive clauses: 1205
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 826
% Trivial inequalities: 52
% Fw subsumption resolutions: 663
% Fw demodulations: 4096
% Bw demodulations: 893
% Condensations: 1091
% Global subsumptions: 2
% 
% Simple tautologies: 67
% Equational tautologies: 129
% Forward subsumptions: 18223
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 45
% 
% Unit resulting resolution: 774
% Factoring: 1
% Forward superposition: 8960
% Backward superposition: 8914
% Self superposition: 286
% Equality factoring: 114
% Equality resolution: 2
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 2483
% SAT solver unit clauses: 534
% SAT solver binary clauses: 1174
% 
% Sat splits: 26
% Sat splitting refutations: 36
% 
% Memory used [KB]: 37483
% Time elapsed: 1.325 s
% ------------------------------
% ------------------------------
% Success in time 2.67 s
