Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.type_eDeterm
% (139460)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1003
% Passive clauses: 6679
% Generated clauses: 19795
% Final active clauses: 860
% Final passive clauses: 5102
% Input formulas: 133
% Initial clauses: 871
% 
% Duplicate literals: 294
% Trivial inequalities: 124
% Fw subsumption resolutions: 1350
% Fw demodulations: 1077
% Bw demodulations: 112
% 
% Simple tautologies: 253
% Equational tautologies: 65
% Forward subsumptions: 9408
% 
% Binary resolution: 8531
% Factoring: 1462
% Forward superposition: 2715
% Backward superposition: 2880
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 30
% 
% Split clauses: 631
% Split components: 1284
% SAT solver clauses: 897
% SAT solver unit clauses: 11
% SAT solver binary clauses: 153
% 
% Sat splits: 631
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 894
% TWLsolver calls for satisfiability: 403
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.type_eDeterm
% (139481)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 694
% Passive clauses: 3005
% Generated clauses: 14870
% Final active clauses: 648
% Final passive clauses: 1715
% Input formulas: 133
% Initial clauses: 871
% 
% Duplicate literals: 145
% Trivial inequalities: 16
% Fw subsumption resolutions: 964
% Fw demodulations: 1140
% Bw demodulations: 626
% Condensations: 143
% Global subsumptions: 87
% 
% Simple tautologies: 221
% Equational tautologies: 13
% Forward subsumptions: 5733
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 4028
% Factoring: 462
% Forward superposition: 1814
% Backward superposition: 1332
% Self superposition: 29
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 237
% Split components: 482
% SAT solver clauses: 4367
% SAT solver unit clauses: 1094
% SAT solver binary clauses: 515
% SAT solver learnt clauses: 183
% SAT solver learnt literals: 171
% 
% Sat splits: 237
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 4002
% TWLsolver calls for satisfiability: 16781
% 
% Memory used [KB]: 18293
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.type_eDeterm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.type_eDeterm
% SZS output start Proof for thm.type_eDeterm
fof(f12591,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f878,f1771,f1769,f877,f1770,f1768,f887,f886,f1775,f885,f1773,f884,f1781,f1779,f883,f882,f893,f1782,f892,f1783,f1776,f891,f890,f1787,f889,f1785,f888,f1777,f1791,f1789,f2375,f3745,f3756,f3754,f3752,f3762,f3760,f8104,f8201,f8199,f8112,f8205,f8203,f9069,f9070,f9079,f9071,f9080,f9084,f9086,f9088,f9092,f9094,f9096,f9098,f9102,f9104,f9106,f9108,f9122,f9124,f9136,f9138,f9142,f9146,f9205,f9350,f9352,f9368,f9374,f9429,f9442,f9477,f9489,f9711,f9713,f9729,f9735,f9790,f9803,f9838,f9850,f9970,f9972,f9974,f9986,f9988,f9990,f10013,f10015,f10026,f10028,f10051,f10053,f10064,f10066,f10089,f10091,f10102,f10104,f10155,f10157,f10159,f10161,f10978,f10979,f10980,f10981,f10982,f10983,f10984,f10234,f10236,f10240,f10243,f10245,f10247,f10249,f10251,f10253,f10255,f10257,f10261,f10264,f10266,f10268,f10270,f10272,f10274,f10298,f10322,f10360,f10373,f10375,f10393,f10397,f10405,f10407,f10409,f10411,f10413,f10415,f10417,f10431,f10441,f10454,f10464,f10466,f10479,f10481,f10488,f10490,f10503,f10505,f10512,f10514,f10529,f10531,f10538,f10540,f10560,f10584,f10626,f10628,f10649,f10651,f10672,f10674,f10695,f10715,f10730,f10732,f10734,f10747,f10749,f10759,f10766,f10773,f10792,f10795,f10797,f10801,f10806,f10808,f10810,f10812,f10840,f10842,f10858,f10860,f10869,f10871,f10878,f10880,f10882,f10884,f10886,f10888,f10892,f10894,f10896,f10898,f10900,f10902,f10906,f10908,f10910,f10912,f10914,f10916,f10920,f10922,f10924,f10926,f10928,f10930,f10935,f10937,f10939,f10944,f10946,f10948,f10985,f10986,f10987,f10988,f10989,f10990,f10991,f10992,f10993,f10994,f10995,f10996,f10997,f10998,f10999,f11000,f11001,f11002,f11003,f11004,f11005,f11006,f11007,f11008,f11009,f11010,f11011,f11012,f11013,f11014,f11015,f11016,f11017,f11018,f11019,f11020,f11021,f11022,f11023,f11024,f11025,f11026,f11027,f11028,f11029,f11030,f11031,f11032,f11033,f11034,f11035,f11036,f11037,f11038,f11039,f11040,f11041,f11042,f11043,f11044,f11045,f11046,f11047,f11048,f11049,f11050,f11051,f11052,f11053,f11054,f11055,f11056,f11057,f11058,f11059,f11060,f11061,f11062,f11063,f11064,f11065,f11066,f11067,f11068,f11069,f11070,f11071,f11072,f11073,f11074,f11075,f11076,f11077,f11078,f11079,f11080,f11081,f11082,f11083,f11084,f11085,f11086,f11087,f11088,f11089,f11090,f11091,f11092,f11093,f11094,f11095,f11096,f11097,f11098,f11099,f11100,f11101,f11102,f11103,f11104,f11105,f11106,f11107,f11108,f11109,f11110,f11111,f11112,f11113,f11114,f11115,f11116,f11117,f11118,f11119,f11120,f11121,f11122,f11123,f11124,f11125,f11126,f11127,f11128,f11129,f11130,f11131,f11132,f11133,f11134,f11135,f11136,f11137,f11138,f11139,f11140,f11141,f11142,f11143,f11144,f11145,f11146,f11147,f11150,f11151,f11161,f11162,f11164,f11166,f11168,f11172,f11173,f11176,f11178,f11181,f11183,f11186,f11196,f11197,f11198,f11199,f11200,f11201,f11226,f11227,f11228,f11229,f11230,f11231,f11232,f11233,f11234,f11235,f11236,f11237,f11238,f11239,f11240,f11241,f11242,f11243,f11244,f11245,f11246,f11247,f11248,f11249,f11250,f11251,f11252,f11253,f11254,f11255,f11256,f11257,f11258,f11259,f11260,f11261,f11262,f11263,f11264,f11265,f11266,f11267,f11268,f11269,f11270,f11271,f11272,f11273,f11274,f11275,f11276,f11277,f11278,f11279,f11280,f11281,f11282,f11283,f11284,f11285,f11286,f11287,f11288,f11289,f11290,f11291,f11292,f11293,f11294,f11295,f11296,f11297,f11298,f11299,f11300,f11301,f11302,f11303,f11304,f11305,f11306,f11307,f11308,f11309,f11310,f11311,f11312,f11313,f11314,f11315,f11316,f11317,f11318,f11319,f11320,f11321,f11322,f11323,f11324,f11325,f11326,f11327,f11328,f11329,f11330,f11331,f11332,f11333,f11334,f11335,f11336,f11337,f11338,f11339,f11340,f11341,f11342,f11343,f11344,f11345,f11346,f11347,f11348,f11349,f11350,f11351,f11352,f11353,f11354,f11355,f11356,f11357,f11358,f11359,f11360,f11361,f11362,f11363,f11364,f11365,f11366,f11367,f11368,f11369,f11370,f11371,f11372,f11373,f11374,f11375,f11376,f11377,f11378,f11379,f11380,f11381,f11382,f11383,f11384,f11385,f11386,f11387,f11388,f11389,f11390,f11391,f11392,f11393,f11394,f11395,f11396,f11397,f11398,f11399,f11400,f11401,f11402,f11403,f11404,f11405,f11406,f11407,f11408,f11409,f11410,f12590])).
fof(f12590,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12589,f1771])).
fof(f12589,plain,(
  s('type.ast.t',sK97) = s('type.ast.t',sK98) | $spl0),
  inference(forward_demodulation,[],[f12588,f12586])).
fof(f12586,plain,(
  s('type.ast.t',sK97) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97)))))),
  inference(subsumption_resolution,[],[f12578,f879])).
fof(f879,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  p(s(bool,'T_0'))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.TRUTH')).
fof(f12578,plain,(
  ~p(s(bool,'T_0')) | s('type.ast.t',sK97) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97)))))),
  inference(superposition,[],[f996,f2434])).
fof(f2434,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),
  inference(unit_resulting_resolution,[],[f875,f916])).
fof(f916,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f875,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),
  inference(cnf_transformation,[],[f491])).
fof(f491,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) | s('type.ast.t',sK97) = s('type.ast.t',sK98)) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) | s('type.ast.t',sK97) != s('type.ast.t',sK98))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK97,sK98])],[f490])).
fof(f490,plain,(
  ? [X0,X1] : (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X0)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) | s('type.ast.t',X0) = s('type.ast.t',X1)) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) | s('type.ast.t',X0) != s('type.ast.t',X1)))),
  inference(flattening,[],[f489])).
fof(f489,plain,(
  ? [X0,X1] : (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X0)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) & ((s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) | s('type.ast.t',X0) = s('type.ast.t',X1)) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) | s('type.ast.t',X0) != s('type.ast.t',X1))))),
  inference(nnf_transformation,[],[f319])).
fof(f319,plain,(
  ? [X0,X1] : (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X0)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) <~> s('type.ast.t',X0) = s('type.ast.t',X1)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ? [X0,X1] : ((p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X0)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1))))) & (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) <~> s('type.ast.t',X0) = s('type.ast.t',X1)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : ((p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X0)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1))))) => (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X0))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))) <=> s('type.ast.t',X0) = s('type.ast.t',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X6)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X7))))) => (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X6))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X7))) <=> s('type.ast.t',X6) = s('type.ast.t',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X6)))) & p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X7))))) => (s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X6))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X7))) <=> s('type.ast.t',X6) = s('type.ast.t',X7)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.type_eDeterm.type_eDeterm_2_prove_ltb')).
fof(f996,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) | s('type.ast.t',X1) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1)))))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1] : (~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) | s('type.ast.t',X1) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))))))),
  inference(ennf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X1)))) => s('type.ast.t',X1) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X1))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12,X8] : (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num',X12),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X8)))) => s('type.ast.t',X8) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',X8))))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.inferProps.check_freevars_empty_convert_unconvert_id')).
fof(f12588,plain,(
  s('type.ast.t',sK98) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))))) | $spl0),
  inference(forward_demodulation,[],[f12587,f1768])).
fof(f12587,plain,(
  s('type.ast.t',sK98) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98)))))),
  inference(subsumption_resolution,[],[f12579,f879])).
fof(f12579,plain,(
  ~p(s(bool,'T_0')) | s('type.ast.t',sK98) = s('type.ast.t','const.inferProps.convert_t_1'(s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98)))))),
  inference(superposition,[],[f996,f2435])).
fof(f2435,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),
  inference(unit_resulting_resolution,[],[f876,f916])).
fof(f876,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),
  inference(cnf_transformation,[],[f491])).
fof(f11410,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9095,f1792])).
fof(f1792,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f895,f879])).
fof(f895,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f500])).
fof(f500,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f61])).
fof(f61,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f9095,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2425,f9071,f1004])).
fof(f1004,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((p(s(bool,X6)) => p(s(bool,X7))) => ((p(s(bool,X7)) => p(s(bool,X6))) => s(bool,X6) = s(bool,X7)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f2425,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1792,f915])).
fof(f915,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f502])).
fof(f11409,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9097,f9071])).
fof(f9097,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2426,f9071,f1004])).
fof(f2426,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f963,f915])).
fof(f963,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X12] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.prim_rec.NOT_LESS_|48|')).
fof(f11408,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9105,f1792])).
fof(f9105,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2425,f9071,f1004])).
fof(f11407,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9107,f9071])).
fof(f9107,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2426,f9071,f1004])).
fof(f11406,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9111,f1792])).
fof(f9111,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2488,f9071,f1032])).
fof(f1032,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0,X1] : ((~sP3(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1)))),
  inference(rectify,[],[f564])).
fof(f564,plain,(
  ! [X1,X0] : ((~sP3(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0)))),
  inference(flattening,[],[f563])).
fof(f563,plain,(
  ! [X1,X0] : ((~sP3(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP3(X1,X0)))),
  inference(nnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X1,X0] : (sP3(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f2488,plain,(
  sP3('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1792,f879,f1034])).
fof(f1034,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP3(X0,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f11405,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9113,f9071])).
fof(f9113,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2527,f9071,f1032])).
fof(f2527,plain,(
  sP3('F_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),
  inference(unit_resulting_resolution,[],[f1792,f2503])).
fof(f2503,plain,(
  ( ! [X0] : (sP3(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2500,f879])).
fof(f2500,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP3(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(superposition,[],[f1037,f2434])).
fof(f1037,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP3(X0,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f11404,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9114,f9071])).
fof(f9114,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2552,f9071,f1032])).
fof(f2552,plain,(
  sP3('F_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),
  inference(unit_resulting_resolution,[],[f1792,f2504])).
fof(f2504,plain,(
  ( ! [X1] : (sP3(X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f2501,f879])).
fof(f2501,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP3(X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(superposition,[],[f1037,f2435])).
fof(f11403,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9115,f9071])).
fof(f9115,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5852,f9071,f1032])).
fof(f5852,plain,(
  ( ! [X0,X1] : (sP3('F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f1792,f3251])).
fof(f3251,plain,(
  ( ! [X30,X28,X29] : (sP3(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0'))) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f3219,f879])).
fof(f3219,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP3(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1037,f966])).
fof(f966,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X28] : ! [X81] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X28,bool),X81),s('type.list.list'(X28),'const.list.NIL_0')))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.list.EVERY_DEF|split|0')).
fof(f11402,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9116,f9071])).
fof(f9116,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f2916,f9071,f1032])).
fof(f2916,plain,(
  sP3('F_0','const.combin.I_1'(s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f879,f1792,f2502])).
fof(f2502,plain,(
  ( ! [X2,X3] : (sP3(X3,'const.combin.I_1'(s(bool,X2))) | p(s(bool,X3)) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1037,f964])).
fof(f964,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X28] : ! [X82] : s(X28,X82) = s(X28,'const.combin.I_1'(s(X28,X82)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.combin.I_THM')).
fof(f11401,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9117,f1792])).
fof(f9117,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2496,f9071,f1032])).
fof(f2496,plain,(
  sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0')),
  inference(unit_resulting_resolution,[],[f1792,f2492])).
fof(f2492,plain,(
  ( ! [X0] : (sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2489,f879])).
fof(f2489,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(superposition,[],[f1034,f2434])).
fof(f11400,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9118,f1792])).
fof(f9118,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2522,f9071,f1032])).
fof(f2522,plain,(
  sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0')),
  inference(unit_resulting_resolution,[],[f1792,f2493])).
fof(f2493,plain,(
  ( ! [X1] : (sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f2490,f879])).
fof(f2490,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP3('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X1)) )),
  inference(superposition,[],[f1034,f2435])).
fof(f11399,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9119,f1792])).
fof(f9119,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f5790,f9071,f1032])).
fof(f5790,plain,(
  ( ! [X0,X1] : (sP3('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1792,f3250])).
fof(f3250,plain,(
  ( ! [X26,X27,X25] : (sP3('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f3218,f879])).
fof(f3218,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP3('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27)) )),
  inference(superposition,[],[f1034,f966])).
fof(f11398,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9120,f1792])).
fof(f9120,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2819,f9071,f1032])).
fof(f2819,plain,(
  sP3('const.combin.I_1'(s(bool,'T_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f2491])).
fof(f2491,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | sP3('const.combin.I_1'(s(bool,X2)),X3)) )),
  inference(superposition,[],[f1034,f964])).
fof(f11397,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9126,f1792])).
fof(f9126,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2499,f9071,f1032])).
fof(f2499,plain,(
  sP3('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1792,f879,f1037])).
fof(f11396,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9127,f1792])).
fof(f9127,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2527,f9071,f1032])).
fof(f11395,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9128,f1792])).
fof(f9128,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2552,f9071,f1032])).
fof(f11394,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9129,f1792])).
fof(f9129,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f5852,f9071,f1032])).
fof(f11393,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9130,f1792])).
fof(f9130,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2916,f9071,f1032])).
fof(f11392,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9131,f9071])).
fof(f9131,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2496,f9071,f1032])).
fof(f11391,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9132,f9071])).
fof(f9132,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2522,f9071,f1032])).
fof(f11390,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9133,f9071])).
fof(f9133,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5790,f9071,f1032])).
fof(f11389,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9134,f9071])).
fof(f9134,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f2819,f9071,f1032])).
fof(f11388,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9204,f9706])).
fof(f9706,plain,(
  ( ! [X0] : (sP7('T_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2382,f2387,f9071,f1086])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (~sP5(X1,X2,X0) | p(s(bool,X0)) | ~sP6(X0,X2,X1) | sP7(X0,X1,X2) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f579])).
fof(f579,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f578])).
fof(f578,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f2387,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1092])).
fof(f1092,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f586])).
fof(f586,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f585])).
fof(f585,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f584])).
fof(f584,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f2382,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1088])).
fof(f1088,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f582])).
fof(f582,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f581])).
fof(f581,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f2392,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1096])).
fof(f1096,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f588])).
fof(f588,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f587])).
fof(f587,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f9204,plain,(
  ( ! [X0] : (~sP7('T_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f879,f9071,f1083])).
fof(f1083,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f580])).
fof(f11387,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9206,f9989])).
fof(f9989,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f1094])).
fof(f1094,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f586])).
fof(f9206,plain,(
  ~sP5('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f9071,f1085])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (~sP5(X1,X2,X0) | p(s(bool,X1)) | ~sP6(X0,X2,X1) | sP7(X0,X1,X2) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f580])).
fof(f2987,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f879,f2425,f1101])).
fof(f1101,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f263,f391,f390,f389,f388])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X121,X122,X123] : ((p(s(bool,X123)) <=> s(bool,X121) = s(bool,X122)) <=> ((p(s(bool,X123)) | p(s(bool,X122)) | p(s(bool,X121))) & (p(s(bool,X123)) | ~p(s(bool,X121)) | ~p(s(bool,X122))) & (p(s(bool,X122)) | ~p(s(bool,X121)) | ~p(s(bool,X123))) & (p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X123)))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f2395,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f879,f1097])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f588])).
fof(f11386,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9207,f9989])).
fof(f9207,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f9071,f1085])).
fof(f2988,plain,(
  ( ! [X0] : (~sP7('T_0','T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f2426,f1101])).
fof(f11385,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9208,f9989])).
fof(f9208,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2382,f9071,f1085])).
fof(f11384,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9222,f9989])).
fof(f9222,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3004,f2383,f9071,f1085])).
fof(f2383,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f875,f1088])).
fof(f3004,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1792,f964,f1102])).
fof(f1102,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f589])).
fof(f2396,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1)) )),
  inference(unit_resulting_resolution,[],[f875,f1097])).
fof(f11383,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9223,f9989])).
fof(f9223,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3005,f2383,f9071,f1085])).
fof(f3005,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1792,f965,f1102])).
fof(f965,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X22] : ! [X25] : s(X22,X25) = s(X22,happ(s(fun(X22,X22),'_i_0_0'),s(X22,X25)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','_i0_ax')).
fof(f11382,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9224,f9989])).
fof(f9224,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3090,f2383,f9071,f1085])).
fof(f3090,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1792,f1794,f1102])).
fof(f1794,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f967,f964])).
fof(f967,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X28] : ! [X18] : s(X28,'const.combin.I_1'(s(X28,X18))) = s(X28,happ(s(fun(X28,X28),'const.combin.I_0'),s(X28,X18)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','ARITY_0_1_const.combin.I')).
fof(f11381,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9225,f9989])).
fof(f9225,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f4234,f2383,f9071,f1085])).
fof(f4234,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1792,f1060,f1102])).
fof(f1060,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X22,X21] : ! [X24,X26] : s(X21,X24) = s(X21,happ(s(fun(X22,X21),'_dst_x0_1'(s(X21,X24))),s(X22,X26)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','_dst_x0_ax')).
fof(f11380,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9231,f9989])).
fof(f9231,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3004,f2384,f9071,f1085])).
fof(f2384,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f876,f1088])).
fof(f2397,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X1)) )),
  inference(unit_resulting_resolution,[],[f876,f1097])).
fof(f11379,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9232,f9989])).
fof(f9232,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3005,f2384,f9071,f1085])).
fof(f11378,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9233,f9989])).
fof(f9233,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3090,f2384,f9071,f1085])).
fof(f11377,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9234,f9989])).
fof(f9234,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f4234,f2384,f9071,f1085])).
fof(f11376,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9235,f9989])).
fof(f9235,plain,(
  ~sP5('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f9071,f1085])).
fof(f2385,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1089])).
fof(f1089,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f583])).
fof(f11375,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9236,f9989])).
fof(f9236,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f9071,f1085])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f963,f1089])).
fof(f11374,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9237,f9989])).
fof(f9237,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2413,f9071,f1085])).
fof(f11373,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9239,f9989])).
fof(f9239,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3004,f2386,f9071,f1085])).
fof(f2386,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1090])).
fof(f1090,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f11372,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9240,f9989])).
fof(f9240,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3004,f2386,f9071,f1085])).
fof(f11371,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9242,f9989])).
fof(f9242,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3005,f2386,f9071,f1085])).
fof(f11370,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9243,f9989])).
fof(f9243,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3005,f2386,f9071,f1085])).
fof(f11369,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9245,f9989])).
fof(f9245,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3090,f2386,f9071,f1085])).
fof(f11368,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9246,f9989])).
fof(f9246,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3090,f2386,f9071,f1085])).
fof(f11367,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9248,f9989])).
fof(f9248,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f4234,f2386,f9071,f1085])).
fof(f11366,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9249,f9989])).
fof(f9249,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f4234,f2386,f9071,f1085])).
fof(f11365,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9266,f9989])).
fof(f9266,plain,(
  ~sP5(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f3000,f2386,f9071,f1085])).
fof(f3000,plain,(
  ( ! [X0] : (~sP7('F_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1792,f965,f1102])).
fof(f2399,plain,(
  ( ! [X0,X1] : (sP4('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f875,f1098])).
fof(f1098,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f588])).
fof(f11364,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9267,f9989])).
fof(f9267,plain,(
  ~sP5(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f3000,f2386,f9071,f1085])).
fof(f2400,plain,(
  ( ! [X0,X1] : (sP4('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f876,f1098])).
fof(f11363,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9271,f9989])).
fof(f9271,plain,(
  ~sP5(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f3087,f2386,f9071,f1085])).
fof(f3087,plain,(
  ( ! [X0] : (~sP7('F_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1792,f1794,f1102])).
fof(f11362,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9272,f9989])).
fof(f9272,plain,(
  ~sP5(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f3087,f2386,f9071,f1085])).
fof(f11361,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9276,f9989])).
fof(f9276,plain,(
  ( ! [X0,X1] : (~sP5(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f4230,f2386,f9071,f1085])).
fof(f4230,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1060,f1102])).
fof(f11360,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9277,f9989])).
fof(f9277,plain,(
  ( ! [X0,X1] : (~sP5(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f4230,f2386,f9071,f1085])).
fof(f11359,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9281,f9989])).
fof(f9281,plain,(
  ~sP5('const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f2998,f2386,f9071,f1085])).
fof(f2998,plain,(
  ( ! [X0] : (~sP7('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1792,f964,f1102])).
fof(f11358,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9282,f9989])).
fof(f9282,plain,(
  ~sP5('const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f2998,f2386,f9071,f1085])).
fof(f11357,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9286,f9973])).
fof(f9973,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f1090])).
fof(f9286,plain,(
  ~sP6('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2387,f9071,f1085])).
fof(f11356,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9287,f9973])).
fof(f9287,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2387,f9071,f1085])).
fof(f11355,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9288,f9973])).
fof(f9288,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2387,f9071,f1085])).
fof(f11354,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9421,f9973])).
fof(f9421,plain,(
  ~sP6('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2390,f9071,f1085])).
fof(f2390,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1093])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f586])).
fof(f11353,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9468,f9973])).
fof(f9468,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2415,f9071,f1085])).
fof(f2415,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f963,f1093])).
fof(f11352,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9469,f9973])).
fof(f9469,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2415,f9071,f1085])).
fof(f11351,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9567,f9989])).
fof(f9567,plain,(
  ~sP5('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f9071,f1086])).
fof(f11350,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9568,f9989])).
fof(f9568,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f9071,f1086])).
fof(f11349,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9569,f9989])).
fof(f9569,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2382,f9071,f1086])).
fof(f11348,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9583,f9989])).
fof(f9583,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3004,f2383,f9071,f1086])).
fof(f11347,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9584,f9989])).
fof(f9584,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3005,f2383,f9071,f1086])).
fof(f11346,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9585,f9989])).
fof(f9585,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3090,f2383,f9071,f1086])).
fof(f11345,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9586,f9989])).
fof(f9586,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f4234,f2383,f9071,f1086])).
fof(f11344,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9592,f9989])).
fof(f9592,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3004,f2384,f9071,f1086])).
fof(f11343,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9593,f9989])).
fof(f9593,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3005,f2384,f9071,f1086])).
fof(f11342,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9594,f9989])).
fof(f9594,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3090,f2384,f9071,f1086])).
fof(f11341,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9595,f9989])).
fof(f9595,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f4234,f2384,f9071,f1086])).
fof(f11340,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9596,f9989])).
fof(f9596,plain,(
  ~sP5('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f9071,f1086])).
fof(f11339,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9597,f9989])).
fof(f9597,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f9071,f1086])).
fof(f11338,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9598,f9989])).
fof(f9598,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2413,f9071,f1086])).
fof(f11337,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9600,f9989])).
fof(f9600,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3004,f2386,f9071,f1086])).
fof(f11336,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9601,f9989])).
fof(f9601,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3004,f2386,f9071,f1086])).
fof(f11335,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9603,f9989])).
fof(f9603,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3005,f2386,f9071,f1086])).
fof(f11334,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9604,f9989])).
fof(f9604,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3005,f2386,f9071,f1086])).
fof(f11333,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9606,f9989])).
fof(f9606,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f3090,f2386,f9071,f1086])).
fof(f11332,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9607,f9989])).
fof(f9607,plain,(
  ~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f3090,f2386,f9071,f1086])).
fof(f11331,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9609,f9989])).
fof(f9609,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f4234,f2386,f9071,f1086])).
fof(f11330,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9610,f9989])).
fof(f9610,plain,(
  ( ! [X0,X1] : (~sP5('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f4234,f2386,f9071,f1086])).
fof(f11329,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9627,f9989])).
fof(f9627,plain,(
  ~sP5(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f3000,f2386,f9071,f1086])).
fof(f11328,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9628,f9989])).
fof(f9628,plain,(
  ~sP5(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f3000,f2386,f9071,f1086])).
fof(f11327,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9632,f9989])).
fof(f9632,plain,(
  ~sP5(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f3087,f2386,f9071,f1086])).
fof(f11326,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9633,f9989])).
fof(f9633,plain,(
  ~sP5(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f3087,f2386,f9071,f1086])).
fof(f11325,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9637,f9989])).
fof(f9637,plain,(
  ( ! [X0,X1] : (~sP5(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))))),s(X0,X1)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f4230,f2386,f9071,f1086])).
fof(f11324,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9638,f9989])).
fof(f9638,plain,(
  ( ! [X0,X1] : (~sP5(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))))),s(X0,X1)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f4230,f2386,f9071,f1086])).
fof(f11323,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9642,f9989])).
fof(f9642,plain,(
  ~sP5('const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f2998,f2386,f9071,f1086])).
fof(f11322,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9643,f9989])).
fof(f9643,plain,(
  ~sP5('const.combin.I_1'(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f2998,f2386,f9071,f1086])).
fof(f11321,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9647,f9973])).
fof(f9647,plain,(
  ~sP6('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2387,f9071,f1086])).
fof(f11320,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9648,f9973])).
fof(f9648,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2387,f9071,f1086])).
fof(f11319,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9649,f9973])).
fof(f9649,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2387,f9071,f1086])).
fof(f11318,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9739,f10254])).
fof(f10254,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2434,f9071,f1102])).
fof(f9739,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f2382,f2388,f9071,f1086])).
fof(f2388,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f875,f1092])).
fof(f11317,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9740,f10256])).
fof(f10256,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2435,f9071,f1102])).
fof(f9740,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f2382,f2388,f9071,f1086])).
fof(f11316,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9762,f10254])).
fof(f9762,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f2382,f2389,f9071,f1086])).
fof(f2389,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f876,f1092])).
fof(f11315,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9763,f10256])).
fof(f9763,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f2382,f2389,f9071,f1086])).
fof(f11314,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9782,f9973])).
fof(f9782,plain,(
  ~sP6('T_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2390,f9071,f1086])).
fof(f11313,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9829,f9973])).
fof(f9829,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2415,f9071,f1086])).
fof(f11312,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9830,f9973])).
fof(f9830,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f2988,f2415,f9071,f1086])).
fof(f11311,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9905,f10244])).
fof(f10244,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2473,f9071,f1102])).
fof(f2473,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f963,f920])).
fof(f920,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f9905,plain,(
  ( ! [X0] : (sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f2385,f2416,f9071,f1086])).
fof(f2416,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f963,f1094])).
fof(f2393,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f875,f1096])).
fof(f11310,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9906,f10244])).
fof(f9906,plain,(
  ( ! [X0] : (sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f2385,f2416,f9071,f1086])).
fof(f2394,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f876,f1096])).
fof(f11309,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9969,f9973])).
fof(f9969,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f879,f9071,f1087])).
fof(f1087,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f11308,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9985,f9989])).
fof(f9985,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f879,f9071,f1091])).
fof(f1091,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f586])).
fof(f11307,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9992,f9071])).
fof(f9992,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f1792,f9071,f1095])).
fof(f1095,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f588])).
fof(f11306,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9993,f9071])).
fof(f9993,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f1792,f9071,f1095])).
fof(f11305,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9994,f9071])).
fof(f9994,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f1792,f9071,f1095])).
fof(f2398,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f879,f1098])).
fof(f11304,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9995,f9071])).
fof(f9995,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f1792,f9071,f1095])).
fof(f11303,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9996,f9071])).
fof(f9996,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f1792,f9071,f1095])).
fof(f11302,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9997,f9071])).
fof(f9997,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f1095])).
fof(f11301,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9998,f9071])).
fof(f9998,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11300,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f9999,f9071])).
fof(f9999,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11299,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10001,f9071])).
fof(f10001,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11298,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10002,f9071])).
fof(f10002,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11297,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10003,f9071])).
fof(f10003,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f9071,f9071,f1095])).
fof(f11296,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10004,f9071])).
fof(f10004,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11295,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10005,f9071])).
fof(f10005,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11294,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10007,f9071])).
fof(f10007,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f1792,f9071,f1095])).
fof(f11293,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10008,f9071])).
fof(f10008,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f1792,f9071,f1095])).
fof(f11292,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10009,f9071])).
fof(f10009,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f1792,f9071,f1095])).
fof(f11291,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10010,f9071])).
fof(f10010,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f1792,f9071,f1095])).
fof(f11290,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10011,f9071])).
fof(f10011,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f1792,f9071,f1095])).
fof(f11289,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10017,f9071])).
fof(f10017,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11288,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10018,f9071])).
fof(f10018,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11287,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10019,f9071])).
fof(f10019,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f9071,f9071,f1095])).
fof(f11286,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10020,f9071])).
fof(f10020,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11285,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10021,f9071])).
fof(f10021,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11284,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10022,f9071])).
fof(f10022,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f9071,f9071,f1095])).
fof(f11283,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10023,f9071])).
fof(f10023,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11282,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10024,f9071])).
fof(f10024,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11281,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10030,f9071])).
fof(f10030,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f1792,f9071,f1095])).
fof(f11280,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10031,f9071])).
fof(f10031,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f1792,f9071,f1095])).
fof(f11279,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10032,f9071])).
fof(f10032,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f1792,f9071,f1095])).
fof(f11278,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10033,f9071])).
fof(f10033,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f1792,f9071,f1095])).
fof(f11277,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10034,f9071])).
fof(f10034,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f1792,f9071,f1095])).
fof(f11276,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10036,f9071])).
fof(f10036,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11275,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10037,f9071])).
fof(f10037,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11274,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10038,f9071])).
fof(f10038,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f9071,f9071,f1095])).
fof(f11273,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10039,f9071])).
fof(f10039,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11272,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10040,f9071])).
fof(f10040,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11271,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10041,f9071])).
fof(f10041,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f9071,f9071,f1095])).
fof(f11270,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10042,f9071])).
fof(f10042,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11269,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10043,f9071])).
fof(f10043,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11268,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10044,f9071])).
fof(f10044,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f1792,f9071,f1095])).
fof(f11267,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10045,f9071])).
fof(f10045,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f1792,f9071,f1095])).
fof(f11266,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10046,f9071])).
fof(f10046,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f1792,f9071,f1095])).
fof(f11265,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10048,f9071])).
fof(f10048,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f1792,f9071,f1095])).
fof(f11264,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10049,f9071])).
fof(f10049,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f1792,f9071,f1095])).
fof(f11263,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10054,f9071])).
fof(f10054,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f1095])).
fof(f11262,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10055,f9071])).
fof(f10055,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11261,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10056,f9071])).
fof(f10056,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11260,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10057,f9071])).
fof(f10057,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f9071,f9071,f1095])).
fof(f11259,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10058,f9071])).
fof(f10058,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11258,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10059,f9071])).
fof(f10059,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11257,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10061,f9071])).
fof(f10061,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11256,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10062,f9071])).
fof(f10062,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11255,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10067,f9071])).
fof(f10067,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f1792,f9071,f1095])).
fof(f11254,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10068,f9071])).
fof(f10068,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f1792,f9071,f1095])).
fof(f11253,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10069,f9071])).
fof(f10069,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f1792,f9071,f1095])).
fof(f11252,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10071,f9071])).
fof(f10071,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f1792,f9071,f1095])).
fof(f11251,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10072,f9071])).
fof(f10072,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f1792,f9071,f1095])).
fof(f11250,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10073,f9071])).
fof(f10073,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f1095])).
fof(f11249,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10074,f9071])).
fof(f10074,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11248,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10075,f9071])).
fof(f10075,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11247,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10077,f9071])).
fof(f10077,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11246,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10078,f9071])).
fof(f10078,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11245,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10079,f9071])).
fof(f10079,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f9071,f9071,f1095])).
fof(f11244,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10080,f9071])).
fof(f10080,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11243,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10081,f9071])).
fof(f10081,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11242,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10082,f9071])).
fof(f10082,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f1792,f9071,f1095])).
fof(f11241,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10083,f9071])).
fof(f10083,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f1792,f9071,f1095])).
fof(f11240,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10084,f9071])).
fof(f10084,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f1792,f9071,f1095])).
fof(f11239,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10086,f9071])).
fof(f10086,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f1792,f9071,f1095])).
fof(f11238,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10087,f9071])).
fof(f10087,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f1792,f9071,f1095])).
fof(f11237,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10092,f9071])).
fof(f10092,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f1095])).
fof(f11236,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10093,f9071])).
fof(f10093,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2393,f9071,f9071,f1095])).
fof(f11235,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10094,f9071])).
fof(f10094,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2394,f9071,f9071,f1095])).
fof(f11234,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10095,f9071])).
fof(f10095,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f9071,f9071,f1095])).
fof(f11233,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10096,f9071])).
fof(f10096,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2396,f9071,f9071,f1095])).
fof(f11232,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10097,f9071])).
fof(f10097,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2397,f9071,f9071,f1095])).
fof(f11231,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10099,f9071])).
fof(f10099,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2399,f9071,f9071,f1095])).
fof(f11230,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10100,f9071])).
fof(f10100,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2400,f9071,f9071,f1095])).
fof(f11229,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10105,f2425])).
fof(f10105,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2984,f9071,f1099])).
fof(f1099,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f589])).
fof(f2984,plain,(
  ~sP7('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f2425,f1101])).
fof(f11228,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10106,f2425])).
fof(f10106,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2987,f9071,f1099])).
fof(f11227,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10107,f2426])).
fof(f10107,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2988,f9071,f1099])).
fof(f11226,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10108,f2426])).
fof(f10108,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2985,f9071,f1099])).
fof(f2985,plain,(
  ( ! [X0] : (~sP7('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f2426,f1101])).
fof(f11201,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10162,f964])).
fof(f10162,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2967,f9071,f1102])).
fof(f2967,plain,(
  ( ! [X0] : (sP7('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f879,f964,f1100])).
fof(f1100,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f589])).
fof(f11200,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10163,f965])).
fof(f10163,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2968,f9071,f1102])).
fof(f2968,plain,(
  ( ! [X0] : (sP7('T_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f879,f965,f1100])).
fof(f11199,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10164,f1794])).
fof(f10164,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3091,f9071,f1102])).
fof(f3091,plain,(
  ( ! [X0] : (sP7('T_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f879,f1794,f1100])).
fof(f11198,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10165,f1060])).
fof(f10165,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4235,f9071,f1102])).
fof(f4235,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f879,f1060,f1100])).
fof(f11197,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10170,f2473])).
fof(f10170,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2966,f9071,f1102])).
fof(f2966,plain,(
  ( ! [X0] : (sP7('T_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f2473,f1100])).
fof(f11196,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10172,f968])).
fof(f968,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X28] : ! [X82] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X28,X82),s(fun(X28,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X28),'const.list.NIL_0')))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.list.MEM|split|0')).
fof(f10172,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5254,f9071,f1102])).
fof(f5254,plain,(
  ( ! [X0,X1] : (sP7('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(unit_resulting_resolution,[],[f879,f968,f1100])).
fof(f11186,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10185,f965])).
fof(f10185,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2963,f9071,f1102])).
fof(f2963,plain,(
  ( ! [X0] : (sP7('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f965,f1100])).
fof(f11183,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11182,f964])).
fof(f11182,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(forward_demodulation,[],[f10188,f965])).
fof(f10188,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3017,f9071,f1102])).
fof(f3017,plain,(
  ( ! [X4] : (sP7('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4)),'const.combin.I_1'(s(bool,X4)))) )),
  inference(superposition,[],[f2967,f965])).
fof(f11181,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10189,f1794])).
fof(f10189,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3088,f9071,f1102])).
fof(f3088,plain,(
  ( ! [X0] : (sP7('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f1794,f1100])).
fof(f11178,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11177,f964])).
fof(f11177,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(forward_demodulation,[],[f10192,f1794])).
fof(f10192,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3131,f9071,f1102])).
fof(f3131,plain,(
  ( ! [X52] : (sP7('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X52)),'const.combin.I_1'(s(bool,X52)))) )),
  inference(superposition,[],[f2967,f1794])).
fof(f11176,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10193,f1060])).
fof(f10193,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4231,f9071,f1102])).
fof(f4231,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f879,f1060,f1100])).
fof(f11173,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10196,f968])).
fof(f10196,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5251,f9071,f1102])).
fof(f5251,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f879,f968,f1100])).
fof(f11172,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10197,f964])).
fof(f10197,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2962,f9071,f1102])).
fof(f2962,plain,(
  ( ! [X0] : (sP7('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f964,f1100])).
fof(f11168,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11167,f964])).
fof(f11167,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(forward_demodulation,[],[f10201,f965])).
fof(f10201,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2981,f9071,f1102])).
fof(f2981,plain,(
  ( ! [X4] : (sP7('T_0','const.combin.I_1'(s(bool,X4)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4)))) )),
  inference(superposition,[],[f2962,f965])).
fof(f11166,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11165,f964])).
fof(f11165,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(forward_demodulation,[],[f10202,f1794])).
fof(f10202,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3130,f9071,f1102])).
fof(f3130,plain,(
  ( ! [X51] : (sP7('T_0','const.combin.I_1'(s(bool,X51)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X51)))) )),
  inference(superposition,[],[f2962,f1794])).
fof(f11164,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11163,f964])).
fof(f11163,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl44),
  inference(forward_demodulation,[],[f10204,f2473])).
fof(f10204,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2979,f9071,f1102])).
fof(f2979,plain,(
  ( ! [X2] : (sP7('T_0','const.combin.I_1'(s(bool,'F_0')),'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2962,f2473])).
fof(f11162,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10205,f2473])).
fof(f10205,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2961,f9071,f1102])).
fof(f2961,plain,(
  ( ! [X0] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f879,f2473,f1100])).
fof(f11161,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f11160,f964])).
fof(f11160,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl44),
  inference(forward_demodulation,[],[f10206,f2473])).
fof(f10206,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3015,f9071,f1102])).
fof(f3015,plain,(
  ( ! [X2] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f2967,f2473])).
fof(f11151,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10223,f2473])).
fof(f10223,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f6170,f9071,f1102])).
fof(f6170,plain,(
  ( ! [X0] : (sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f2393,f2413,f2415,f1792,f1085])).
fof(f11150,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10224,f2473])).
fof(f10224,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f6171,f9071,f1102])).
fof(f6171,plain,(
  ( ! [X0] : (sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f2394,f2413,f2415,f1792,f1085])).
fof(f11147,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10297,f10539])).
fof(f10539,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f9071,f9071,f1114])).
fof(f1114,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f596])).
fof(f596,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f265,f394,f393])).
fof(f393,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f394,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X121,X122,X123] : ((p(s(bool,X123)) <=> (p(s(bool,X122)) | p(s(bool,X121)))) <=> ((p(s(bool,X123)) | ~p(s(bool,X122))) & (p(s(bool,X123)) | ~p(s(bool,X121))) & (p(s(bool,X122)) | p(s(bool,X121)) | ~p(s(bool,X123)))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f10297,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f9071,f1103])).
fof(f1103,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f592])).
fof(f592,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f591])).
fof(f591,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f394])).
fof(f11146,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10321,f10539])).
fof(f10321,plain,(
  ( ! [X0,X1] : (~sP9(X0,'T_0',X1)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f9071,f1104])).
fof(f1104,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f592])).
fof(f11145,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10349,f9071])).
fof(f10349,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2407,f2526,f9071,f1109])).
fof(f1109,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f592])).
fof(f2526,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f1792,f1104])).
fof(f2407,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1113])).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f594])).
fof(f594,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f593])).
fof(f593,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f393])).
fof(f11144,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10350,f9071])).
fof(f10350,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2404,f2526,f9071,f1109])).
fof(f2404,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f879,f1112])).
fof(f1112,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f595])).
fof(f11143,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10352,f9071])).
fof(f10352,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2402,f2526,f9071,f1109])).
fof(f2402,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f875,f1111])).
fof(f1111,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f595])).
fof(f11142,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10353,f9071])).
fof(f10353,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2403,f2526,f9071,f1109])).
fof(f2403,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f876,f1111])).
fof(f11141,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10359,f10465])).
fof(f10465,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f1113])).
fof(f10359,plain,(
  ~sP8('F_0','T_0','F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f1792,f9071,f1109])).
fof(f11140,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10372,f10465])).
fof(f10372,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1109])).
fof(f2523,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1792,f1103])).
fof(f11139,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10374,f10465])).
fof(f10374,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1109])).
fof(f11138,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10379,f9071])).
fof(f10379,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2401,f9071,f1109])).
fof(f2401,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1111])).
fof(f11137,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10384,f9071])).
fof(f10384,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2405,f9071,f1109])).
fof(f2405,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1)) )),
  inference(unit_resulting_resolution,[],[f875,f1112])).
fof(f11136,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10385,f9071])).
fof(f10385,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2406,f9071,f1109])).
fof(f2406,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X1)) )),
  inference(unit_resulting_resolution,[],[f876,f1112])).
fof(f11135,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10386,f9071])).
fof(f10386,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2407,f9071,f1109])).
fof(f11134,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10388,f10465])).
fof(f10388,plain,(
  ~sP8('F_0','F_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f9071,f1109])).
fof(f11133,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10399,f10465])).
fof(f10399,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1109])).
fof(f11132,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10400,f10465])).
fof(f10400,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1109])).
fof(f11131,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10420,f9071])).
fof(f10420,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2402,f879,f9071,f1110])).
fof(f1110,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f595])).
fof(f11130,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10421,f9071])).
fof(f10421,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2403,f879,f9071,f1110])).
fof(f11129,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10422,f9071])).
fof(f10422,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2404,f879,f9071,f1110])).
fof(f11128,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10423,f9071])).
fof(f10423,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2405,f879,f9071,f1110])).
fof(f11127,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10424,f9071])).
fof(f10424,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2406,f879,f9071,f1110])).
fof(f11126,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10430,f10465])).
fof(f10430,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f1792,f9071,f1110])).
fof(f11125,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10440,f10465])).
fof(f10440,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f9071,f9071,f1110])).
fof(f11124,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10442,f9071])).
fof(f10442,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2401,f879,f9071,f1110])).
fof(f11123,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10443,f9071])).
fof(f10443,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2402,f879,f9071,f1110])).
fof(f11122,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10444,f9071])).
fof(f10444,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2403,f879,f9071,f1110])).
fof(f11121,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10446,f9071])).
fof(f10446,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2405,f879,f9071,f1110])).
fof(f11120,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10447,f9071])).
fof(f10447,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2406,f879,f9071,f1110])).
fof(f11119,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10453,f10465])).
fof(f10453,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f1792,f9071,f1110])).
fof(f11118,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10463,f10465])).
fof(f10463,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f9071,f9071,f1110])).
fof(f11117,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10468,f9071])).
fof(f10468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f1792,f9071,f1114])).
fof(f11116,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10470,f9071])).
fof(f10470,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1114])).
fof(f11115,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10471,f1792])).
fof(f10471,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f9071,f9071,f1114])).
fof(f3395,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f3384,f1105])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f592])).
fof(f3384,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f1792,f1110])).
fof(f11114,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10472,f1792])).
fof(f10472,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f3632,plain,(
  ~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1792,f1792,f3429])).
fof(f3429,plain,(
  ( ! [X6,X7] : (~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f3418,f879])).
fof(f3418,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X7,X6)) )),
  inference(superposition,[],[f1117,f2434])).
fof(f1117,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f597])).
fof(f11113,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10473,f1792])).
fof(f10473,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f3645,plain,(
  ~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1792,f1792,f3430])).
fof(f3430,plain,(
  ( ! [X8,X9] : (~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X9,X8) | p(s(bool,X9)) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f3419,f879])).
fof(f3419,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | p(s(bool,X9)) | ~sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X9,X8)) )),
  inference(superposition,[],[f1117,f2435])).
fof(f11112,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10474,f1792])).
fof(f10474,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f1792,f9071,f1114])).
fof(f11111,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10476,f9071])).
fof(f10476,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f1792,f9071,f1114])).
fof(f11110,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10477,f9071])).
fof(f10477,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f1792,f9071,f1114])).
fof(f11109,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10482,f1792])).
fof(f10482,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1114])).
fof(f11108,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10483,f1792])).
fof(f10483,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1114])).
fof(f11107,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10485,f9071])).
fof(f10485,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f11106,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10486,f9071])).
fof(f10486,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f11105,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10491,f9071])).
fof(f10491,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f9071,f1114])).
fof(f11104,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10493,f9071])).
fof(f10493,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1114])).
fof(f11103,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10495,f1792])).
fof(f10495,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f9071,f9071,f1114])).
fof(f11102,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10496,f1792])).
fof(f10496,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f11101,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10497,f1792])).
fof(f10497,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f11100,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10498,f1792])).
fof(f10498,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f9071,f1114])).
fof(f11099,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10500,f9071])).
fof(f10500,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f1792,f9071,f1114])).
fof(f11098,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10501,f9071])).
fof(f10501,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f1792,f9071,f1114])).
fof(f11097,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10506,f1792])).
fof(f10506,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1114])).
fof(f11096,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10507,f1792])).
fof(f10507,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1114])).
fof(f11095,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10509,f9071])).
fof(f10509,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f11094,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10510,f9071])).
fof(f10510,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f11093,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10516,f1792])).
fof(f10516,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f1792,f9071,f1114])).
fof(f11092,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10517,f1792])).
fof(f10517,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f1792,f9071,f1114])).
fof(f11091,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10518,f1792])).
fof(f10518,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f1792,f9071,f1114])).
fof(f11090,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10519,f9071])).
fof(f10519,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f1114])).
fof(f11089,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10521,f1792])).
fof(f10521,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f9071,f9071,f1114])).
fof(f11088,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10522,f1792])).
fof(f10522,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f11087,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10523,f1792])).
fof(f10523,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f11086,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10525,f1792])).
fof(f10525,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f1792,f9071,f1114])).
fof(f11085,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10526,f1792])).
fof(f10526,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f1792,f9071,f1114])).
fof(f11084,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10527,f1792])).
fof(f10527,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f1792,f9071,f1114])).
fof(f11083,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10533,f9071])).
fof(f10533,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f9071,f1114])).
fof(f11082,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10534,f1792])).
fof(f10534,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f9071,f9071,f1114])).
fof(f11081,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10535,f1792])).
fof(f10535,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3632,f9071,f9071,f1114])).
fof(f11080,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10536,f1792])).
fof(f10536,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3645,f9071,f9071,f1114])).
fof(f11079,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10542,f9071])).
fof(f10542,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2626,f879,f9071,f1117])).
fof(f2626,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f2402,f879,f1759])).
fof(f1759,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1106])).
fof(f1106,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f592])).
fof(f11078,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10543,f9071])).
fof(f10543,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2627,f879,f9071,f1117])).
fof(f2627,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f2403,f879,f1759])).
fof(f11077,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10544,f9071])).
fof(f10544,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2543,f879,f9071,f1117])).
fof(f2543,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f879,f1116])).
fof(f1116,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f597])).
fof(f11076,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10545,f9071])).
fof(f10545,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2629,f879,f9071,f1117])).
fof(f2629,plain,(
  ( ! [X0] : (sP9('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2405,f879,f1759])).
fof(f11075,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10546,f9071])).
fof(f10546,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2630,f879,f9071,f1117])).
fof(f2630,plain,(
  ( ! [X0] : (sP9('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2406,f879,f1759])).
fof(f11074,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10619,f9071])).
fof(f10619,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2530,f879,f9071,f1117])).
fof(f2530,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f879,f1115])).
fof(f1115,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f597])).
fof(f11073,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10620,f9071])).
fof(f10620,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2626,f879,f9071,f1117])).
fof(f11072,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10621,f9071])).
fof(f10621,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2627,f879,f9071,f1117])).
fof(f11071,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10623,f9071])).
fof(f10623,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2629,f879,f9071,f1117])).
fof(f11070,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10624,f9071])).
fof(f10624,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2630,f879,f9071,f1117])).
fof(f11069,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10629,f1792])).
fof(f10629,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2560,f9071,f1120])).
fof(f1120,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f599])).
fof(f599,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f598])).
fof(f598,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f2560,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1792,f879,f1132])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f356,f397,f396])).
fof(f396,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f356,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X121,X122,X123] : ((p(s(bool,X123)) <=> (p(s(bool,X122)) => p(s(bool,X121)))) <=> ((p(s(bool,X123)) | p(s(bool,X122))) & (p(s(bool,X123)) | ~p(s(bool,X121))) & (~p(s(bool,X122)) | p(s(bool,X121)) | ~p(s(bool,X123)))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f11068,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10630,f9071])).
fof(f10630,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2603,f9071,f1120])).
fof(f2603,plain,(
  ( ! [X0,X1] : (sP11('T_0',X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f2418,f879,f1758])).
fof(f1758,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1123])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f600])).
fof(f2418,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f963,f1128])).
fof(f1128,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f602])).
fof(f602,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f601])).
fof(f601,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f396])).
fof(f11067,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10631,f9071])).
fof(f10631,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2576,f9071,f1120])).
fof(f2576,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f879,f1133])).
fof(f1133,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f605])).
fof(f11066,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10632,f9071])).
fof(f10632,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2605,f9071,f1120])).
fof(f2605,plain,(
  ( ! [X0] : (sP11('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2410,f879,f1758])).
fof(f2410,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1)) )),
  inference(unit_resulting_resolution,[],[f875,f1129])).
fof(f1129,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f603])).
fof(f11065,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10633,f9071])).
fof(f10633,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2606,f9071,f1120])).
fof(f2606,plain,(
  ( ! [X0] : (sP11('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2411,f879,f1758])).
fof(f2411,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X1)) )),
  inference(unit_resulting_resolution,[],[f876,f1129])).
fof(f11064,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10635,f1792])).
fof(f10635,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2691,f9071,f1120])).
fof(f2691,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1792,f2564])).
fof(f2564,plain,(
  ( ! [X0,X1] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2561,f879])).
fof(f2561,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1,X0)) )),
  inference(superposition,[],[f1132,f2434])).
fof(f11063,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10636,f1792])).
fof(f10636,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2697,f9071,f1120])).
fof(f2697,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1792,f2565])).
fof(f2565,plain,(
  ( ! [X2,X3] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X3,X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2562,f879])).
fof(f2562,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X3,X2)) )),
  inference(superposition,[],[f1132,f2435])).
fof(f11062,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10637,f9071])).
fof(f10637,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4889,f9071,f1120])).
fof(f4889,plain,(
  ( ! [X0,X1] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f2418,f2610])).
fof(f2610,plain,(
  ( ! [X0,X1] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | ~sP10('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(subsumption_resolution,[],[f2607,f879])).
fof(f2607,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP10('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(superposition,[],[f1758,f2434])).
fof(f11061,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10638,f9071])).
fof(f10638,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4984,f9071,f1120])).
fof(f4984,plain,(
  ( ! [X0,X1] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f2418,f2611])).
fof(f2611,plain,(
  ( ! [X2,X3] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | ~sP10('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(subsumption_resolution,[],[f2608,f879])).
fof(f2608,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~sP10('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(superposition,[],[f1758,f2435])).
fof(f11060,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10639,f9071])).
fof(f10639,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2703,f9071,f1120])).
fof(f2703,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2580])).
fof(f2580,plain,(
  ( ! [X0,X1] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2577,f879])).
fof(f2577,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(superposition,[],[f1133,f2434])).
fof(f11059,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10640,f9071])).
fof(f10640,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2709,f9071,f1120])).
fof(f2709,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2581])).
fof(f2581,plain,(
  ( ! [X2,X3] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2578,f879])).
fof(f2578,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(superposition,[],[f1133,f2435])).
fof(f11058,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10641,f9071])).
fof(f10641,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4891,f9071,f1120])).
fof(f4891,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2410,f2610])).
fof(f11057,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10642,f9071])).
fof(f10642,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4892,f9071,f1120])).
fof(f4892,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2411,f2610])).
fof(f11056,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10643,f9071])).
fof(f10643,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4986,f9071,f1120])).
fof(f4986,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2410,f2611])).
fof(f11055,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10644,f9071])).
fof(f10644,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4987,f9071,f1120])).
fof(f4987,plain,(
  ( ! [X0] : (sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2411,f2611])).
fof(f11054,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10645,f1792])).
fof(f10645,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f4436,f9071,f1120])).
fof(f4436,plain,(
  ( ! [X0] : (sP11('const.combin.I_1'(s(bool,'T_0')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1792,f2563])).
fof(f2563,plain,(
  ( ! [X6,X4,X5] : (sP11('const.combin.I_1'(s(bool,X4)),X6,X5) | p(s(bool,X5)) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f1132,f964])).
fof(f11053,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10646,f9071])).
fof(f10646,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4459,f9071,f1120])).
fof(f4459,plain,(
  ( ! [X0] : (sP11('const.combin.I_1'(s(bool,'T_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f879,f2579])).
fof(f2579,plain,(
  ( ! [X6,X4,X5] : (sP11('const.combin.I_1'(s(bool,X4)),X5,X6) | ~p(s(bool,X5)) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f1133,f964])).
fof(f11052,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10657,f1792])).
fof(f10657,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3504,f9071,f1120])).
fof(f3504,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2412,f879,f1792,f1126])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f600])).
fof(f2412,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1792,f1130])).
fof(f1130,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f603])).
fof(f11051,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10658,f9071])).
fof(f10658,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2691,f9071,f1120])).
fof(f11050,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10659,f9071])).
fof(f10659,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2697,f9071,f1120])).
fof(f11049,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10660,f9071])).
fof(f10660,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4889,f9071,f1120])).
fof(f11048,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10661,f9071])).
fof(f10661,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4984,f9071,f1120])).
fof(f11047,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10662,f9071])).
fof(f10662,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2703,f9071,f1120])).
fof(f11046,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10663,f9071])).
fof(f10663,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2709,f9071,f1120])).
fof(f11045,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10664,f9071])).
fof(f10664,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4891,f9071,f1120])).
fof(f11044,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10665,f9071])).
fof(f10665,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4892,f9071,f1120])).
fof(f11043,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10666,f9071])).
fof(f10666,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4986,f9071,f1120])).
fof(f11042,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10667,f9071])).
fof(f10667,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f4987,f9071,f1120])).
fof(f11041,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10668,f9071])).
fof(f10668,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4436,f9071,f1120])).
fof(f11040,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10669,f9071])).
fof(f10669,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4459,f9071,f1120])).
fof(f11039,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10670,f9071])).
fof(f10670,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3505,f9071,f1120])).
fof(f3505,plain,(
  ( ! [X0] : (sP11('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f2419,f879,f1792,f1126])).
fof(f2419,plain,(
  ( ! [X2,X0,X1] : (sP10('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f963,f1130])).
fof(f11038,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10714,f10733])).
fof(f10733,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f1130])).
fof(f10714,plain,(
  ~sP10('F_0','T_0','T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f2556,f879,f9071,f1126])).
fof(f2556,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f1792,f1121])).
fof(f1121,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f600])).
fof(f11037,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10729,f10733])).
fof(f10729,plain,(
  ( ! [X0] : (~sP10('T_0',X0,'T_0')) ) | $spl44),
  inference(unit_resulting_resolution,[],[f879,f879,f9071,f1127])).
fof(f1127,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f603])).
fof(f11036,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10742,f1792])).
fof(f10742,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f2556,f879,f9071,f1131])).
fof(f1131,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f605])).
fof(f11035,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10744,f9071])).
fof(f10744,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3658,f879,f9071,f1131])).
fof(f3658,plain,(
  ~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f3581])).
fof(f3581,plain,(
  ( ! [X6,X7] : (~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f3571,f879])).
fof(f3571,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X7,X6)) )),
  inference(superposition,[],[f1134,f2434])).
fof(f1134,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f605])).
fof(f11034,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10745,f9071])).
fof(f10745,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3671,f879,f9071,f1131])).
fof(f3671,plain,(
  ~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f3582])).
fof(f3582,plain,(
  ( ! [X8,X9] : (~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X9,X8) | p(s(bool,X9)) | ~p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f3572,f879])).
fof(f3572,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP11('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X9,X8)) )),
  inference(superposition,[],[f1134,f2435])).
fof(f11033,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10751,f1792])).
fof(f10751,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3533,f879,f9071,f1131])).
fof(f3533,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f3523,f1122])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f600])).
fof(f3523,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f879,f1792,f1127])).
fof(f11032,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10752,f1792])).
fof(f10752,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3658,f879,f9071,f1131])).
fof(f11031,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10753,f1792])).
fof(f10753,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(unit_resulting_resolution,[],[f3671,f879,f9071,f1131])).
fof(f11030,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10826,f9071])).
fof(f10826,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2379,f9071,f1795])).
fof(f1795,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1017,f1792])).
fof(f1017,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP2(X0,X1)))),
  inference(flattening,[],[f557])).
fof(f557,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP2(X0,X1)))),
  inference(nnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0,X1] : (sP2(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f2379,plain,(
  ( ! [X0] : (sP2(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f875,f1018])).
fof(f1018,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f558])).
fof(f11029,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10827,f9071])).
fof(f10827,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2381,f9071,f1795])).
fof(f2381,plain,(
  ( ! [X0] : (sP2(X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f876,f1018])).
fof(f11028,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10828,f9071])).
fof(f10828,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3248,f9071,f1795])).
fof(f3248,plain,(
  ( ! [X21,X19,X20] : (sP2(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(subsumption_resolution,[],[f3216,f879])).
fof(f3216,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP2(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1018,f966])).
fof(f11027,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10829,f9071])).
fof(f10829,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4482,f9071,f1795])).
fof(f4482,plain,(
  ( ! [X0] : (sP2(X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f2851])).
fof(f2851,plain,(
  ( ! [X8,X7] : (sP2(X8,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X7))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f1018,f965])).
fof(f11026,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10830,f9071])).
fof(f10830,plain,(
  p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4516,f9071,f1795])).
fof(f4516,plain,(
  ( ! [X0] : (sP2(X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f3115])).
fof(f3115,plain,(
  ( ! [X15,X16] : (sP2(X16,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X15))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f1018,f1794])).
fof(f11025,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10831,f9071])).
fof(f10831,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f2721,f9071,f1795])).
fof(f2721,plain,(
  ( ! [X0] : (sP2(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f2423])).
fof(f2423,plain,(
  ( ! [X2,X3] : (sP2(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1018,f964])).
fof(f11024,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10832,f9071])).
fof(f10832,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2377,f9071,f1795])).
fof(f2377,plain,(
  ( ! [X0] : (sP2('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f1019])).
fof(f1019,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f558])).
fof(f11023,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10833,f9071])).
fof(f10833,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2378,f9071,f1795])).
fof(f2378,plain,(
  ( ! [X0] : (sP2('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f875,f1019])).
fof(f11022,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10834,f9071])).
fof(f10834,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2380,f9071,f1795])).
fof(f2380,plain,(
  ( ! [X0] : (sP2('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f876,f1019])).
fof(f11021,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10835,f9071])).
fof(f10835,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3249,f9071,f1795])).
fof(f3249,plain,(
  ( ! [X24,X23,X22] : (sP2('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(subsumption_resolution,[],[f3217,f879])).
fof(f3217,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP2('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(superposition,[],[f1019,f966])).
fof(f11020,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10836,f9071])).
fof(f10836,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4499,f9071,f1795])).
fof(f4499,plain,(
  ( ! [X0] : (sP2(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2852])).
fof(f2852,plain,(
  ( ! [X10,X9] : (sP2(happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9)),X10) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f1019,f965])).
fof(f11019,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10837,f9071])).
fof(f10837,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4533,f9071,f1795])).
fof(f4533,plain,(
  ( ! [X0] : (sP2(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f3116])).
fof(f3116,plain,(
  ( ! [X17,X18] : (sP2(happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X17)),X18) | ~p(s(bool,X17))) )),
  inference(superposition,[],[f1019,f1794])).
fof(f11018,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10838,f9071])).
fof(f10838,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2715,f9071,f1795])).
fof(f2715,plain,(
  ( ! [X0] : (sP2('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2422])).
fof(f2422,plain,(
  ( ! [X0,X1] : (sP2('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1019,f964])).
fof(f11017,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10843,f9071])).
fof(f10843,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2376,f9071,f1795])).
fof(f2376,plain,(
  ( ! [X0] : (sP2(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f1018])).
fof(f11016,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10844,f9071])).
fof(f10844,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2379,f9071,f1795])).
fof(f11015,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10845,f9071])).
fof(f10845,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2381,f9071,f1795])).
fof(f11014,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10846,f9071])).
fof(f10846,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3248,f9071,f1795])).
fof(f11013,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10847,f9071])).
fof(f10847,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4482,f9071,f1795])).
fof(f11012,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10848,f9071])).
fof(f10848,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f4516,f9071,f1795])).
fof(f11011,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10849,f9071])).
fof(f10849,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2721,f9071,f1795])).
fof(f11010,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10851,f9071])).
fof(f10851,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2378,f9071,f1795])).
fof(f11009,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10852,f9071])).
fof(f10852,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f2380,f9071,f1795])).
fof(f11008,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10853,f9071])).
fof(f10853,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f3249,f9071,f1795])).
fof(f11007,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10854,f9071])).
fof(f10854,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4499,f9071,f1795])).
fof(f11006,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10855,f9071])).
fof(f10855,plain,(
  p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f4533,f9071,f1795])).
fof(f11005,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10856,f9071])).
fof(f10856,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl44),
  inference(unit_resulting_resolution,[],[f2715,f9071,f1795])).
fof(f11004,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10877,f9071])).
fof(f10877,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2557,f9071,f3429])).
fof(f2557,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f2534])).
fof(f2534,plain,(
  ( ! [X0,X1] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1,X0) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2531,f879])).
fof(f2531,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X1,X0)) )),
  inference(superposition,[],[f1115,f2434])).
fof(f11003,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10879,f9071])).
fof(f10879,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5155,f9071,f3429])).
fof(f5155,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f2402,f2634])).
fof(f2634,plain,(
  ( ! [X0,X1] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | ~sP8('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(subsumption_resolution,[],[f2631,f879])).
fof(f2631,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(superposition,[],[f1759,f2434])).
fof(f11002,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10881,f9071])).
fof(f10881,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5156,f9071,f3429])).
fof(f5156,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f2403,f2634])).
fof(f11001,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10885,f9071])).
fof(f10885,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5158,f9071,f3429])).
fof(f5158,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2405,f2634])).
fof(f11000,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10887,f9071])).
fof(f10887,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5159,f9071,f3429])).
fof(f5159,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2406,f2634])).
fof(f10999,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10893,f9071])).
fof(f10893,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5155,f9071,f3429])).
fof(f10998,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10895,f9071])).
fof(f10895,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5156,f9071,f3429])).
fof(f10997,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10897,f9071])).
fof(f10897,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2593,f9071,f3429])).
fof(f2593,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2547])).
fof(f2547,plain,(
  ( ! [X0,X1] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f2544,f879])).
fof(f2544,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0,X1)) )),
  inference(superposition,[],[f1116,f2434])).
fof(f10996,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10899,f9071])).
fof(f10899,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5158,f9071,f3429])).
fof(f10995,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10901,f9071])).
fof(f10901,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5159,f9071,f3429])).
fof(f10994,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10905,f9071])).
fof(f10905,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2585,f9071,f3430])).
fof(f2585,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f2535])).
fof(f2535,plain,(
  ( ! [X2,X3] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X3,X2) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2532,f879])).
fof(f2532,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X3,X2)) )),
  inference(superposition,[],[f1115,f2435])).
fof(f10993,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10907,f9071])).
fof(f10907,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5382,f9071,f3430])).
fof(f5382,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) )),
  inference(unit_resulting_resolution,[],[f2402,f2635])).
fof(f2635,plain,(
  ( ! [X2,X3] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | ~sP8('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(subsumption_resolution,[],[f2632,f879])).
fof(f2632,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~sP8('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(superposition,[],[f1759,f2435])).
fof(f10992,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10909,f9071])).
fof(f10909,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5383,f9071,f3430])).
fof(f5383,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) )),
  inference(unit_resulting_resolution,[],[f2403,f2635])).
fof(f10991,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10913,f9071])).
fof(f10913,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5385,f9071,f3430])).
fof(f5385,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),X0)) )),
  inference(unit_resulting_resolution,[],[f2405,f2635])).
fof(f10990,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10915,f9071])).
fof(f10915,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5386,f9071,f3430])).
fof(f5386,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X0)) )),
  inference(unit_resulting_resolution,[],[f2406,f2635])).
fof(f10989,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10921,f9071])).
fof(f10921,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5382,f9071,f3430])).
fof(f10988,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10923,f9071])).
fof(f10923,plain,(
  p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))) | $spl44),
  inference(unit_resulting_resolution,[],[f5383,f9071,f3430])).
fof(f10987,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10925,f9071])).
fof(f10925,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2671,f9071,f3430])).
fof(f2671,plain,(
  ( ! [X0] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f879,f2548])).
fof(f2548,plain,(
  ( ! [X2,X3] : (sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2545,f879])).
fof(f2545,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP9('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),X2,X3)) )),
  inference(superposition,[],[f1116,f2435])).
fof(f10986,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10927,f9071])).
fof(f10927,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5385,f9071,f3430])).
fof(f10985,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f10929,f9071])).
fof(f10929,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f5386,f9071,f3430])).
fof(f10948,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f4987,f9071,f3582])).
fof(f10946,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f4986,f9071,f3582])).
fof(f10944,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f2709,f9071,f3582])).
fof(f10939,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f4892,f9071,f3581])).
fof(f10937,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f4891,f9071,f3581])).
fof(f10935,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f2703,f9071,f3581])).
fof(f10930,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5386,f9071,f3430])).
fof(f10928,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5385,f9071,f3430])).
fof(f10926,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2671,f9071,f3430])).
fof(f10924,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5383,f9071,f3430])).
fof(f10922,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5382,f9071,f3430])).
fof(f10920,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2585,f9071,f3430])).
fof(f10916,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5386,f9071,f3430])).
fof(f10914,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5385,f9071,f3430])).
fof(f10912,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2671,f9071,f3430])).
fof(f10910,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5383,f9071,f3430])).
fof(f10908,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5382,f9071,f3430])).
fof(f10906,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2585,f9071,f3430])).
fof(f10902,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5159,f9071,f3429])).
fof(f10900,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5158,f9071,f3429])).
fof(f10898,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2593,f9071,f3429])).
fof(f10896,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5156,f9071,f3429])).
fof(f10894,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5155,f9071,f3429])).
fof(f10892,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2557,f9071,f3429])).
fof(f10888,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5159,f9071,f3429])).
fof(f10886,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f5158,f9071,f3429])).
fof(f10884,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2593,f9071,f3429])).
fof(f10882,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5156,f9071,f3429])).
fof(f10880,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f5155,f9071,f3429])).
fof(f10878,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2557,f9071,f3429])).
fof(f10871,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3671,f9071,f2565])).
fof(f10869,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3658,f9071,f2564])).
fof(f10860,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2376,f9071,f9071,f1795])).
fof(f10858,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2376,f1792,f9071,f1795])).
fof(f10842,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2376,f9071,f9071,f1795])).
fof(f10840,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2377,f1792,f9071,f1795])).
fof(f10812,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f1060,f9071,f1757])).
fof(f1757,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1756])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1152])).
fof(f1152,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X82,X119,X84,X120] : ((s(bool,X82) = s(bool,X119) & (p(s(bool,X119)) => s(bool,X84) = s(bool,X120))) => ((p(s(bool,X82)) => p(s(bool,X84))) <=> (p(s(bool,X119)) => p(s(bool,X120)))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.bool.IMP_CONG')).
fof(f10810,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f1794,f9071,f1757])).
fof(f10808,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f965,f9071,f1757])).
fof(f10806,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f964,f9071,f1757])).
fof(f10801,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f966,f9071,f1757])).
fof(f10797,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f2435,f9071,f1757])).
fof(f10795,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f2434,f9071,f1757])).
fof(f10792,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2576,f879,f879,f9071,f1134])).
fof(f10773,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3533,f879,f9071,f1132])).
fof(f10766,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2556,f879,f9071,f9071,f1131])).
fof(f10759,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3533,f879,f1792,f9071,f1131])).
fof(f10749,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2556,f9071,f879,f9071,f1131])).
fof(f10747,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2556,f1792,f879,f9071,f1131])).
fof(f10734,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3523,f9071,f1130])).
fof(f10732,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3523,f9071,f1128])).
fof(f10730,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2409,f879,f879,f9071,f1127])).
fof(f2409,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f879,f1129])).
fof(f10715,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2412,f2556,f879,f9071,f1126])).
fof(f10695,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2576,f879,f9071,f1121])).
fof(f10674,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2560,f9071,f9071,f1120])).
fof(f10672,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3504,f1792,f9071,f1120])).
fof(f10651,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2560,f9071,f9071,f1120])).
fof(f10649,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2560,f1792,f9071,f1120])).
fof(f10628,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2530,f9071,f879,f9071,f1117])).
fof(f10626,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2530,f1792,f879,f9071,f1117])).
fof(f10584,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2530,f879,f9071,f9071,f1117])).
fof(f10560,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2543,f879,f1792,f9071,f1117])).
fof(f10540,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f9071,f1114])).
fof(f10538,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f1792,f9071,f9071,f1114])).
fof(f10531,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f1792,f9071,f1114])).
fof(f10529,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3395,f1792,f1792,f9071,f1114])).
fof(f10514,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f9071,f1114])).
fof(f10512,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f9071,f9071,f1114])).
fof(f10505,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f1792,f9071,f1114])).
fof(f10503,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f1792,f9071,f1114])).
fof(f10490,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f9071,f9071,f9071,f1114])).
fof(f10488,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f1792,f9071,f9071,f1114])).
fof(f10481,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f9071,f1792,f9071,f1114])).
fof(f10479,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f1792,f1792,f9071,f1114])).
fof(f10466,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3384,f9071,f1113])).
fof(f10464,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2401,f879,f9071,f9071,f1110])).
fof(f10454,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2401,f879,f1792,f9071,f1110])).
fof(f10441,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2401,f879,f9071,f9071,f1110])).
fof(f10431,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2404,f879,f1792,f9071,f1110])).
fof(f10417,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2407,f9071,f9071,f1109])).
fof(f10415,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2406,f9071,f9071,f1109])).
fof(f10413,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2405,f9071,f9071,f1109])).
fof(f10411,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f2404,f9071,f9071,f1109])).
fof(f10409,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f2403,f9071,f9071,f1109])).
fof(f10407,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2526,f2402,f9071,f9071,f1109])).
fof(f10405,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2401,f9071,f9071,f1109])).
fof(f10397,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2407,f1792,f9071,f1109])).
fof(f10393,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2523,f2401,f1792,f9071,f1109])).
fof(f10375,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2407,f2526,f9071,f9071,f1109])).
fof(f10373,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2407,f2523,f9071,f9071,f1109])).
fof(f10360,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2407,f2526,f1792,f9071,f1109])).
fof(f10322,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2543,f879,f9071,f1104])).
fof(f10298,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2530,f879,f9071,f1103])).
fof(f10274,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f4235,f1060,f9071,f1102])).
fof(f10272,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3091,f1794,f9071,f1102])).
fof(f10270,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2968,f965,f9071,f1102])).
fof(f10268,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2967,f964,f9071,f1102])).
fof(f10266,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2966,f2473,f9071,f1102])).
fof(f10264,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f5254,f968,f9071,f1102])).
fof(f10261,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3203,f966,f9071,f1102])).
fof(f3203,plain,(
  ( ! [X0,X1] : (sP7('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f879,f966,f1100])).
fof(f10257,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2965,f2435,f9071,f1102])).
fof(f2965,plain,(
  sP7('T_0','T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)))),
  inference(unit_resulting_resolution,[],[f879,f2435,f1100])).
fof(f10255,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2964,f2434,f9071,f1102])).
fof(f2964,plain,(
  sP7('T_0','T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)))),
  inference(unit_resulting_resolution,[],[f879,f2434,f1100])).
fof(f10253,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f4231,f1060,f9071,f1102])).
fof(f10251,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3088,f1794,f9071,f1102])).
fof(f10249,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2963,f965,f9071,f1102])).
fof(f10247,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2962,f964,f9071,f1102])).
fof(f10245,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2961,f2473,f9071,f1102])).
fof(f10243,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f5251,f968,f9071,f1102])).
fof(f10240,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3201,f966,f9071,f1102])).
fof(f3201,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f879,f966,f1100])).
fof(f10236,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2960,f2435,f9071,f1102])).
fof(f2960,plain,(
  sP7('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0')),
  inference(unit_resulting_resolution,[],[f879,f2435,f1100])).
fof(f10234,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2959,f2434,f9071,f1102])).
fof(f2959,plain,(
  sP7('T_0','const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0')),
  inference(unit_resulting_resolution,[],[f879,f2434,f1100])).
fof(f10984,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10222])).
fof(f10222,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6166,f9071,f1102])).
fof(f6166,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2394,f2385,f2390,f1792,f1085])).
fof(f10983,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10221])).
fof(f10221,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6165,f9071,f1102])).
fof(f6165,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2393,f2385,f2390,f1792,f1085])).
fof(f10982,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10220])).
fof(f10220,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6092,f9071,f1102])).
fof(f6092,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK98)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2398,f2382,f2389,f879,f1084])).
fof(f1084,plain,(
  ( ! [X2,X0,X1] : (~sP5(X1,X2,X0) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | sP7(X0,X1,X2) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f580])).
fof(f10981,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10219])).
fof(f10219,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6089,f9071,f1102])).
fof(f6089,plain,(
  sP7('const.typeSystem.check_freevars_3'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',sK97)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2398,f2382,f2388,f879,f1084])).
fof(f10980,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10203])).
fof(f10203,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl44),
  inference(unit_resulting_resolution,[],[f2980,f9071,f1102])).
fof(f2980,plain,(
  ( ! [X3] : (sP7('T_0','const.combin.I_1'(s(bool,X3)),'const.combin.I_1'(s(bool,X3)))) )),
  inference(superposition,[],[f2962,f964])).
fof(f10979,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10171])).
fof(f10171,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6164,f9071,f1102])).
fof(f6164,plain,(
  sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2392,f2385,f2390,f1792,f1085])).
fof(f10978,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f10169])).
fof(f10169,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f6086,f9071,f1102])).
fof(f6086,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2398,f2382,f2387,f879,f1084])).
fof(f10161,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2988,f2426,f9071,f1099])).
fof(f10159,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2987,f2425,f9071,f1099])).
fof(f10157,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2985,f2426,f9071,f1099])).
fof(f10155,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2984,f2425,f9071,f1099])).
fof(f10104,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f9071,f1095])).
fof(f10102,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f1792,f9071,f9071,f1095])).
fof(f10091,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f1792,f9071,f1095])).
fof(f10089,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f1792,f1792,f9071,f1095])).
fof(f10066,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f9071,f1095])).
fof(f10064,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f1792,f9071,f9071,f1095])).
fof(f10053,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f9071,f1792,f9071,f1095])).
fof(f10051,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f1792,f1792,f9071,f1095])).
fof(f10028,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f9071,f9071,f1095])).
fof(f10026,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f1792,f9071,f9071,f1095])).
fof(f10015,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2392,f9071,f1792,f9071,f1095])).
fof(f10013,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2398,f1792,f1792,f9071,f1095])).
fof(f9990,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3314,f9071,f1094])).
fof(f3314,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f879,f1091])).
fof(f9988,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3314,f9071,f1093])).
fof(f9986,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2387,f879,f879,f9071,f1091])).
fof(f9974,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3291,f9071,f1090])).
fof(f3291,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f879,f1792,f879,f1087])).
fof(f9972,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f3291,f9071,f1089])).
fof(f9970,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f4020,f879,f879,f9071,f1087])).
fof(f4020,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f3201,f1082])).
fof(f1082,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f580])).
fof(f9850,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f2415,f9071,f1086])).
fof(f9838,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f2415,f9071,f1086])).
fof(f9803,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f2390,f9071,f1086])).
fof(f9790,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f2390,f9071,f1086])).
fof(f9735,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f2387,f9071,f1086])).
fof(f9729,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f2387,f9071,f1086])).
fof(f9713,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f2387,f9071,f1086])).
fof(f9711,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f2387,f9071,f1086])).
fof(f9489,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f2415,f9071,f1085])).
fof(f9477,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f2415,f9071,f1085])).
fof(f9442,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f2390,f9071,f1085])).
fof(f9429,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f2390,f9071,f1085])).
fof(f9374,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2413,f2387,f9071,f1085])).
fof(f9368,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2385,f2387,f9071,f1085])).
fof(f9352,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2988,f2382,f2387,f9071,f1085])).
fof(f9350,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2395,f2987,f2382,f2387,f9071,f1085])).
fof(f9205,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2964,f879,f879,f9071,f1083])).
fof(f9146,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2484,f879,f9071,f1037])).
fof(f2484,plain,(
  ~sP3('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f879,f1033])).
fof(f1033,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f565])).
fof(f9142,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2484,f879,f9071,f1034])).
fof(f9138,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2488,f9071,f9071,f1032])).
fof(f9136,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2499,f1792,f9071,f1032])).
fof(f9124,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2488,f9071,f9071,f1032])).
fof(f9122,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2488,f1792,f9071,f1032])).
fof(f9108,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2426,f9071,f1004])).
fof(f9106,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2425,f9071,f1004])).
fof(f9104,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2426,f9071,f1004])).
fof(f9102,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2425,f9071,f1004])).
fof(f9098,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2426,f9071,f1004])).
fof(f9096,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f1792,f2425,f9071,f1004])).
fof(f9094,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2426,f9071,f1004])).
fof(f9092,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f9071,f2425,f9071,f1004])).
fof(f9088,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2425,f9071,f920])).
fof(f9086,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f2434,f9071,f915])).
fof(f9084,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f879,f9071])).
fof(f9080,plain,(
  ( ! [X68,X69] : (sP11('const.combin.I_1'(s(bool,'T_0')),X68,X69) | p(s(bool,X69))) ) | $spl46),
  inference(cnf_transformation,[],[f9080_D])).
fof(f9080_D,plain,(
  ( ! [X68,X69] : (sP11('const.combin.I_1'(s(bool,'T_0')),X68,X69) | p(s(bool,X69))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9071,plain,(
  ( ! [X59] : (~p(s(bool,X59))) ) | $spl44),
  inference(cnf_transformation,[],[f9071_D])).
fof(f9071_D,plain,(
  ( ! [X59] : (~p(s(bool,X59))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f9079,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X67)) | sP11('const.combin.I_1'(s(bool,'T_0')),X68,X69) | p(s(bool,X69))) )),
  inference(subsumption_resolution,[],[f9078,f919])).
fof(f919,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f504])).
fof(f9078,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X67)) | sP11('const.combin.I_1'(s(bool,'T_0')),X68,X69) | p(s(bool,X69)) | s(bool,'F_0') = s(bool,X67)) )),
  inference(forward_demodulation,[],[f8946,f964])).
fof(f8946,plain,(
  ( ! [X68,X69,X67] : (sP11('const.combin.I_1'(s(bool,'T_0')),X68,X69) | p(s(bool,X69)) | ~p(s(bool,'const.combin.I_1'(s(bool,X67)))) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f2563,f2739])).
fof(f2739,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f907,f964])).
fof(f907,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb',aHL_BOOL_CASES)).
fof(f9070,plain,(
  ( ! [X60] : (sP3(X60,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X60))) ) | $spl42),
  inference(cnf_transformation,[],[f9070_D])).
fof(f9070_D,plain,(
  ( ! [X60] : (sP3(X60,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X60))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9069,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | sP3(X60,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X60))) )),
  inference(subsumption_resolution,[],[f9068,f919])).
fof(f9068,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | sP3(X60,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X60)) | s(bool,'F_0') = s(bool,X59)) )),
  inference(forward_demodulation,[],[f8943,f964])).
fof(f8943,plain,(
  ( ! [X59,X60] : (sP3(X60,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X60)) | ~p(s(bool,'const.combin.I_1'(s(bool,X59)))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f2502,f2739])).
fof(f8203,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl38),
  inference(cnf_transformation,[],[f8203_D])).
fof(f8203_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f8205,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl41),
  inference(cnf_transformation,[],[f8205_D])).
fof(f8205_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f8112,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f923,f5097])).
fof(f5097,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f905,f3776])).
fof(f3776,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f3766,f1794])).
fof(f3766,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(unit_resulting_resolution,[],[f1794,f924])).
fof(f924,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X12))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.numeral.numeral_distrib|split|18')).
fof(f905,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X12] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X12))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.numeral.numeral_distrib|split|16')).
fof(f923,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f506])).
fof(f8199,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl34),
  inference(cnf_transformation,[],[f8199_D])).
fof(f8199_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f8201,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(cnf_transformation,[],[f8201_D])).
fof(f8201_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f8104,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f923,f5097])).
fof(f3760,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl30),
  inference(cnf_transformation,[],[f3760_D])).
fof(f3760_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3762,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl33),
  inference(cnf_transformation,[],[f3762_D])).
fof(f3762_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3752,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f923,f881])).
fof(f881,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f124])).
fof(f124,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f3754,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(cnf_transformation,[],[f3754_D])).
fof(f3754_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3756,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl29),
  inference(cnf_transformation,[],[f3756_D])).
fof(f3756_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f3745,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f923,f1793])).
fof(f1793,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f904,f881])).
fof(f904,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X12] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.numeral.numeral_distrib|split|11')).
fof(f2375,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f2374])).
fof(f2374,plain,(
  s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1769,f1770])).
fof(f1789,plain,(
  p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK101)))) | $spl22),
  inference(cnf_transformation,[],[f1789_D])).
fof(f1789_D,plain,(
  p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK101)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1791,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))) | $spl24),
  inference(cnf_transformation,[],[f1791_D])).
fof(f1791_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1777,plain,(
  sP0 | $spl8),
  inference(cnf_transformation,[],[f1777_D])).
fof(f1777_D,plain,(
  sP0 <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f888,plain,(
  sP0 | p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))) | p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK101))))),
  inference(cnf_transformation,[],[f499])).
fof(f499,plain,(
  ((p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK101)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK101))))))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102))))))) | sP0) & (~sP0 | (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK101,sK102])],[f498])).
fof(f498,plain,(
  (? [X3] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) | ? [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))) | sP0) & (~sP0 | (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))))),
  inference(rectify,[],[f497])).
fof(f497,plain,(
  (? [X1] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) | ? [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0))))))) | sP0) & (~sP0 | (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))))),
  inference(flattening,[],[f496])).
fof(f496,plain,(
  ((? [X1] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) | ? [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))) | sP0) & (~sP0 | (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))) <=> sP0),
  inference(definition_folding,[],[f320,f380])).
fof(f380,plain,(
  sP0 <=> (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f320,plain,(
  (! [X1] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))) <=> (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  (! [X1] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) => p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) & ! [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))))) <=> (! [X3] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) => p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  (! [X8] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X8)))) => p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X8))))))) & ! [X49] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X49)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X49)))))))) <=> (! [X8] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X8)))) => p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X8))))))) & ! [X49] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X49)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X49))))))))),
  file('thm.type_eDeterm.type_eDeterm_2_prove_ltb','thm.inferProps.inferProps_|49|33_prove_ltb')).
fof(f1785,plain,(
  ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK101)))))) | $spl19),
  inference(cnf_transformation,[],[f1785_D])).
fof(f1785_D,plain,(
  ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK101)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f889,plain,(
  sP0 | p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))) | ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK101))))))),
  inference(cnf_transformation,[],[f499])).
fof(f1787,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))))) | $spl21),
  inference(cnf_transformation,[],[f1787_D])).
fof(f1787_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f890,plain,(
  sP0 | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))))) | p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK101))))),
  inference(cnf_transformation,[],[f499])).
fof(f891,plain,(
  sP0 | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK102)))))) | ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK101))))))),
  inference(cnf_transformation,[],[f499])).
fof(f1776,plain,(
  ~sP0 | $spl9),
  inference(cnf_transformation,[],[f1776_D])).
fof(f1776_D,plain,(
  ~sP0 <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1783,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3)))))) | ~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3))))) ) | $spl16),
  inference(cnf_transformation,[],[f1783_D])).
fof(f1783_D,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3)))))) | ~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f892,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1)))))) | ~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) | ~sP0) )),
  inference(cnf_transformation,[],[f499])).
fof(f1782,plain,(
  ( ! [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))) ) | $spl14),
  inference(cnf_transformation,[],[f1782_D])).
fof(f1782_D,plain,(
  ( ! [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f893,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) | ~sP0) )),
  inference(cnf_transformation,[],[f499])).
fof(f882,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3)))))) | ~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | ~sP0) )),
  inference(cnf_transformation,[],[f495])).
fof(f495,plain,(
  (~sP0 | (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))) & ((p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK99)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK99))))))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100))))))) | sP0)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK99,sK100])],[f494])).
fof(f494,plain,(
  (~sP0 | (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))) & (? [X1] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X1)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X1))))))) | ? [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X0))))))) | sP0)),
  inference(rectify,[],[f493])).
fof(f493,plain,(
  (~sP0 | (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))) & (? [X3] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) | ? [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))) | sP0)),
  inference(flattening,[],[f492])).
fof(f492,plain,(
  (~sP0 | (! [X3] : (~p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) | p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) & ! [X2] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2))))))))) & ((? [X3] : (p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3)))) & ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',X3))))))) | ? [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))))))) | sP0)),
  inference(nnf_transformation,[],[f380])).
fof(f883,plain,(
  ( ! [X2] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X2)))) | ~sP0) )),
  inference(cnf_transformation,[],[f495])).
fof(f1779,plain,(
  p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK99)))) | $spl10),
  inference(cnf_transformation,[],[f1779_D])).
fof(f1779_D,plain,(
  p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK99)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1781,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))) | $spl12),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f884,plain,(
  sP0 | p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))) | p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK99))))),
  inference(cnf_transformation,[],[f495])).
fof(f1773,plain,(
  ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK99)))))) | $spl5),
  inference(cnf_transformation,[],[f1773_D])).
fof(f1773_D,plain,(
  ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK99)))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f885,plain,(
  sP0 | p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))) | ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK99))))))),
  inference(cnf_transformation,[],[f495])).
fof(f1775,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))))) | $spl7),
  inference(cnf_transformation,[],[f1775_D])).
fof(f1775_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f886,plain,(
  sP0 | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))))) | p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',sK99))))),
  inference(cnf_transformation,[],[f495])).
fof(f887,plain,(
  sP0 | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK100)))))) | ~p(s(bool,happ(s(fun('type.infer_t.infer_t',bool),'const.infer.check_t_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))),s('type.infer_t.infer_t',happ(s(fun('type.ast.t','type.infer_t.infer_t'),'const.infer.infer_type_subst_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.ast.t',sK99))))))),
  inference(cnf_transformation,[],[f495])).
fof(f1768,plain,(
  s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) | $spl0),
  inference(cnf_transformation,[],[f1768_D])).
fof(f1768_D,plain,(
  s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1770,plain,(
  s('type.ast.t',sK97) = s('type.ast.t',sK98) | $spl2),
  inference(cnf_transformation,[],[f1770_D])).
fof(f1770_D,plain,(
  s('type.ast.t',sK97) = s('type.ast.t',sK98) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f877,plain,(
  s('type.ast.t',sK97) = s('type.ast.t',sK98) | s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) = s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98)))),
  inference(cnf_transformation,[],[f491])).
fof(f1769,plain,(
  s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) | $spl1),
  inference(cnf_transformation,[],[f1769_D])).
fof(f1769_D,plain,(
  s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1771,plain,(
  s('type.ast.t',sK97) != s('type.ast.t',sK98) | $spl3),
  inference(cnf_transformation,[],[f1771_D])).
fof(f1771_D,plain,(
  s('type.ast.t',sK97) != s('type.ast.t',sK98) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f878,plain,(
  s('type.ast.t',sK97) != s('type.ast.t',sK98) | s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK97))) != s('type.infer_t.infer_t','const.inferProps.unconvert_t_1'(s('type.ast.t',sK98)))),
  inference(cnf_transformation,[],[f491])).
% SZS output end Proof for thm.type_eDeterm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 683
% Passive clauses: 1890
% Generated clauses: 11684
% Final active clauses: 677
% Final passive clauses: 1080
% Input formulas: 133
% Initial clauses: 871
% 
% Duplicate literals: 13
% Trivial inequalities: 17
% Fw subsumption resolutions: 480
% Fw demodulations: 1455
% Bw demodulations: 6
% Condensations: 21
% 
% Simple tautologies: 73
% Equational tautologies: 22
% Forward subsumptions: 7103
% Fw demodulations to eq. taut.: 46
% 
% Unit resulting resolution: 3635
% Factoring: 1
% Forward superposition: 2833
% Backward superposition: 2298
% Self superposition: 33
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 20
% Split components: 48
% SAT solver clauses: 2476
% SAT solver unit clauses: 1371
% SAT solver binary clauses: 753
% 
% Sat splits: 20
% Sat splitting refutations: 544
% 
% Memory used [KB]: 12537
% Time elapsed: 0.384 s
% ------------------------------
% ------------------------------
% Success in time 1.72 s
