Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.decSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.decSem
% SZS output start Proof for thm.decSem
fof(f16660,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f16656,f12876,f119,f170,f120,f172,f121,f174,f122,f176,f168,f15597,f123,f177,f175,f173,f171,f169,f14465,f14467])).
fof(f14467,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f14466,f171])).
fof(f14466,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f14417,f140])).
fof(f140,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num',X1) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X22,X40,X21,X41] : s('type.num.num',X22) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X22),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X40),s('type.ffi.ffi_state'(X6),X21),s('type.list.list'('type.option.option'('type.conSem.v')),X41)))))),
  file('thm.decSem.state_|49||49|','thm.decSem.state_clock')).
fof(f14417,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))))) | ($spl0 | $spl6 | $spl8)),
  inference(superposition,[],[f140,f12878])).
fof(f12878,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f12877,f174])).
fof(f12877,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f168,f176])).
fof(f14465,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f14464,f173])).
fof(f14464,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f14414,f143])).
fof(f143,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X22,X40,X21,X41] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X40) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X22),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X40),s('type.ffi.ffi_state'(X6),X21),s('type.list.list'('type.option.option'('type.conSem.v')),X41)))))),
  file('thm.decSem.state_|49||49|','thm.decSem.state_refs')).
fof(f14414,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))))) | ($spl0 | $spl6 | $spl8)),
  inference(superposition,[],[f143,f12878])).
fof(f169,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) | $spl1),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f171,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) | $spl3),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f173,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | $spl5),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f175,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | $spl7),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f177,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) != s('type.list.list'('type.option.option'('type.conSem.v')),sK8) | $spl9),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) != s('type.list.list'('type.option.option'('type.conSem.v')),sK8) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f123,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) != s('type.list.list'('type.option.option'('type.conSem.v')),sK8) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8)))),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  (s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) | (s('type.num.num',sK1) = s('type.num.num',sK5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) & s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) & s('type.list.list'('type.option.option'('type.conSem.v')),sK4) = s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) & (s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | s('type.list.list'('type.option.option'('type.conSem.v')),sK4) != s('type.list.list'('type.option.option'('type.conSem.v')),sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f112])).
fof(f112,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) = s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),X8))) & (s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) != s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) | s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X7) | s('type.list.list'('type.option.option'('type.conSem.v')),X4) != s('type.list.list'('type.option.option'('type.conSem.v')),X8)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) = s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),X8))) & (s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) != s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) | (s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X7) | s('type.list.list'('type.option.option'('type.conSem.v')),X4) != s('type.list.list'('type.option.option'('type.conSem.v')),X8))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) = s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) <~> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),X8)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) = s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),X8)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4))) = s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.conSem.v')),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),X8)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.list.list'('type.option.option'('type.conSem.v')),X10))) = s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12),s('type.ffi.ffi_state'(X6),X13),s('type.list.list'('type.option.option'('type.conSem.v')),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X13) & s('type.list.list'('type.option.option'('type.conSem.v')),X10) = s('type.list.list'('type.option.option'('type.conSem.v')),X14)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.list.list'('type.option.option'('type.conSem.v')),X10))) = s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12),s('type.ffi.ffi_state'(X6),X13),s('type.list.list'('type.option.option'('type.conSem.v')),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X13) & s('type.list.list'('type.option.option'('type.conSem.v')),X10) = s('type.list.list'('type.option.option'('type.conSem.v')),X14)))),
  file('thm.decSem.state_|49||49|','thm.decSem.state_|49||49|')).
fof(f15597,plain,(
  $false | ($spl0 | $spl7 | $spl8)),
  inference(subsumption_resolution,[],[f15596,f175])).
fof(f15596,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f15544,f141])).
fof(f141,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X22,X40,X21,X41] : s('type.ffi.ffi_state'(X6),X21) = s('type.ffi.ffi_state'(X6),'const.decSem.state_ffi_1'(s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X22),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X40),s('type.ffi.ffi_state'(X6),X21),s('type.list.list'('type.option.option'('type.conSem.v')),X41)))))),
  file('thm.decSem.state_|49||49|','thm.decSem.state_ffi')).
fof(f15544,plain,(
  s('type.ffi.ffi_state'(sK0),sK7) = s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))))) | ($spl0 | $spl8)),
  inference(superposition,[],[f141,f12877])).
fof(f168,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) | $spl0),
  inference(cnf_transformation,[],[f168_D])).
fof(f168_D,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f176,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) = s('type.list.list'('type.option.option'('type.conSem.v')),sK8) | $spl8),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) = s('type.list.list'('type.option.option'('type.conSem.v')),sK8) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f122,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) = s('type.list.list'('type.option.option'('type.conSem.v')),sK8) | s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8)))),
  inference(cnf_transformation,[],[f113])).
fof(f174,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | $spl6),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f121,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8)))),
  inference(cnf_transformation,[],[f113])).
fof(f172,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | $spl4),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f120,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8)))),
  inference(cnf_transformation,[],[f113])).
fof(f170,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | $spl2),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f119,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) = s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK8)))),
  inference(cnf_transformation,[],[f113])).
fof(f12876,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f12875])).
fof(f12875,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f12874,f170])).
fof(f12874,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f12873,f172])).
fof(f12873,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl1 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f12872,f174])).
fof(f12872,plain,(
  s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) != s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f169,f176])).
fof(f16656,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f16655,f177])).
fof(f16655,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK4) = s('type.list.list'('type.option.option'('type.conSem.v')),sK8) | $spl0),
  inference(forward_demodulation,[],[f16600,f142])).
fof(f142,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.conSem.v')),X4) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.conSem.v')),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X22,X40,X21,X41] : s('type.list.list'('type.option.option'('type.conSem.v')),X41) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X6),'const.decSem.recordtype|46|state_4'(s('type.num.num',X22),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X40),s('type.ffi.ffi_state'(X6),X21),s('type.list.list'('type.option.option'('type.conSem.v')),X41)))))),
  file('thm.decSem.state_|49||49|','thm.decSem.state_globals')).
fof(f16600,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK8) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),'const.decSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.conSem.v')),sK4))))) | $spl0),
  inference(superposition,[],[f142,f168])).
% SZS output end Proof for thm.decSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 115
% Passive clauses: 1925
% Generated clauses: 16295
% Final active clauses: 96
% Final passive clauses: 1713
% Input formulas: 37
% Initial clauses: 49
% 
% Duplicate literals: 539
% Trivial inequalities: 234
% Fw subsumption resolutions: 1715
% Fw demodulations: 1748
% Bw demodulations: 6
% 
% Simple tautologies: 195
% Equational tautologies: 39
% Forward subsumptions: 9478
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 113
% Factoring: 1
% Forward superposition: 6947
% Backward superposition: 4492
% Self superposition: 234
% Equality factoring: 95
% Equality resolution: 20
% 
% Split clauses: 185
% Split components: 373
% SAT solver clauses: 513
% SAT solver unit clauses: 27
% SAT solver binary clauses: 103
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 185
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 507
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 11897
% Time elapsed: 0.464 s
% ------------------------------
% ------------------------------
% Success in time 0.47 s
