Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (33323)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 3112
% Generated clauses: 22443
% Final active clauses: 123
% Final passive clauses: 2854
% Input formulas: 37
% Initial clauses: 51
% 
% Duplicate literals: 704
% Trivial inequalities: 354
% Fw subsumption resolutions: 2822
% Fw demodulations: 4597
% Bw demodulations: 4
% 
% Simple tautologies: 368
% Equational tautologies: 45
% Forward subsumptions: 10213
% 
% Binary resolution: 105
% Factoring: 1
% Forward superposition: 7074
% Backward superposition: 6630
% Self superposition: 201
% Equality factoring: 96
% Equality resolution: 37
% 
% Split clauses: 166
% Split components: 338
% SAT solver clauses: 359
% SAT solver unit clauses: 39
% SAT solver binary clauses: 129
% 
% Sat splits: 166
% 
% TWLsolver clauses: 359
% TWLsolver calls for satisfiability: 89
% 
% Memory used [KB]: 14967
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (33336)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 1070
% Generated clauses: 8137
% Final active clauses: 95
% Final passive clauses: 817
% Input formulas: 37
% Initial clauses: 51
% 
% Duplicate literals: 22
% Trivial inequalities: 62
% Fw subsumption resolutions: 156
% Fw demodulations: 1842
% Bw demodulations: 120
% Condensations: 18
% 
% Simple tautologies: 16
% Equational tautologies: 6
% Forward subsumptions: 3611
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 17
% 
% Binary resolution: 104
% Factoring: 1
% Forward superposition: 2387
% Backward superposition: 2000
% Self superposition: 57
% Equality factoring: 2
% Equality resolution: 28
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 1230
% SAT solver unit clauses: 68
% SAT solver binary clauses: 34
% SAT solver learnt clauses: 6
% SAT solver learnt literals: 21
% 
% Sat splits: 15
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1224
% TWLsolver calls for satisfiability: 358
% 
% Memory used [KB]: 18933
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% (33382)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 390
% Passive clauses: 5560
% Generated clauses: 157598
% Final active clauses: 155
% Final passive clauses: 792
% Input formulas: 37
% Initial clauses: 51
% Discarded non-redundant clauses: 4391
% 
% Duplicate literals: 5699
% Trivial inequalities: 1935
% Fw subsumption resolutions: 18562
% Fw demodulations: 26311
% Bw demodulations: 2895
% Condensations: 15759
% Global subsumptions: 78
% 
% Simple tautologies: 1648
% Equational tautologies: 303
% Forward subsumptions: 81269
% Fw demodulations to eq. taut.: 43
% Bw demodulations to eq. taut.: 183
% 
% Unit resulting resolution: 1661
% Factoring: 1
% Forward superposition: 42381
% Backward superposition: 37602
% Self superposition: 800
% Equality factoring: 398
% Equality resolution: 113
% 
% Split clauses: 67
% Split components: 134
% SAT solver clauses: 5909
% SAT solver unit clauses: 192
% SAT solver binary clauses: 964
% 
% Sat splits: 67
% Sat splitting refutations: 135
% 
% Memory used [KB]: 78804
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f6151,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f83,f112,f111,f82,f114,f81,f116,f110,f821,f824,f822,f828,f831,f834,f835,f1309,f1311,f1314,f1315,f1310,f1447,f2656,f2664,f2665,f2706,f2711,f2076,f2067,f2077,f2064,f2718,f2024,f2015,f2025,f2013,f2719,f1254,f1252,f898,f2716,f2724,f774,f2759,f3152,f3153,f3156,f3155,f3157,f3185,f5132,f5133,f5137,f5138,f830,f829,f5466,f5498,f5527,f5525,f5650,f5654,f5656,f5660,f5662,f5664,f5730,f5731,f5732,f5733,f5734,f5735,f5886,f5887,f5891,f5889,f5892,f5893,f5963,f5651,f5964,f5653,f5655,f5965,f5657,f5966,f5661,f5967,f5663,f5968,f5659,f5665,f6081,f6088,f6109,f4966,f4954,f2717,f6110,f2712,f6111,f4736,f6113,f6114,f2714,f6112,f6150])).
fof(f6150,plain,(
  $false | ($spl0 | $spl32)),
  inference(equality_resolution,[],[f6080])).
fof(f6080,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),X0)) ) | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f6021,f99])).
fof(f99,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.closProps.ref_rel_simp|split|0',aHL_FALSITY)).
fof(f6021,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),X0)) ) | ($spl0 | $spl32)),
  inference(backward_demodulation,[],[f5908,f2778])).
fof(f2778,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),X0) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),X0))),s('type.closSem.ref'(sK1),sK2))))) ) | $spl0),
  inference(superposition,[],[f110,f90])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) != s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X15,X37] : (s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X15))) = s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X37))) <=> s('type.list.list'(X6),X15) = s('type.list.list'(X6),X37))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closSem.ref_|49||49||split|0')).
fof(f5908,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X13,fun(sK1,bool)),X14),s('type.closSem.ref'(X13),'const.closSem.ValueArray_1'(s('type.list.list'(X13),X15))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl32),
  inference(superposition,[],[f101,f5525])).
fof(f101,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3)))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X12,X13,X14] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X14),s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X12))),s('type.closSem.ref'(X7),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X13)))))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closProps.ref_rel_def|split|2')).
fof(f6112,plain,(
  ( ! [X94,X97,X95,X98,X96] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X94,fun(X95,bool)),X96),s('type.closSem.ref'(X94),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X97))),s('type.closSem.ref'(X95),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X98))))))) ) | $spl58),
  inference(cnf_transformation,[],[f6112_D])).
fof(f6112_D,plain,(
  ( ! [X94,X97,X95,X98,X96] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X94,fun(X95,bool)),X96),s('type.closSem.ref'(X94),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X97))),s('type.closSem.ref'(X95),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X98))))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2714,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl21),
  inference(cnf_transformation,[],[f2714_D])).
fof(f2714_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f6114,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X28,fun(X29,bool)),X30),s('type.closSem.ref'(X28),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X31))),s('type.closSem.ref'(X29),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X32))))))) )),
  inference(subsumption_resolution,[],[f4486,f94])).
fof(f94,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)))))) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) != s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)))))) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3)) & (s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) != s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))))))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)))))) <=> s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)))))) <=> s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X7] : ! [X30,X31,X10] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.closSem.ref'(X6),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X31))),s('type.closSem.ref'(X7),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X30)))))) <=> s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X30) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X31))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closProps.ref_rel_def|split|1')).
fof(f4486,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') != s(bool,'F_0') | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X31) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X32) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X28,fun(X29,bool)),X30),s('type.closSem.ref'(X28),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X31))),s('type.closSem.ref'(X29),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X32))))))) )),
  inference(superposition,[],[f1737,f96])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.closProps.ref_rel_simp|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f1737,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(X2,fun(X3,bool)),X4),s('type.closSem.ref'(X2),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1))),s('type.closSem.ref'(X3),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0))))) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1)) )),
  inference(resolution,[],[f93,f97])).
fof(f97,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.closProps.ref_rel_simp|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f93,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f6113,plain,(
  ( ! [X4,X3] : (s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4)) ) | $spl60),
  inference(cnf_transformation,[],[f6113_D])).
fof(f6113_D,plain,(
  ( ! [X4,X3] : (s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4736,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4)) )),
  inference(duplicate_literal_removal,[],[f4713])).
fof(f4713,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4)) )),
  inference(superposition,[],[f1737,f1738])).
fof(f1738,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X7,fun(X8,bool)),X9),s('type.closSem.ref'(X7),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X6))),s('type.closSem.ref'(X8),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))))) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X6)) )),
  inference(resolution,[],[f93,f96])).
fof(f6111,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X94,fun(X95,bool)),X96),s('type.closSem.ref'(X94),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X97))),s('type.closSem.ref'(X95),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X98))))))) )),
  inference(subsumption_resolution,[],[f4730,f94])).
fof(f4730,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X94,fun(X95,bool)),X96),s('type.closSem.ref'(X94),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X97))),s('type.closSem.ref'(X95),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X98)))))) | s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X97) = s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X98)) )),
  inference(superposition,[],[f97,f1738])).
fof(f2712,plain,(
  ( ! [X26,X24] : (s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24))))) ) | $spl18),
  inference(cnf_transformation,[],[f2712_D])).
fof(f2712_D,plain,(
  ( ! [X26,X24] : (s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6110,plain,(
  ( ! [X57,X55] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X55),X57) = s('type.closSem.ref'(X55),'const.closSem.ValueArray_1'(s('type.list.list'(X55),sK8(X57,X55))))) )),
  inference(subsumption_resolution,[],[f4903,f3624])).
fof(f3624,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X0,bool)),X2),s('type.closSem.ref'(X0),X1),s('type.closSem.ref'(X0),X1)))) | s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8(X1,X0))))) )),
  inference(duplicate_literal_removal,[],[f3614])).
fof(f3614,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X0,bool)),X2),s('type.closSem.ref'(X0),X1),s('type.closSem.ref'(X0),X1)))) | s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8(X1,X0)))) | s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8(X1,X0))))) )),
  inference(superposition,[],[f2204,f108])).
fof(f108,plain,(
  ( ! [X0,X1] : (s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X1,X0)))) | s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8(X1,X0))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8(X1,X0)))) | s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (? [X3] : s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))) | ? [X2] : s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))) | ? [X2] : s('type.closSem.ref'(X0),X1) = s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X35] : (? [X36] : s('type.closSem.ref'(X6),X35) = s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X36))) | ? [X36] : s('type.closSem.ref'(X6),X35) = s('type.closSem.ref'(X6),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X36))))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closSem.ref_nchotomy')).
fof(f2204,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X20,fun(X22,bool)),X23),s('type.closSem.ref'(X20),X21),s('type.closSem.ref'(X22),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X21,X20))))))) | s('type.closSem.ref'(X20),X21) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X21,X20))))) )),
  inference(superposition,[],[f2171,f108])).
fof(f2171,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X3))))))) )),
  inference(equality_resolution,[],[f94])).
fof(f4903,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X55,fun(X55,bool)),X56),s('type.closSem.ref'(X55),X57),s('type.closSem.ref'(X55),X57)))) | s('type.closSem.ref'(X55),X57) = s('type.closSem.ref'(X55),'const.closSem.ValueArray_1'(s('type.list.list'(X55),sK8(X57,X55))))) )),
  inference(superposition,[],[f95,f3976])).
fof(f3976,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X2,fun(X2,bool)),X4),s('type.closSem.ref'(X2),X3),s('type.closSem.ref'(X2),X3))) | s('type.closSem.ref'(X2),X3) = s('type.closSem.ref'(X2),'const.closSem.ValueArray_1'(s('type.list.list'(X2),sK8(X3,X2))))) )),
  inference(resolution,[],[f3624,f98])).
fof(f98,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f78])).
fof(f95,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f2717,plain,(
  ( ! [X2,X0] : (s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) | $spl24),
  inference(cnf_transformation,[],[f2717_D])).
fof(f2717_D,plain,(
  ( ! [X2,X0] : (s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4954,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X3),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)),X3))))) )),
  inference(superposition,[],[f3977,f2187])).
fof(f2187,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X1,fun(X2,bool)),X3),s('type.closSem.ref'(X1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4))),s('type.closSem.ref'(X2),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X4)))))) )),
  inference(resolution,[],[f2171,f98])).
fof(f3977,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(X5,fun(X5,bool)),X7),s('type.closSem.ref'(X5),X6),s('type.closSem.ref'(X5),X6))) | s('type.closSem.ref'(X5),X6) = s('type.closSem.ref'(X5),'const.closSem.ValueArray_1'(s('type.list.list'(X5),sK8(X6,X5))))) )),
  inference(resolution,[],[f3624,f95])).
fof(f4966,plain,(
  ( ! [X11,X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X9),X11) = s('type.closSem.ref'(X9),'const.closSem.ValueArray_1'(s('type.list.list'(X9),sK8(X11,X9))))) )),
  inference(duplicate_literal_removal,[],[f4956])).
fof(f4956,plain,(
  ( ! [X11,X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X9),X11) = s('type.closSem.ref'(X9),'const.closSem.ValueArray_1'(s('type.list.list'(X9),sK8(X11,X9)))) | s('type.closSem.ref'(X9),X11) = s('type.closSem.ref'(X9),'const.closSem.ValueArray_1'(s('type.list.list'(X9),sK8(X11,X9))))) )),
  inference(superposition,[],[f3977,f3976])).
fof(f6109,plain,(
  ( ! [X34,X32] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X32),X34) = s('type.closSem.ref'(X32),'const.closSem.ValueArray_1'(s('type.list.list'(X32),sK8(X34,X32))))) )),
  inference(subsumption_resolution,[],[f4963,f3624])).
fof(f4963,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X32),X34) = s('type.closSem.ref'(X32),'const.closSem.ValueArray_1'(s('type.list.list'(X32),sK8(X34,X32)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X32,fun(X32,bool)),X33),s('type.closSem.ref'(X32),X34),s('type.closSem.ref'(X32),X34))))) )),
  inference(superposition,[],[f3977,f98])).
fof(f6088,plain,(
  $false | ($spl0 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6060,f2714])).
fof(f6060,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl32)),
  inference(superposition,[],[f2772,f5908])).
fof(f2772,plain,(
  s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))) | $spl0),
  inference(resolution,[],[f110,f98])).
fof(f6081,plain,(
  $false | ($spl0 | $spl21 | $spl32)),
  inference(subsumption_resolution,[],[f6053,f2714])).
fof(f6053,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl32)),
  inference(superposition,[],[f5908,f2772])).
fof(f5665,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl50),
  inference(cnf_transformation,[],[f5665_D])).
fof(f5665_D,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5659,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | $spl44),
  inference(cnf_transformation,[],[f5659_D])).
fof(f5659_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5968,plain,(
  ( ! [X144] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl32),
  inference(forward_demodulation,[],[f5938,f5525])).
fof(f5938,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3977,f5525])).
fof(f5663,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl48),
  inference(cnf_transformation,[],[f5663_D])).
fof(f5663_D,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5967,plain,(
  ( ! [X143] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl32),
  inference(forward_demodulation,[],[f5937,f5525])).
fof(f5937,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3976,f5525])).
fof(f5661,plain,(
  ( ! [X142] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))))) ) | $spl46),
  inference(cnf_transformation,[],[f5661_D])).
fof(f5661_D,plain,(
  ( ! [X142] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5966,plain,(
  ( ! [X142] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))))) ) | $spl32),
  inference(forward_demodulation,[],[f5936,f5525])).
fof(f5936,plain,(
  ( ! [X142] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3624,f5525])).
fof(f5657,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))))))) ) | $spl42),
  inference(cnf_transformation,[],[f5657_D])).
fof(f5657_D,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5965,plain,(
  ( ! [X140,X141] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))))))) ) | $spl32),
  inference(forward_demodulation,[],[f5935,f5525])).
fof(f5935,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f2204,f5525])).
fof(f5655,plain,(
  ( ! [X138,X139,X137] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl40),
  inference(cnf_transformation,[],[f5655_D])).
fof(f5655_D,plain,(
  ( ! [X138,X139,X137] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5653,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | $spl38),
  inference(cnf_transformation,[],[f5653_D])).
fof(f5653_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5964,plain,(
  ( ! [X138,X139,X137] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(forward_demodulation,[],[f5934,f5525])).
fof(f5934,plain,(
  ( ! [X138,X139,X137] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f1534,f5525])).
fof(f1534,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'const.list.LIST_REL_3'(s(fun(X48,fun(X46,bool)),X49),s('type.list.list'(X48),X50),s('type.list.list'(X46),sK8(X47,X46)))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X48,fun(X46,bool)),X49),s('type.closSem.ref'(X48),'const.closSem.ValueArray_1'(s('type.list.list'(X48),X50))),s('type.closSem.ref'(X46),X47))) | s('type.closSem.ref'(X46),X47) = s('type.closSem.ref'(X46),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X47,X46))))) )),
  inference(superposition,[],[f84,f108])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2)))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X9,X10] : s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X9),s('type.list.list'(X7),X11))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X9))),s('type.closSem.ref'(X7),'const.closSem.ValueArray_1'(s('type.list.list'(X7),X11)))))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closProps.ref_rel_def|split|0')).
fof(f5651,plain,(
  ( ! [X134,X136,X135] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136)))) ) | $spl36),
  inference(cnf_transformation,[],[f5651_D])).
fof(f5651_D,plain,(
  ( ! [X134,X136,X135] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5963,plain,(
  ( ! [X134,X136,X135] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136)))) ) | $spl32),
  inference(forward_demodulation,[],[f5933,f5525])).
fof(f5933,plain,(
  ( ! [X134,X136,X135] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f1513,f5525])).
fof(f1513,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'const.list.LIST_REL_3'(s(fun(X46,fun(X48,bool)),X49),s('type.list.list'(X46),sK8(X47,X46)),s('type.list.list'(X48),X50))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X46,fun(X48,bool)),X49),s('type.closSem.ref'(X46),X47),s('type.closSem.ref'(X48),'const.closSem.ValueArray_1'(s('type.list.list'(X48),X50))))) | s('type.closSem.ref'(X46),X47) = s('type.closSem.ref'(X46),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X47,X46))))) )),
  inference(superposition,[],[f84,f108])).
fof(f5893,plain,(
  $false | ($spl2 | $spl33 | $spl34)),
  inference(subsumption_resolution,[],[f5564,f5524])).
fof(f5524,plain,(
  s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | $spl33),
  inference(cnf_transformation,[],[f5524_D])).
fof(f5524_D,plain,(
  s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f5564,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | ($spl2 | $spl34)),
  inference(subsumption_resolution,[],[f5557,f100])).
fof(f100,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.closProps.ref_rel_simp|split|0',aHL_TRUTH)).
fof(f5557,plain,(
  ~p(s(bool,'T_0')) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | ($spl2 | $spl34)),
  inference(trivial_inequality_removal,[],[f5538])).
fof(f5538,plain,(
  ~p(s(bool,'T_0')) | s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),sK2) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | ($spl2 | $spl34)),
  inference(superposition,[],[f2877,f5527])).
fof(f2877,plain,(
  ( ! [X18] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(X18,sK1))))) | s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),X18) | s('type.closSem.ref'(sK1),X18) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(X18,sK1))))) ) | $spl2),
  inference(superposition,[],[f112,f108])).
fof(f5892,plain,(
  ( ! [X134,X132,X133] : (s(bool,'const.list.LIST_REL_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.list.list'(X132),X134),s('type.list.list'(sK1),sK8(sK2,sK1)))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.closSem.ref'(X132),'const.closSem.ValueArray_1'(s('type.list.list'(X132),X134))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl56),
  inference(cnf_transformation,[],[f5892_D])).
fof(f5892_D,plain,(
  ( ! [X134,X132,X133] : (s(bool,'const.list.LIST_REL_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.list.list'(X132),X134),s('type.list.list'(sK1),sK8(sK2,sK1)))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.closSem.ref'(X132),'const.closSem.ValueArray_1'(s('type.list.list'(X132),X134))),s('type.closSem.ref'(sK1),sK2)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5889,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)))) | $spl54),
  inference(cnf_transformation,[],[f5889_D])).
fof(f5889_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5891,plain,(
  ( ! [X134,X132,X133] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)))) | s(bool,'const.list.LIST_REL_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.list.list'(X132),X134),s('type.list.list'(sK1),sK8(sK2,sK1)))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.closSem.ref'(X132),'const.closSem.ValueArray_1'(s('type.list.list'(X132),X134))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl33),
  inference(forward_demodulation,[],[f5890,f5799])).
fof(f5799,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl33),
  inference(trivial_inequality_removal,[],[f5798])).
fof(f5798,plain,(
  s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),sK2) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl33),
  inference(superposition,[],[f5524,f108])).
fof(f5890,plain,(
  ( ! [X134,X132,X133] : (s(bool,'const.list.LIST_REL_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.list.list'(X132),X134),s('type.list.list'(sK1),sK8(sK2,sK1)))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.closSem.ref'(X132),'const.closSem.ValueArray_1'(s('type.list.list'(X132),X134))),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1))))) ) | $spl33),
  inference(forward_demodulation,[],[f5845,f3826])).
fof(f3826,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(X0),X1)),X0))) )),
  inference(equality_resolution,[],[f3654])).
fof(f3654,plain,(
  ( ! [X14,X12,X13] : (s('type.closSem.ref'(X12),'const.closSem.ValueArray_1'(s('type.list.list'(X12),X13))) != s('type.closSem.ref'(X12),'const.closSem.ValueArray_1'(s('type.list.list'(X12),X14))) | s('type.list.list'(X12),X14) = s('type.list.list'(X12),sK8('const.closSem.ValueArray_1'(s('type.list.list'(X12),X13)),X12))) )),
  inference(superposition,[],[f89,f3631])).
fof(f3631,plain,(
  ( ! [X0,X3] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(X0),X3)),X0))))) )),
  inference(subsumption_resolution,[],[f3615,f99])).
fof(f3615,plain,(
  ( ! [X0,X3] : (p(s(bool,'F_0')) | s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(X0),X3)),X0))))) )),
  inference(superposition,[],[f2204,f101])).
fof(f89,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) | s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) != s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f5845,plain,(
  ( ! [X134,X132,X133] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.closSem.ref'(X132),'const.closSem.ValueArray_1'(s('type.list.list'(X132),X134))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X132,fun(sK1,bool)),X133),s('type.list.list'(X132),X134),s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1))))) ) | $spl33),
  inference(superposition,[],[f1534,f5799])).
fof(f5887,plain,(
  ( ! [X130,X131,X129] : (s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.list.list'(sK1),sK8(sK2,sK1)),s('type.list.list'(X129),X131))) = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X129),'const.closSem.ValueArray_1'(s('type.list.list'(X129),X131)))))) ) | $spl52),
  inference(cnf_transformation,[],[f5887_D])).
fof(f5887_D,plain,(
  ( ! [X130,X131,X129] : (s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.list.list'(sK1),sK8(sK2,sK1)),s('type.list.list'(X129),X131))) = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X129),'const.closSem.ValueArray_1'(s('type.list.list'(X129),X131)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5886,plain,(
  ( ! [X130,X131,X129] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)))) | s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.list.list'(sK1),sK8(sK2,sK1)),s('type.list.list'(X129),X131))) = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X129),'const.closSem.ValueArray_1'(s('type.list.list'(X129),X131)))))) ) | $spl33),
  inference(forward_demodulation,[],[f5885,f5799])).
fof(f5885,plain,(
  ( ! [X130,X131,X129] : (s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.list.list'(sK1),sK8(sK2,sK1)),s('type.list.list'(X129),X131))) = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X129),'const.closSem.ValueArray_1'(s('type.list.list'(X129),X131))))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1))))) ) | $spl33),
  inference(forward_demodulation,[],[f5844,f3826])).
fof(f5844,plain,(
  ( ! [X130,X131,X129] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X129),'const.closSem.ValueArray_1'(s('type.list.list'(X129),X131))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X129,bool)),X130),s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1)),s('type.list.list'(X129),X131))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1))),sK1))))) ) | $spl33),
  inference(superposition,[],[f1513,f5799])).
fof(f5735,plain,(
  ( ! [X144] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl32),
  inference(forward_demodulation,[],[f5705,f5525])).
fof(f5705,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3977,f5525])).
fof(f5734,plain,(
  ( ! [X143] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl32),
  inference(forward_demodulation,[],[f5704,f5525])).
fof(f5704,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3976,f5525])).
fof(f5733,plain,(
  ( ! [X142] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))))) ) | $spl32),
  inference(forward_demodulation,[],[f5703,f5525])).
fof(f5703,plain,(
  ( ! [X142] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f3624,f5525])).
fof(f5732,plain,(
  ( ! [X140,X141] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))))))) ) | $spl32),
  inference(forward_demodulation,[],[f5702,f5525])).
fof(f5702,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f2204,f5525])).
fof(f5731,plain,(
  ( ! [X138,X139,X137] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(forward_demodulation,[],[f5701,f5525])).
fof(f5701,plain,(
  ( ! [X138,X139,X137] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f1534,f5525])).
fof(f5730,plain,(
  ( ! [X134,X136,X135] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136)))) ) | $spl32),
  inference(forward_demodulation,[],[f5700,f5525])).
fof(f5700,plain,(
  ( ! [X134,X136,X135] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | $spl32),
  inference(superposition,[],[f1513,f5525])).
fof(f5664,plain,(
  ( ! [X144] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5625,f5564])).
fof(f5625,plain,(
  ( ! [X144] : (s(bool,'F_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X144),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f3977,f5564])).
fof(f5662,plain,(
  ( ! [X143] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5624,f5564])).
fof(f5624,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X143),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f3976,f5564])).
fof(f5660,plain,(
  ( ! [X142] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5623,f5564])).
fof(f5623,plain,(
  ( ! [X142] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X142),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f3624,f5564])).
fof(f5656,plain,(
  ( ! [X140,X141] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5622,f5564])).
fof(f5622,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X140,bool)),X141),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X140),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f2204,f5564])).
fof(f5654,plain,(
  ( ! [X138,X139,X137] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5621,f5564])).
fof(f5621,plain,(
  ( ! [X138,X139,X137] : (s(bool,'const.closProps.ref_rel_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.closSem.ref'(X137),'const.closSem.ValueArray_1'(s('type.list.list'(X137),X139))),s('type.closSem.ref'(sK1),sK2))) = s(bool,'const.list.LIST_REL_3'(s(fun(X137,fun(sK1,bool)),X138),s('type.list.list'(X137),X139),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f1534,f5564])).
fof(f5650,plain,(
  ( ! [X134,X136,X135] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)))) | s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136)))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f5620,f5564])).
fof(f5620,plain,(
  ( ! [X134,X136,X135] : (s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X134),'const.closSem.ValueArray_1'(s('type.list.list'(X134),X136))))) = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X134,bool)),X135),s('type.list.list'(sK1),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1)),s('type.list.list'(X134),X136))) | s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1))),sK1))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f1513,f5564])).
fof(f5525,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | $spl32),
  inference(cnf_transformation,[],[f5525_D])).
fof(f5525_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5527,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl34),
  inference(cnf_transformation,[],[f5527_D])).
fof(f5527_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK2,sK1)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5498,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK2,sK1)))) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | $spl0),
  inference(superposition,[],[f2772,f1534])).
fof(f5466,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK8(sK2,sK1)))) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK9(sK2,sK1)))) | $spl0),
  inference(superposition,[],[f1534,f2772])).
fof(f829,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22)))) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20))))) ) | $spl12),
  inference(cnf_transformation,[],[f829_D])).
fof(f829_D,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22)))) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f830,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,X23)) ) | $spl14),
  inference(cnf_transformation,[],[f830_D])).
fof(f830_D,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5138,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,X36) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X33,bool)),X34),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X33),X35)))) | s('type.closSem.ref'(X33),X35) = s('type.closSem.ref'(X33),'const.closSem.ValueArray_1'(s('type.list.list'(X33),sK8(X35,X33))))) ) | $spl8),
  inference(subsumption_resolution,[],[f5063,f96])).
fof(f5063,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,X36) | ~p(s(bool,X36)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X33,bool)),X34),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X33),X35)))) | s('type.closSem.ref'(X33),X35) = s('type.closSem.ref'(X33),'const.closSem.ValueArray_1'(s('type.list.list'(X33),sK8(X35,X33))))) ) | $spl8),
  inference(superposition,[],[f88,f4394])).
fof(f4394,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X18,bool)),X20),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X18),X19))) | s('type.closSem.ref'(X18),X19) = s('type.closSem.ref'(X18),'const.closSem.ValueArray_1'(s('type.list.list'(X18),sK8(X19,X18))))) ) | $spl8),
  inference(superposition,[],[f4042,f108])).
fof(f4042,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X14,bool)),X15),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X14),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X16)))))) ) | $spl8),
  inference(superposition,[],[f101,f4023])).
fof(f4023,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl8),
  inference(subsumption_resolution,[],[f4010,f99])).
fof(f4010,plain,(
  p(s(bool,'F_0')) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl8),
  inference(superposition,[],[f3624,f822])).
fof(f88,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X24,X25] : ((p(s(bool,X24)) => p(s(bool,X25))) => ((p(s(bool,X25)) => p(s(bool,X24))) => s(bool,X24) = s(bool,X25)))),
  file('thm.closProps.ref_rel_simp|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f5137,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') = s(bool,X32) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X29,bool)),X30),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X29),X31)))) | s('type.closSem.ref'(X29),X31) = s('type.closSem.ref'(X29),'const.closSem.ValueArray_1'(s('type.list.list'(X29),sK8(X31,X29))))) ) | $spl8),
  inference(subsumption_resolution,[],[f5062,f96])).
fof(f5062,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') = s(bool,X32) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X29,bool)),X30),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X29),X31)))) | ~p(s(bool,X32)) | s('type.closSem.ref'(X29),X31) = s('type.closSem.ref'(X29),'const.closSem.ValueArray_1'(s('type.list.list'(X29),sK8(X31,X29))))) ) | $spl8),
  inference(superposition,[],[f88,f4394])).
fof(f5133,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,X23) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22))))) ) | $spl8),
  inference(subsumption_resolution,[],[f5044,f96])).
fof(f5044,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,X23) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20)))) | ~p(s(bool,X23)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22))))) ) | $spl8),
  inference(superposition,[],[f4394,f88])).
fof(f5132,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X19) | s('type.closSem.ref'(X16),X18) = s('type.closSem.ref'(X16),'const.closSem.ValueArray_1'(s('type.list.list'(X16),sK8(X18,X16)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X16,bool)),X17),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X16),X18))))) ) | $spl8),
  inference(subsumption_resolution,[],[f5043,f96])).
fof(f5043,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X19) | s('type.closSem.ref'(X16),X18) = s('type.closSem.ref'(X16),'const.closSem.ValueArray_1'(s('type.list.list'(X16),sK8(X18,X16)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X16,bool)),X17),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X16),X18)))) | ~p(s(bool,X19))) ) | $spl8),
  inference(superposition,[],[f4394,f88])).
fof(f3185,plain,(
  $false | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f3176,f100])).
fof(f3176,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl30)),
  inference(superposition,[],[f3157,f2772])).
fof(f3157,plain,(
  ( ! [X116] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),X116),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))) ) | $spl30),
  inference(cnf_transformation,[],[f3157_D])).
fof(f3157_D,plain,(
  ( ! [X116] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),X116),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3155,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0)))) | $spl28),
  inference(cnf_transformation,[],[f3155_D])).
fof(f3155_D,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3156,plain,(
  ( ! [X116] : (s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),X116),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))) ) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f3116,f3053])).
fof(f3053,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0)))) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f3042,f100])).
fof(f3042,plain,(
  ~p(s(bool,'T_0')) | s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f1310,f2772])).
fof(f3116,plain,(
  ( ! [X116] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),X116),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0)))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0))))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f1310,f3053])).
fof(f3153,plain,(
  ( ! [X115] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK0,bool)),X115),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))))))) ) | $spl26),
  inference(cnf_transformation,[],[f3153_D])).
fof(f3153_D,plain,(
  ( ! [X115] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK0,bool)),X115),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3152,plain,(
  ( ! [X115] : (s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK0,bool)),X115),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))))))) ) | ($spl0 | $spl12 | $spl16)),
  inference(forward_demodulation,[],[f3115,f3053])).
fof(f3115,plain,(
  ( ! [X115] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK0,bool)),X115),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)))))) | s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0)))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3)),sK0))),sK0))))) ) | ($spl0 | $spl12 | $spl16)),
  inference(superposition,[],[f829,f3053])).
fof(f2759,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2683,f113])).
fof(f113,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) | $spl5),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f2683,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f110,f122])).
fof(f122,plain,(
  ( ! [X4,X5,X3] : (s(bool,'const.list.LIST_REL_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.list.list'(X3),X5),s('type.list.list'(sK1),sK5))) = s(bool,'const.closProps.ref_rel_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),X5))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl6),
  inference(superposition,[],[f84,f116])).
fof(f774,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) | $spl6),
  inference(superposition,[],[f420,f383])).
fof(f383,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X18,bool)),X20),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X18),X19))) | s('type.closSem.ref'(X18),X19) = s('type.closSem.ref'(X18),'const.closSem.ValueArray_1'(s('type.list.list'(X18),sK8(X19,X18))))) ) | $spl6),
  inference(superposition,[],[f127,f108])).
fof(f127,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X10,bool)),X11),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X10),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X12)))))) ) | $spl6),
  inference(superposition,[],[f101,f116])).
fof(f420,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X47,bool)),X48),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X47),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X49)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f401,f99])).
fof(f401,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X47,bool)),X48),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X47),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X49)))))) ) | $spl6),
  inference(superposition,[],[f97,f127])).
fof(f2724,plain,(
  ( ! [X52,X50] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X50),X52) = s('type.closSem.ref'(X50),'const.closSem.ValueArray_1'(s('type.list.list'(X50),sK8(X52,X50))))) ) | $spl6),
  inference(subsumption_resolution,[],[f786,f820])).
fof(f820,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X38,bool)),X39),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X38),X40)))) | s('type.closSem.ref'(X38),X40) = s('type.closSem.ref'(X38),'const.closSem.ValueArray_1'(s('type.list.list'(X38),sK8(X40,X38))))) ) | $spl6),
  inference(trivial_inequality_removal,[],[f783])).
fof(f783,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X38,bool)),X39),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X38),X40)))) | s('type.closSem.ref'(X38),X40) = s('type.closSem.ref'(X38),'const.closSem.ValueArray_1'(s('type.list.list'(X38),sK8(X40,X38))))) ) | $spl6),
  inference(superposition,[],[f95,f383])).
fof(f786,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X50,bool)),X51),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X50),X52)))) | s('type.closSem.ref'(X50),X52) = s('type.closSem.ref'(X50),'const.closSem.ValueArray_1'(s('type.list.list'(X50),sK8(X52,X50))))) ) | $spl6),
  inference(superposition,[],[f97,f383])).
fof(f2716,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl22),
  inference(cnf_transformation,[],[f2716_D])).
fof(f2716_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f898,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | ($spl6 | $spl8)),
  inference(superposition,[],[f884,f383])).
fof(f884,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X22),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f865,f99])).
fof(f865,plain,(
  ( ! [X22] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X22),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl8),
  inference(superposition,[],[f97,f822])).
fof(f1252,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) | $spl6),
  inference(superposition,[],[f486,f449])).
fof(f449,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X18,fun(sK1,bool)),X20),s('type.closSem.ref'(X18),X19),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(X18),X19) = s('type.closSem.ref'(X18),'const.closSem.ValueArray_1'(s('type.list.list'(X18),sK8(X19,X18))))) ) | $spl6),
  inference(superposition,[],[f128,f108])).
fof(f128,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X13,fun(sK1,bool)),X14),s('type.closSem.ref'(X13),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X15))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl6),
  inference(superposition,[],[f102,f116])).
fof(f102,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X3)))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X3)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))),s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X3)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : ! [X28,X29,X14] : s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X14),s('type.closSem.ref'(X6),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X28))),s('type.closSem.ref'(X7),'const.closSem.ValueArray_1'(s('type.list.list'(X7),X29)))))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closProps.ref_rel_def|split|3')).
fof(f486,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(X47,fun(sK1,bool)),X48),s('type.closSem.ref'(X47),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X49))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f467,f99])).
fof(f467,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(X47,fun(sK1,bool)),X48),s('type.closSem.ref'(X47),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X49))),s('type.closSem.ref'(sK1),sK2)))) ) | $spl6),
  inference(superposition,[],[f97,f128])).
fof(f1254,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | ($spl6 | $spl8)),
  inference(superposition,[],[f884,f449])).
fof(f2719,plain,(
  ( ! [X54,X56] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X54),X56) = s('type.closSem.ref'(X54),'const.closSem.ValueArray_1'(s('type.list.list'(X54),sK8(X56,X54))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1267,f1303])).
fof(f1303,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X42,fun(sK1,bool)),X43),s('type.closSem.ref'(X42),X44),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X42),X44) = s('type.closSem.ref'(X42),'const.closSem.ValueArray_1'(s('type.list.list'(X42),sK8(X44,X42))))) ) | $spl6),
  inference(trivial_inequality_removal,[],[f1264])).
fof(f1264,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X42,fun(sK1,bool)),X43),s('type.closSem.ref'(X42),X44),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X42),X44) = s('type.closSem.ref'(X42),'const.closSem.ValueArray_1'(s('type.list.list'(X42),sK8(X44,X42))))) ) | $spl6),
  inference(superposition,[],[f95,f449])).
fof(f1267,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X54,fun(sK1,bool)),X55),s('type.closSem.ref'(X54),X56),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X54),X56) = s('type.closSem.ref'(X54),'const.closSem.ValueArray_1'(s('type.list.list'(X54),sK8(X56,X54))))) ) | $spl6),
  inference(superposition,[],[f97,f449])).
fof(f2013,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) | $spl6),
  inference(superposition,[],[f533,f127])).
fof(f533,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X18,bool)),X20),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X18),X19))) | s('type.closSem.ref'(X18),X19) = s('type.closSem.ref'(X18),'const.closSem.ValueArray_1'(s('type.list.list'(X18),sK8(X19,X18))))) ) | $spl6),
  inference(superposition,[],[f420,f108])).
fof(f2025,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2014])).
fof(f2014,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3)))) | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3))))) ) | $spl6),
  inference(superposition,[],[f533,f383])).
fof(f2015,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | ($spl6 | $spl8)),
  inference(superposition,[],[f533,f822])).
fof(f2024,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl6),
  inference(duplicate_literal_removal,[],[f2016])).
fof(f2016,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl6),
  inference(superposition,[],[f533,f449])).
fof(f2718,plain,(
  ( ! [X26,X24] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2021,f820])).
fof(f2021,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X24,bool)),X25),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X24),X26))))) ) | $spl6),
  inference(superposition,[],[f533,f96])).
fof(f2064,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X0),'const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),sK8('const.closSem.ByteArray_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X2)),X0))))) ) | $spl6),
  inference(superposition,[],[f573,f128])).
fof(f573,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(X18,fun(sK1,bool)),X20),s('type.closSem.ref'(X18),X19),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(X18),X19) = s('type.closSem.ref'(X18),'const.closSem.ValueArray_1'(s('type.list.list'(X18),sK8(X19,X18))))) ) | $spl6),
  inference(superposition,[],[f486,f108])).
fof(f2077,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2065])).
fof(f2065,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3)))) | s('type.closSem.ref'(X3),X5) = s('type.closSem.ref'(X3),'const.closSem.ValueArray_1'(s('type.list.list'(X3),sK8(X5,X3))))) ) | $spl6),
  inference(superposition,[],[f573,f449])).
fof(f2067,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | ($spl6 | $spl8)),
  inference(superposition,[],[f573,f822])).
fof(f2076,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl6),
  inference(duplicate_literal_removal,[],[f2068])).
fof(f2068,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8(sK2,sK1)))) | $spl6),
  inference(superposition,[],[f573,f383])).
fof(f2711,plain,(
  ( ! [X26,X24] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2073,f1303])).
fof(f2073,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(X24,fun(sK1,bool)),X25),s('type.closSem.ref'(X24),X26),s('type.closSem.ref'(sK1),sK2))))) ) | $spl6),
  inference(superposition,[],[f573,f96])).
fof(f2706,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f2698])).
fof(f2698,plain,(
  s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),sK2) | ($spl2 | $spl4 | $spl6)),
  inference(superposition,[],[f2682,f116])).
fof(f2682,plain,(
  ( ! [X5] : (s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X5)))) ) | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f112,f188])).
fof(f188,plain,(
  ( ! [X0] : (s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X0))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X0))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f178,f116])).
fof(f178,plain,(
  ( ! [X0] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X0)))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X0)))) ) | $spl4),
  inference(superposition,[],[f114,f89])).
fof(f2665,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2636,f172])).
fof(f172,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | $spl4),
  inference(resolution,[],[f114,f98])).
fof(f2636,plain,(
  s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | ($spl1 | $spl6)),
  inference(superposition,[],[f227,f122])).
fof(f227,plain,(
  s(bool,'T_0') != s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))) | $spl1),
  inference(resolution,[],[f111,f97])).
fof(f2664,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2663,f286])).
fof(f286,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(forward_demodulation,[],[f285,f172])).
fof(f285,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | $spl4),
  inference(subsumption_resolution,[],[f262,f100])).
fof(f262,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | $spl4),
  inference(superposition,[],[f95,f172])).
fof(f2663,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f2635,f172])).
fof(f2635,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | ($spl1 | $spl6)),
  inference(superposition,[],[f228,f122])).
fof(f228,plain,(
  s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))) | $spl1),
  inference(resolution,[],[f111,f96])).
fof(f2656,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2655,f286])).
fof(f2655,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f2625,f172])).
fof(f2625,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5))) | ($spl1 | $spl6)),
  inference(superposition,[],[f122,f228])).
fof(f1447,plain,(
  $false | $spl14),
  inference(resolution,[],[f1352,f100])).
fof(f1352,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f99,f830])).
fof(f1310,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X22,fun(sK1,bool)),X23),s('type.closSem.ref'(X22),X24),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22))))) ) | $spl16),
  inference(cnf_transformation,[],[f1310_D])).
fof(f1310_D,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X22,fun(sK1,bool)),X23),s('type.closSem.ref'(X22),X24),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1315,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X33) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X30,fun(sK1,bool)),X31),s('type.closSem.ref'(X30),X32),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X30),X32) = s('type.closSem.ref'(X30),'const.closSem.ValueArray_1'(s('type.list.list'(X30),sK8(X32,X30))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1261,f96])).
fof(f1261,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X33)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X30,fun(sK1,bool)),X31),s('type.closSem.ref'(X30),X32),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X30),X32) = s('type.closSem.ref'(X30),'const.closSem.ValueArray_1'(s('type.list.list'(X30),sK8(X32,X30))))) ) | $spl6),
  inference(superposition,[],[f88,f449])).
fof(f1314,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X26,fun(sK1,bool)),X27),s('type.closSem.ref'(X26),X28),s('type.closSem.ref'(sK1),sK2)))) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1260,f96])).
fof(f1260,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X26,fun(sK1,bool)),X27),s('type.closSem.ref'(X26),X28),s('type.closSem.ref'(sK1),sK2)))) | ~p(s(bool,X29)) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26))))) ) | $spl6),
  inference(superposition,[],[f88,f449])).
fof(f1311,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X26,fun(sK1,bool)),X27),s('type.closSem.ref'(X26),X28),s('type.closSem.ref'(sK1),sK2))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1245,f96])).
fof(f1245,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X26,fun(sK1,bool)),X27),s('type.closSem.ref'(X26),X28),s('type.closSem.ref'(sK1),sK2)))) | ~p(s(bool,X29))) ) | $spl6),
  inference(superposition,[],[f449,f88])).
fof(f1309,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X25) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X22,fun(sK1,bool)),X23),s('type.closSem.ref'(X22),X24),s('type.closSem.ref'(sK1),sK2))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1244,f96])).
fof(f1244,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X25) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22)))) | ~p(s(bool,X25)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X22,fun(sK1,bool)),X23),s('type.closSem.ref'(X22),X24),s('type.closSem.ref'(sK1),sK2))))) ) | $spl6),
  inference(superposition,[],[f449,f88])).
fof(f835,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X26,bool)),X27),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X26),X28)))) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26))))) ) | $spl6),
  inference(subsumption_resolution,[],[f780,f96])).
fof(f780,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X29) | ~p(s(bool,X29)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X26,bool)),X27),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X26),X28)))) | s('type.closSem.ref'(X26),X28) = s('type.closSem.ref'(X26),'const.closSem.ValueArray_1'(s('type.list.list'(X26),sK8(X28,X26))))) ) | $spl6),
  inference(superposition,[],[f88,f383])).
fof(f834,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X22,bool)),X23),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X22),X24)))) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22))))) ) | $spl6),
  inference(subsumption_resolution,[],[f779,f96])).
fof(f779,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X25) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X22,bool)),X23),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X22),X24)))) | ~p(s(bool,X25)) | s('type.closSem.ref'(X22),X24) = s('type.closSem.ref'(X22),'const.closSem.ValueArray_1'(s('type.list.list'(X22),sK8(X24,X22))))) ) | $spl6),
  inference(superposition,[],[f88,f383])).
fof(f831,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X27) | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X24,bool)),X25),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X24),X26))))) ) | $spl6),
  inference(subsumption_resolution,[],[f767,f96])).
fof(f767,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X27) | s('type.closSem.ref'(X24),X26) = s('type.closSem.ref'(X24),'const.closSem.ValueArray_1'(s('type.list.list'(X24),sK8(X26,X24)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X24,bool)),X25),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X24),X26)))) | ~p(s(bool,X27))) ) | $spl6),
  inference(superposition,[],[f383,f88])).
fof(f828,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,X23) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20)))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22))))) ) | $spl6),
  inference(subsumption_resolution,[],[f766,f96])).
fof(f766,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,X23) | s('type.closSem.ref'(X20),X22) = s('type.closSem.ref'(X20),'const.closSem.ValueArray_1'(s('type.list.list'(X20),sK8(X22,X20)))) | ~p(s(bool,X23)) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(X20,bool)),X21),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(X20),X22))))) ) | $spl6),
  inference(superposition,[],[f383,f88])).
fof(f822,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X54),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl8),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X54),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f824,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5)),sK1)))) | $spl10),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5)),sK1)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f821,plain,(
  ( ! [X54] : (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5)),sK1)))) | s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X54),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2)))) ) | $spl6),
  inference(forward_demodulation,[],[f752,f116])).
fof(f752,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.closProps.ref_rel_3'(s(fun(sK1,fun(sK1,bool)),X54),s('type.closSem.ref'(sK1),sK2),s('type.closSem.ref'(sK1),sK2))) | s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK8('const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5)),sK1))))) ) | $spl6),
  inference(superposition,[],[f383,f116])).
fof(f110,plain,(
  p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f116,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) | $spl6),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f81,plain,(
  s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  (p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) | (s('type.closSem.ref'(sK1),sK2) = s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),sK5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))))) & (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) | ! [X5] : (s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) | ? [X6] : (s('type.closSem.ref'(X1),X2) = s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X6)))))) & (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) | ! [X5] : (s('type.closSem.ref'(X1),X2) != s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5)))))))),
  inference(rectify,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) | ? [X5] : (s('type.closSem.ref'(X1),X2) = s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5)))))) & (~p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) | ! [X5] : (s('type.closSem.ref'(X1),X2) != s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5)))))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) <~> ? [X5] : (s('type.closSem.ref'(X1),X2) = s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) <=> ? [X5] : (s('type.closSem.ref'(X1),X2) = s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X3))),s('type.closSem.ref'(X1),X2)))) <=> ? [X5] : (s('type.closSem.ref'(X1),X2) = s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X9))),s('type.closSem.ref'(X7),X8)))) <=> ? [X11] : (s('type.closSem.ref'(X7),X8) = s('type.closSem.ref'(X7),'const.closSem.ValueArray_1'(s('type.list.list'(X7),X11))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X9),s('type.list.list'(X7),X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.closProps.ref_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X9))),s('type.closSem.ref'(X7),X8)))) <=> ? [X11] : (s('type.closSem.ref'(X7),X8) = s('type.closSem.ref'(X7),'const.closSem.ValueArray_1'(s('type.list.list'(X7),X11))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X9),s('type.list.list'(X7),X11))))))),
  file('thm.closProps.ref_rel_simp|split|0','thm.closProps.ref_rel_simp|split|0')).
fof(f114,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) | $spl4),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f82,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK5)))) | p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))),
  inference(cnf_transformation,[],[f73])).
fof(f111,plain,(
  ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f112,plain,(
  ( ! [X5] : (s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X5))))) ) | $spl2),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X5] : (s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X5))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f83,plain,(
  ( ! [X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),X5)))) | s('type.closSem.ref'(sK1),sK2) != s('type.closSem.ref'(sK1),'const.closSem.ValueArray_1'(s('type.list.list'(sK1),X5))) | ~p(s(bool,'const.closProps.ref_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK3))),s('type.closSem.ref'(sK1),sK2))))) )),
  inference(cnf_transformation,[],[f73])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 159
% Passive clauses: 2484
% Generated clauses: 6116
% Final active clauses: 70
% Final passive clauses: 1125
% Input formulas: 37
% Initial clauses: 29
% 
% Selected by SInE selection: 16
% SInE iterations: 1
% 
% Duplicate literals: 72
% Trivial inequalities: 77
% Fw subsumption resolutions: 254
% Fw demodulations: 999
% Bw demodulations: 18
% 
% Simple tautologies: 70
% Equational tautologies: 3
% Forward subsumptions: 2133
% Fw demodulations to eq. taut.: 45
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 115
% Forward superposition: 3199
% Backward superposition: 1202
% Self superposition: 16
% Equality resolution: 33
% 
% Split clauses: 63
% Split components: 126
% SAT solver clauses: 75
% SAT solver unit clauses: 1
% SAT solver binary clauses: 12
% 
% Sat splits: 63
% Sat splitting refutations: 11
% 
% Memory used [KB]: 6268
% Time elapsed: 0.287 s
% ------------------------------
% ------------------------------
% Success in time 12.04 s
