Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.source_to_modProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.source_to_modProof
% SZS output start Proof for thm.source_to_modProof
fof(f7739,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7736,f229,f348,f350,f210,f351])).
fof(f351,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f210,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb',aHL_FALSITY)).
fof(f350,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f348,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f229,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f7736,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f7735,f203])).
fof(f203,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK12))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),sK10)))),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),sK11),s('type.list.list'('type.semanticPrimitives.v'),sK10),s('type.list.list'('type.modSem.v'),sK12)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK12))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),sK10)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f98])).
fof(f98,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X1),s('type.list.list'('type.semanticPrimitives.v'),X0),s('type.list.list'('type.modSem.v'),X2)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X2))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),X0))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X1),s('type.list.list'('type.semanticPrimitives.v'),X0),s('type.list.list'('type.modSem.v'),X2)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),X0))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X7),s('type.list.list'('type.semanticPrimitives.v'),X6),s('type.list.list'('type.modSem.v'),X8)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),X6))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X7),s('type.list.list'('type.semanticPrimitives.v'),X6),s('type.list.list'('type.modSem.v'),X8)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),X6))))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb','thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb')).
fof(f7735,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK12))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),sK10))) | $spl0),
  inference(subsumption_resolution,[],[f7729,f348])).
fof(f7729,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK12))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),sK10)))),
  inference(superposition,[],[f3829,f2914])).
fof(f2914,plain,(
  s(bool,'T_0') = s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),sK11),s('type.list.list'('type.semanticPrimitives.v'),sK10),s('type.list.list'('type.modSem.v'),sK12)))),
  inference(resolution,[],[f202,f219])).
fof(f219,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f202,plain,(
  p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),sK11),s('type.list.list'('type.semanticPrimitives.v'),sK10),s('type.list.list'('type.modSem.v'),sK12))))),
  inference(cnf_transformation,[],[f141])).
fof(f3829,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X0),s('type.list.list'('type.semanticPrimitives.v'),X1),s('type.list.list'('type.modSem.v'),X2)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.semanticPrimitives.v'),X1)))) )),
  inference(superposition,[],[f284,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X0),s('type.list.list'('type.semanticPrimitives.v'),X1),s('type.list.list'('type.modSem.v'),X2))) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.v',fun('type.modSem.v',bool)),'const.source_to_modProof.v_rel_1'(s('type.list.list'('type.option.option'('type.modSem.v')),X0))),s('type.list.list'('type.semanticPrimitives.v'),X1),s('type.list.list'('type.modSem.v'),X2)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X0),s('type.list.list'('type.semanticPrimitives.v'),X1),s('type.list.list'('type.modSem.v'),X2))) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.v',fun('type.modSem.v',bool)),'const.source_to_modProof.v_rel_1'(s('type.list.list'('type.option.option'('type.modSem.v')),X0))),s('type.list.list'('type.semanticPrimitives.v'),X1),s('type.list.list'('type.modSem.v'),X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7,X6,X8] : s(bool,'const.source_to_modProof.vs_rel_3'(s('type.list.list'('type.option.option'('type.modSem.v')),X7),s('type.list.list'('type.semanticPrimitives.v'),X6),s('type.list.list'('type.modSem.v'),X8))) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.v',fun('type.modSem.v',bool)),'const.source_to_modProof.v_rel_1'(s('type.list.list'('type.option.option'('type.modSem.v')),X7))),s('type.list.list'('type.semanticPrimitives.v'),X6),s('type.list.list'('type.modSem.v'),X8)))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb','thm.source_to_modProof.source_to_modProof_|49||48|_prove_ltb')).
fof(f284,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10,X9] : ! [X15,X14,X16] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X9,bool)),X15),s('type.list.list'(X10),X14),s('type.list.list'(X9),X16)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X9),X16))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X10),X14))))),
  file('thm.source_to_modProof.source_to_modProof_|49|2_prove_ltb','thm.list.LIST_REL_LENGTH')).
% SZS output end Proof for thm.source_to_modProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 444
% Passive clauses: 2434
% Generated clauses: 7546
% Final active clauses: 396
% Final passive clauses: 1714
% Input formulas: 39
% Initial clauses: 142
% 
% Duplicate literals: 211
% Trivial inequalities: 73
% Fw subsumption resolutions: 763
% Fw demodulations: 98
% Bw demodulations: 47
% 
% Simple tautologies: 147
% Equational tautologies: 61
% Forward subsumptions: 3638
% 
% Binary resolution: 3410
% Factoring: 858
% Forward superposition: 525
% Backward superposition: 1231
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 32
% 
% Split clauses: 175
% Split components: 351
% SAT solver clauses: 353
% SAT solver unit clauses: 9
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 1
% 
% Sat splits: 175
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 350
% TWLsolver calls for satisfiability: 221
% 
% Memory used [KB]: 3582
% Time elapsed: 0.185 s
% ------------------------------
% ------------------------------
% Success in time 0.21 s
