Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exhSem
% (21353)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 175
% Passive clauses: 3735
% Generated clauses: 20184
% Final active clauses: 166
% Final passive clauses: 3363
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 681
% Trivial inequalities: 514
% Fw subsumption resolutions: 2150
% Fw demodulations: 3162
% Bw demodulations: 6
% 
% Simple tautologies: 538
% Equational tautologies: 56
% Forward subsumptions: 9023
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 319
% Factoring: 1
% Forward superposition: 6880
% Backward superposition: 5977
% Self superposition: 144
% Equality factoring: 90
% Equality resolution: 51
% 
% Split clauses: 216
% Split components: 444
% SAT solver clauses: 501
% SAT solver unit clauses: 10
% SAT solver binary clauses: 99
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 216
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 500
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exhSem
% (21377)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 140
% Passive clauses: 1323
% Generated clauses: 14133
% Final active clauses: 130
% Final passive clauses: 1143
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 405
% Trivial inequalities: 369
% Fw subsumption resolutions: 911
% Fw demodulations: 1751
% Bw demodulations: 42
% Condensations: 727
% Global subsumptions: 8
% 
% Simple tautologies: 245
% Equational tautologies: 44
% Forward subsumptions: 6211
% 
% Binary resolution: 192
% Factoring: 1
% Forward superposition: 4403
% Backward superposition: 3586
% Self superposition: 109
% Equality factoring: 57
% Equality resolution: 32
% 
% Split clauses: 47
% Split components: 94
% SAT solver clauses: 2224
% SAT solver unit clauses: 155
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 47
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2190
% TWLsolver calls for satisfiability: 664
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exhSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exhSem
% SZS output start Proof for thm.exhSem
fof(f120994,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f429,f427,f260,f428,f426,f431,f1384,f2033,f2036,f2038,f2044,f2046,f2049,f2051,f2054,f2056,f2058,f2060,f2062,f2064,f2066,f2082,f2084,f2086,f2088,f2092,f2094,f2096,f2098,f2100,f2102,f2105,f2173,f2174,f2175,f2176,f5121,f5128,f5393,f5394,f5397,f5398,f7685,f7694,f7687,f7736,f7754,f7752,f9660,f9671,f9670,f12345,f12418,f13065,f13067,f13079,f13082,f13085,f13103,f13251,f13253,f13255,f13257,f13259,f13261,f13263,f13265,f13267,f13269,f13271,f13273,f13275,f13277,f13279,f13281,f13283,f13285,f13287,f13289,f20024,f20099,f18432,f20178,f20179,f20181,f20183,f20185,f20187,f20191,f20193,f20195,f20197,f20199,f20201,f20205,f20207,f20209,f20211,f20213,f20215,f20252,f20254,f20256,f20258,f20260,f20262,f20264,f20266,f20318,f20320,f20322,f20330,f21133,f21135,f21143,f21145,f21148,f21155,f21159,f21165,f21166,f21167,f21168,f21169,f21170,f21171,f21172,f21173,f21174,f21175,f21176,f22303,f22352,f28261,f28419,f28420,f28422,f28423,f28424,f28425,f25468,f28426,f24644,f22355,f20177,f22354,f28421,f21950,f16791,f28445,f15989,f14353,f4811,f2034,f4831,f7686,f28486,f28427,f2394,f1395,f1394,f31759,f28338,f12417,f120867,f120884,f120937])).
fof(f120937,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f120936,f1849])).
fof(f1849,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',sK0),s('type.list.list'('type.exhSem.v'),X0))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.exhSem.v'),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f429,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : ((s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.exhSem.v'),X1) != s('type.list.list'('type.exhSem.v'),X3) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : ((s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.exhSem.v'),X1) != s('type.list.list'('type.exhSem.v'),X3)) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3)))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.exhSem.v'),X17))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.exhSem.v'),X19))) <=> (s('type.num.num',X16) = s('type.num.num',X18) & s('type.list.list'('type.exhSem.v'),X17) = s('type.list.list'('type.exhSem.v'),X19)))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_|49||49||split|1')).
fof(f120936,plain,(
  ( ! [X212] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',sK0),s('type.list.list'('type.exhSem.v'),X212))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.exhSem.v'),X212)))) ) | $spl0),
  inference(forward_demodulation,[],[f109832,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.exhSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.exhSem.v'),X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X20] : s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.exhSem.v'),X20))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X8))),s('type.list.list'('type.exhSem.v'),X20)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_2_const.exhSem.Conv')).
fof(f109832,plain,(
  ( ! [X212] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.exhSem.v'),X212))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',sK0))),s('type.list.list'('type.exhSem.v'),X212)))) ) | $spl0),
  inference(superposition,[],[f328,f104072])).
fof(f104072,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f94235,f327])).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X9)),X11),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X9))),X12),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.exhSem.v'),X9),X15)))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_case_def|split|4')).
fof(f94235,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X1),s(fun('type.list.list'('type.exhSem.v'),X0),X6)))) ) | $spl0),
  inference(superposition,[],[f327,f426])).
fof(f120884,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f120883,f429])).
fof(f120883,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f109743,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.exhSem.v_|49||49||split|4','_i0_ax')).
fof(f109743,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f309,f104072])).
fof(f120867,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f120866,f429])).
fof(f120866,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f109718,f309])).
fof(f109718,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f104072,f309])).
fof(f12417,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl30),
  inference(cnf_transformation,[],[f12417_D])).
fof(f12417_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f28338,plain,(
  ( ! [X50] : (s(bool,X50) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X50)) ) | $spl40),
  inference(cnf_transformation,[],[f28338_D])).
fof(f28338_D,plain,(
  ( ! [X50] : (s(bool,X50) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X50)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f31759,plain,(
  ( ! [X80,X79] : (s(bool,X80) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X80)) ) | ($spl10 | $spl12 | $spl26)),
  inference(duplicate_literal_removal,[],[f31587])).
fof(f31587,plain,(
  ( ! [X80,X79] : (s(bool,X80) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X79)) ) | ($spl10 | $spl12 | $spl26)),
  inference(superposition,[],[f24999,f28784])).
fof(f28784,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f28769])).
fof(f28769,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f28418,f445])).
fof(f445,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.exhSem.v_|49||49||split|4',aHL_BOOL_CASES)).
fof(f28418,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1497,f1374,f1376,f1378,f2036,f2199,f2180,f2297,f2515,f2810,f2865,f3011,f3027,f3249,f3250,f3251,f3252,f3254,f3255,f3256,f3257,f3463,f3604,f3675,f3768,f4178,f4364,f4365,f4516,f4517,f4682,f4605,f4611,f4617,f4686,f4813,f4840,f4916,f4915,f4909,f5038,f5122,f5125,f5121,f5380,f5391,f5393,f5395,f5396,f5400,f5787,f6722,f6866,f7280,f7281,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6980,f6981,f6992,f7142,f7147,f7159,f7163,f7165,f7235,f7240,f7252,f7256,f7258,f7420,f7422,f7424,f7425,f7426,f7427,f7428,f7685,f7694,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8060,f8061,f8062,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8096,f8097,f8108,f8109,f8110,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9611,f9660,f9736,f9737,f9826,f9827,f10018,f10064,f10067,f10069,f10071,f10072,f10073,f10074,f10075,f10543,f10544,f10545,f10547,f10548,f10549,f10550,f10551,f10754,f10756,f10758,f10759,f10760,f10761,f10762,f11056,f11347,f11348,f11153,f11159,f11184,f11678,f11695,f11698,f11699,f11701,f11703,f11704,f11705,f11706,f11873,f11875,f11446,f10793,f10464,f10463,f10462,f10461,f10460,f10458,f10456,f10453,f9837,f9670,f264,f18664,f18631,f18775,f19750,f19753,f19756,f19758,f19759,f19760,f19761,f19762,f19763,f20178,f20179,f20181,f9630,f9661,f9666,f11707,f5776,f4609,f8094,f8059,f4715,f4578,f4576,f4571,f4570,f4568,f4565,f4564,f1515,f4529,f6563,f6566,f2803,f4531,f11124,f11128,f7145,f7238,f4685,f4684,f4600,f4536,f4535,f4534,f4533,f4532,f4530,f4528,f4527,f930,f834,f833,f832,f831,f830,f829,f318,f20180,f18623,f10755,f9627,f21177,f7421,f2037,f1486,f1484,f9958,f9676,f9959,f9677,f9971,f9689,f9972,f9690,f9981,f9968,f9699,f9686,f5256,f5093,f398,f350,f349,f2017,f9585,f4153,f1533,f9582,f6588,f18620,f10752,f7418,f7098,f9587,f4155,f2841,f9524,f9525,f11353,f7141,f11070,f7143,f7234,f7236,f18624,f11113,f11112,f11073,f11072,f11042,f11041,f7146,f7239,f11340,f9534,f9591,f9590,f9589,f9588,f9586,f9583,f9584,f9604,f339,f11071,f7664,f4601,f4160,f4159,f4158,f4157,f4156,f4154,f4152,f4151,f420,f311,f7695,f7894,f9536,f9537,f9551,f9552,f7679,f7677,f11337,f7674,f7583,f7584,f7662,f4599,f1391,f1390,f1389,f1388,f1387,f1386,f1385,f319,f1392,f7099,f12005,f12026,f12023,f11994,f11995,f12027,f12029,f9581,f7193,f7190,f18621,f12036,f12033,f12032,f12031,f12030,f12028,f12025,f12024,f11991,f11990,f11989,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11981,f11980,f11979,f11978,f10753,f7663,f7419,f7209,f7207,f7202,f7201,f7199,f7196,f7195,f7192,f7191,f7115,f7113,f7108,f7107,f7105,f7102,f7101,f348,f345,f19691,f19690,f19688,f19687,f19686,f19685,f19684,f19683,f19682,f19681,f19680,f19679,f19678,f19659,f19649,f19648,f19645,f19644,f19643,f19642,f19641,f19640,f19639,f19638,f19637,f19636,f19635,f19634,f19633,f19632,f19300,f19299,f19297,f19296,f19295,f19294,f19293,f19292,f19291,f19290,f19289,f19288,f19287,f19268,f19258,f19257,f19254,f19253,f19252,f19251,f19250,f19249,f19248,f19247,f19246,f19245,f19244,f19243,f19242,f19241,f12035,f11338,f7885,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7867,f7838,f7829,f7828,f7688,f7683,f7682,f7681,f7680,f7678,f7676,f7675,f7593,f4833,f4754,f4681,f4598,f2032,f347,f346,f321,f320,f317,f18626,f10757,f7423,f2766,f2764,f4376,f4377,f20182,f21178,f21179,f21180,f19694,f19303,f18671,f12039,f11868,f11688,f11628,f11627,f11560,f11555,f11553,f11552,f11550,f11549,f11303,f11108,f11096,f11095,f11091,f11090,f11088,f11049,f11048,f11040,f10798,f10021,f9740,f9607,f9599,f9593,f9559,f9558,f9550,f9544,f9543,f9535,f8160,f8156,f8142,f8121,f8120,f8112,f8111,f8104,f8103,f8095,f8073,f8072,f8064,f8063,f7888,f7691,f7457,f7357,f7228,f7218,f7217,f7216,f7215,f7211,f7210,f7208,f7198,f7197,f7134,f7124,f7123,f7122,f7121,f7117,f7116,f7114,f7104,f7103,f6988,f6987,f6979,f6876,f6782,f6781,f6779,f6777,f6776,f6775,f6774,f6732,f6639,f6638,f6636,f6634,f6633,f6632,f6631,f5404,f5392,f5390,f5383,f5259,f5096,f4961,f4836,f4680,f4595,f4585,f4584,f4580,f4579,f4577,f4567,f4566,f4524,f4523,f4515,f4509,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4391,f4388,f4384,f4383,f4375,f4372,f4371,f4363,f4360,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f4219,f4218,f4217,f4216,f4215,f4214,f4213,f4212,f4211,f4165,f4116,f4109,f3674,f3499,f3498,f3497,f3496,f3495,f3494,f3493,f3492,f3428,f3427,f3426,f3425,f3424,f3423,f3422,f3421,f3420,f3013,f2979,f2977,f2975,f2792,f2522,f2521,f2520,f2483,f2482,f2480,f2478,f2477,f2476,f2475,f2474,f2473,f2427,f1488,f909,f743,f686,f633,f540,f513,f472,f276,f275,f5379,f5378,f5377,f5376,f5375,f5374,f5373,f11109,f5372,f5371,f5370,f7229,f7135,f4812,f4854,f4859,f2804,f326,f7893,f7884,f7880,f7848,f7847,f7841,f7840,f7839,f358,f11696,f11700,f11697,f11702,f10546,f11177,f10542,f10541,f3253,f11167,f7265,f7174,f7282,f7267,f7172,f6970,f6959,f5222,f5049,f7149,f7152,f7242,f7245,f356,f7151,f7150,f7244,f7243,f5212,f5059,f355,f354,f21521,f22126,f22129,f22132,f22134,f22135,f22136,f22137,f22138,f28417])).
fof(f28417,plain,(
  ( ! [X87] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X87)) ) | $spl10),
  inference(subsumption_resolution,[],[f26637,f11678])).
fof(f26637,plain,(
  ( ! [X87] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X87)) ) | $spl10),
  inference(superposition,[],[f11678,f2975])).
fof(f22138,plain,(
  ( ! [X127,X128,X129] : (s(bool,X127) = s(bool,sK7('const.pair.FST_0',X129,bool)) | s(bool,'T_0') = s(bool,sK4(X127,X128,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X128)))) ) | $spl26),
  inference(superposition,[],[f21177,f836])).
fof(f22137,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X125) = s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X123,bool),'const.pair.,_2'(s(X123,X124),s(bool,X125)))))) ) | $spl26),
  inference(superposition,[],[f21177,f342])).
fof(f22136,plain,(
  ( ! [X121,X122,X120] : (s(bool,X121) = s(bool,sK7('const.pair.FST_0',X122,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X120),'const.sum.INL_1'(s(bool,X121)))))) ) | $spl26),
  inference(superposition,[],[f21177,f336])).
fof(f22135,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,X117) = s(bool,sK7('const.pair.FST_0',X119,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X116),'const.pair.,_2'(s(bool,X117),s(X116,X118)))))) ) | $spl26),
  inference(superposition,[],[f21177,f341])).
fof(f22134,plain,(
  ( ! [X114,X115,X113] : (s(bool,X114) = s(bool,sK7('const.pair.FST_0',X115,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X114)))))) ) | $spl26),
  inference(superposition,[],[f21177,f337])).
fof(f22132,plain,(
  ( ! [X107,X108,X106] : (s(bool,X106) = s(bool,sK7('const.pair.FST_0',X108,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X106),s(fun('type.num.num',bool),X107),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f21177,f316])).
fof(f22129,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X95) = s(bool,sK7('const.pair.FST_0',X96,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X93),bool),sK18(X94,X95,X93,bool)),s('type.list.list'(X93),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f21177,f357])).
fof(f22126,plain,(
  ( ! [X80,X79] : (s(bool,X79) = s(bool,sK7('const.pair.FST_0',X80,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X79)))) ) | $spl26),
  inference(superposition,[],[f21177,f309])).
fof(f21521,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f21466,f21180])).
fof(f21466,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f271,f21178])).
fof(f354,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0)))))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X31,X9] : ! [X10,X37] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) <=> (! [X38] : s(X31,X38) = s(X31,happ(s(fun(X9,X31),X10),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))) & ! [X43] : s(X9,X43) = s(X9,happ(s(fun(X31,X9),X37),s(X31,happ(s(fun(X9,X31),X10),s(X9,X43)))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.ISO')).
fof(f355,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f5059,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f836])).
fof(f5212,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f7243,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X403,X402),X409),s(fun(X402,X403),happ(s(fun(X405,fun(X402,X403)),X406),s(X405,happ(s(fun(X404,X405),X407),s(X404,X408)))))))) | s(X402,X410) = s(X402,happ(s(fun(X403,X402),X409),s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),'_dst_x0x1_2'(s(fun(X405,fun(X402,X403)),X406),s(fun(X404,X405),X407))),s(X404,X408))),s(X402,X410)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7244,plain,(
  ( ! [X412,X414,X416,X418,X413,X415,X411,X417,X419] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X411,X412),happ(s(fun(X414,fun(X411,X412)),X415),s(X414,happ(s(fun(X413,X414),X416),s(X413,X417))))),s(fun(X412,X411),X418)))) | s(X411,X419) = s(X411,happ(s(fun(X412,X411),X418),s(X412,happ(s(fun(X411,X412),happ(s(fun(X413,fun(X411,X412)),'_dst_x0x1_2'(s(fun(X414,fun(X411,X412)),X415),s(fun(X413,X414),X416))),s(X413,X417))),s(X411,X419)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f7150,plain,(
  ( ! [X420,X422,X416,X418,X415,X421,X423,X417,X419] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X416,X415),X422),s(fun(X415,X416),happ(s(fun(X419,fun(X415,X416)),'_dst_x0x1_2'(s(fun(X417,fun(X415,X416)),X418),s(fun(X419,X417),X420))),s(X419,X421)))))) | s(X415,X423) = s(X415,happ(s(fun(X416,X415),X422),s(X416,happ(s(fun(X415,X416),happ(s(fun(X417,fun(X415,X416)),X418),s(X417,happ(s(fun(X419,X417),X420),s(X419,X421))))),s(X415,X423)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7151,plain,(
  ( ! [X428,X430,X424,X426,X432,X429,X431,X425,X427] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X424,X425),happ(s(fun(X428,fun(X424,X425)),'_dst_x0x1_2'(s(fun(X426,fun(X424,X425)),X427),s(fun(X428,X426),X429))),s(X428,X430))),s(fun(X425,X424),X431)))) | s(X424,X432) = s(X424,happ(s(fun(X425,X424),X431),s(X425,happ(s(fun(X424,X425),happ(s(fun(X426,fun(X424,X425)),X427),s(X426,happ(s(fun(X428,X426),X429),s(X428,X430))))),s(X424,X432)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f7245,plain,(
  ( ! [X420,X422,X428,X424,X426,X421,X423,X425,X427] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X421,X420),X427),s(fun(X420,X421),happ(s(fun(X423,fun(X420,X421)),X424),s(X423,happ(s(fun(X422,X423),X425),s(X422,X426)))))))) | s(X421,X428) = s(X421,happ(s(fun(X420,X421),happ(s(fun(X422,fun(X420,X421)),'_dst_x0x1_2'(s(fun(X423,fun(X420,X421)),X424),s(fun(X422,X423),X425))),s(X422,X426))),s(X420,happ(s(fun(X421,X420),X427),s(X421,X428)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f7242,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X393,X394),happ(s(fun(X396,fun(X393,X394)),X397),s(X396,happ(s(fun(X395,X396),X398),s(X395,X399))))),s(fun(X394,X393),X400)))) | s(X394,X401) = s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X396,fun(X393,X394)),X397),s(fun(X395,X396),X398))),s(X395,X399))),s(X393,happ(s(fun(X394,X393),X400),s(X394,X401)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7152,plain,(
  ( ! [X441,X436,X438,X434,X440,X437,X439,X433,X435] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X434,X433),X440),s(fun(X433,X434),happ(s(fun(X437,fun(X433,X434)),'_dst_x0x1_2'(s(fun(X435,fun(X433,X434)),X436),s(fun(X437,X435),X438))),s(X437,X439)))))) | s(X434,X441) = s(X434,happ(s(fun(X433,X434),happ(s(fun(X435,fun(X433,X434)),X436),s(X435,happ(s(fun(X437,X435),X438),s(X437,X439))))),s(X433,happ(s(fun(X434,X433),X440),s(X434,X441)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f7149,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X409,X411,X406] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X406,X407),happ(s(fun(X410,fun(X406,X407)),'_dst_x0x1_2'(s(fun(X408,fun(X406,X407)),X409),s(fun(X410,X408),X411))),s(X410,X412))),s(fun(X407,X406),X413)))) | s(X407,X414) = s(X407,happ(s(fun(X406,X407),happ(s(fun(X408,fun(X406,X407)),X409),s(X408,happ(s(fun(X410,X408),X411),s(X410,X412))))),s(X406,happ(s(fun(X407,X406),X413),s(X407,X414)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f5049,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f836])).
fof(f5222,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f6959,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X0,X1),X3))),s(X0,X4))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X5,'type.list.list'('type.exhSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X7,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.exhSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f6970,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X5,'type.list.list'('type.exhSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X7,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.exhSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f7172,plain,(
  ( ! [X592,X588,X590,X586,X589,X591,X585,X587] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X591),s('type.list.list'('type.exhSem.v'),X592))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X590),s('type.list.list'('type.exhSem.v'),happ(s(fun(X587,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X585,'type.list.list'('type.exhSem.v')),X586),s(fun(X587,X585),X588))),s(X587,X589))))) | s('type.list.list'('type.exhSem.v'),X592) = s('type.list.list'('type.exhSem.v'),happ(s(fun(X585,'type.list.list'('type.exhSem.v')),X586),s(X585,happ(s(fun(X587,X585),X588),s(X587,X589)))))) )),
  inference(superposition,[],[f395,f369])).
fof(f7267,plain,(
  ( ! [X575,X580,X582,X576,X578,X581,X577,X579] : (s('type.list.list'('type.exhSem.v'),X580) != s('type.list.list'('type.exhSem.v'),happ(s(fun(X576,'type.list.list'('type.exhSem.v')),X577),s(X576,happ(s(fun(X575,X576),X578),s(X575,X579))))) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X582),s('type.list.list'('type.exhSem.v'),X580))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X581),s('type.list.list'('type.exhSem.v'),happ(s(fun(X575,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X576,'type.list.list'('type.exhSem.v')),X577),s(fun(X575,X576),X578))),s(X575,X579))))) | s('type.num.num',X581) != s('type.num.num',X582)) )),
  inference(superposition,[],[f396,f369])).
fof(f7282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),happ(s(fun(X1,'type.list.list'('type.exhSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),happ(s(fun(X3,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.exhSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(forward_demodulation,[],[f6967,f342])).
fof(f6967,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),happ(s(fun(X1,'type.list.list'('type.exhSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X6,'type.num.num'),'const.pair.,_2'(s(X6,X7),s('type.num.num',X0))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X3,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.exhSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f342,f369,f396])).
fof(f7174,plain,(
  ( ! [X606,X602,X605,X607,X601,X603,X608,X604] : (s('type.list.list'('type.exhSem.v'),X606) != s('type.list.list'('type.exhSem.v'),happ(s(fun(X603,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X601,'type.list.list'('type.exhSem.v')),X602),s(fun(X603,X601),X604))),s(X603,X605))) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X608),s('type.list.list'('type.exhSem.v'),X606))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X607),s('type.list.list'('type.exhSem.v'),happ(s(fun(X601,'type.list.list'('type.exhSem.v')),X602),s(X601,happ(s(fun(X603,X601),X604),s(X603,X605))))))) | s('type.num.num',X607) != s('type.num.num',X608)) )),
  inference(superposition,[],[f396,f369])).
fof(f7265,plain,(
  ( ! [X565,X561,X563,X559,X564,X566,X560,X562] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X565),s('type.list.list'('type.exhSem.v'),X566))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X564),s('type.list.list'('type.exhSem.v'),happ(s(fun(X560,'type.list.list'('type.exhSem.v')),X561),s(X560,happ(s(fun(X559,X560),X562),s(X559,X563))))))) | s('type.list.list'('type.exhSem.v'),X566) = s('type.list.list'('type.exhSem.v'),happ(s(fun(X559,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X560,'type.list.list'('type.exhSem.v')),X561),s(fun(X559,X560),X562))),s(X559,X563)))) )),
  inference(superposition,[],[f395,f369])).
fof(f11167,plain,(
  ( ! [X753,X749,X751,X747,X752,X754,X748,X750] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X753),s('type.list.list'('type.exhSem.v'),X754))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X752),s('type.list.list'('type.exhSem.v'),happ(s(fun(X748,'type.list.list'('type.exhSem.v')),happ(s(fun(X747,fun(X748,'type.list.list'('type.exhSem.v'))),X749),s(X747,X750))),s(X748,X751))))) | s('type.list.list'('type.exhSem.v'),X754) = s('type.list.list'('type.exhSem.v'),happ(s(fun('type.pair.prod'(X747,X748),'type.list.list'('type.exhSem.v')),sK19(X749,X748,X747,'type.list.list'('type.exhSem.v'))),s('type.pair.prod'(X747,X748),'const.pair.,_2'(s(X747,X750),s(X748,X751)))))) )),
  inference(superposition,[],[f395,f367])).
fof(f3253,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X24),s('type.list.list'('type.exhSem.v'),X25))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X23),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X21),s('type.num.num',X22))))) | s('type.list.list'('type.exhSem.v'),X25) = s('type.list.list'('type.exhSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.exhSem.v'),X20),s(fun('type.num.num','type.list.list'('type.exhSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f395,f324])).
fof(f10541,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.exhLang.exp',X15))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11),s('type.list.list'('type.string.char'),happ(s(fun(X8,'type.list.list'('type.string.char')),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.string.char')),X7),s(fun(X8,X6),X9))),s(X8,X10))),s('type.exhLang.exp',X12))) | s('type.list.list'('type.string.char'),X14) = s('type.list.list'('type.string.char'),happ(s(fun(X6,'type.list.list'('type.string.char')),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f400,f369])).
fof(f10542,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.exhLang.exp',X25))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X21),s('type.list.list'('type.string.char'),happ(s(fun(X17,'type.list.list'('type.string.char')),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.exhLang.exp',X22))) | s('type.list.list'('type.string.char'),X24) = s('type.list.list'('type.string.char'),happ(s(fun(X16,'type.list.list'('type.string.char')),'_dst_x0x1_2'(s(fun(X17,'type.list.list'('type.string.char')),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f400,f369])).
fof(f11177,plain,(
  ( ! [X831,X835,X833,X839,X837,X840,X834,X832,X838,X836] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X836),s('type.list.list'('type.string.char'),X837),s('type.exhLang.exp',X838))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X839),s('type.list.list'('type.string.char'),happ(s(fun(X832,'type.list.list'('type.string.char')),happ(s(fun(X831,fun(X832,'type.list.list'('type.string.char'))),X833),s(X831,X834))),s(X832,X835))),s('type.exhLang.exp',X840))) | s('type.list.list'('type.string.char'),X837) = s('type.list.list'('type.string.char'),happ(s(fun('type.pair.prod'(X831,X832),'type.list.list'('type.string.char')),sK19(X833,X832,X831,'type.list.list'('type.string.char'))),s('type.pair.prod'(X831,X832),'const.pair.,_2'(s(X831,X834),s(X832,X835)))))) )),
  inference(superposition,[],[f400,f367])).
fof(f10546,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X53),s('type.list.list'('type.string.char'),X54),s('type.exhLang.exp',X55))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X51),s('type.list.list'('type.string.char'),happ(s(fun('type.num.num','type.list.list'('type.string.char')),X49),s('type.num.num',X50))),s('type.exhLang.exp',X52))) | s('type.list.list'('type.string.char'),X54) = s('type.list.list'('type.string.char'),'const.ind_type.FCONS_3'(s('type.list.list'('type.string.char'),X48),s(fun('type.num.num','type.list.list'('type.string.char')),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f400,f324])).
fof(f11702,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X63),s('type.list.list'('type.string.char'),X64),s('type.exhLang.exp',X65))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X59),s('type.num.num',X60))),s('type.list.list'('type.string.char'),X61),s('type.exhLang.exp',X62))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X63) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.ind_type.FCONS_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X58),s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60)))))) )),
  inference(superposition,[],[f399,f324])).
fof(f11697,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.exhLang.exp',X25))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun(X17,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.list.list'('type.string.char'),X21),s('type.exhLang.exp',X22))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun(X16,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'_dst_x0x1_2'(s(fun(X17,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f399,f369])).
fof(f11700,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X49] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X48),s('type.list.list'('type.string.char'),X49),s('type.exhLang.exp',X50))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun(X42,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X41,fun(X42,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X43),s(X41,X44))),s(X42,X45))),s('type.list.list'('type.string.char'),X46),s('type.exhLang.exp',X47))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X48) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.pair.prod'(X41,X42),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK19(X43,X42,X41,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) )),
  inference(superposition,[],[f399,f367])).
fof(f11696,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.exhLang.exp',X15))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun(X8,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7),s(fun(X8,X6),X9))),s(X8,X10))),s('type.list.list'('type.string.char'),X11),s('type.exhLang.exp',X12))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun(X6,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f399,f369])).
fof(f358,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X31,X9] : ! [X32,X11] : ? [X33] : (s(X31,X32) = s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X31,happ(s(fun(X31,X31),happ(s(fun('type.list.list'(X9),fun(X31,X31)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X31,X31))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),X17))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.list.list_Axiom')).
fof(f7839,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f321,f270])).
fof(f7840,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,sK6(X3,X4,bool))))) | s(bool,X6) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f321,f445])).
fof(f7841,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X7,X8,bool))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool))) )),
  inference(superposition,[],[f321,f445])).
fof(f7847,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X30),s(bool,sK6(X29,X30,bool))))) | s(bool,X32) = s(bool,sK5(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f321,f2865])).
fof(f7848,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,sK6(X33,X34,bool))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK5(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2865])).
fof(f7880,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f7845,f2199])).
fof(f7845,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2522])).
fof(f7884,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6(X0,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(forward_demodulation,[],[f7883,f309])).
fof(f7883,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f7852,f262])).
fof(f7852,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f321,f472])).
fof(f7893,plain,(
  ( ! [X64,X65] : (p(s(bool,sK6(X64,'_i_0_0',bool))) | ~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))))) )),
  inference(forward_demodulation,[],[f7866,f309])).
fof(f7866,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X64,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f326,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X8,X20,X23] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X20),s('type.num.num',X23))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X20))),s('type.num.num',X23)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f2804,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X151),s('type.list.list'('type.exhSem.v'),X152))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X149),s('type.num.num',X150))),s('type.list.list'('type.exhSem.v'),X153))) | s('type.num.num',X151) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X148),s(fun('type.num.num','type.num.num'),X149),s('type.num.num','const.num.SUC_1'(s('type.num.num',X150)))))) )),
  inference(superposition,[],[f394,f324])).
fof(f4859,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4850,f264])).
fof(f4850,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4812,f2865])).
fof(f4854,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4842,f264])).
fof(f4842,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4812,f270])).
fof(f4812,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK18(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4681,f357])).
fof(f7135,plain,(
  ( ! [X308,X304,X306,X303,X309,X305,X307] : (~p(s(bool,happ(s(fun(X305,bool),'_dst_x0x1_2'(s(fun(X303,bool),X304),s(fun(X305,X303),X306))),s(X305,X307)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X308),bool),sK18(X309,happ(s(fun(X303,bool),X304),s(X303,happ(s(fun(X305,X303),X306),s(X305,X307)))),X308,bool))))) )),
  inference(superposition,[],[f4854,f369])).
fof(f7229,plain,(
  ( ! [X300,X302,X304,X306,X301,X303,X305] : (~p(s(bool,happ(s(fun(X301,bool),X302),s(X301,happ(s(fun(X300,X301),X303),s(X300,X304)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X305),bool),sK18(X306,happ(s(fun(X300,bool),'_dst_x0x1_2'(s(fun(X301,bool),X302),s(fun(X300,X301),X303))),s(X300,X304)),X305,bool))))) )),
  inference(superposition,[],[f4854,f369])).
fof(f5370,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4854,f309])).
fof(f5371,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK18(X114,happ(s(fun('type.list.list'(X110),bool),sK18(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4854,f357])).
fof(f5372,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK18(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4854,f338])).
fof(f11109,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(X328,bool),happ(s(fun(X327,fun(X328,bool)),X329),s(X327,X330))),s(X328,X331)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X332),bool),sK18(X333,happ(s(fun('type.pair.prod'(X327,X328),bool),sK19(X329,X328,X327,bool)),s('type.pair.prod'(X327,X328),'const.pair.,_2'(s(X327,X330),s(X328,X331)))),X332,bool))))) )),
  inference(superposition,[],[f4854,f367])).
fof(f5373,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK18(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) )),
  inference(superposition,[],[f4854,f316])).
fof(f5374,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun('type.num.num',bool),X124),s('type.num.num',X125)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X126),bool),sK18(X127,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.SUC_1'(s('type.num.num',X125)))),X126,bool))))) )),
  inference(superposition,[],[f4854,f324])).
fof(f5375,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,X129)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK18(X131,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X129)))),X130,bool))))) )),
  inference(superposition,[],[f4854,f337])).
fof(f5376,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4854,f341])).
fof(f5377,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,X138)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X139),bool),sK18(X140,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X138)))),X139,bool))))) )),
  inference(superposition,[],[f4854,f336])).
fof(f5378,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4854,f342])).
fof(f5379,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK18(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4854,f836])).
fof(f275,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|3')).
fof(f276,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f472,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f309,f270])).
fof(f513,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f336,f270])).
fof(f540,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f337,f270])).
fof(f633,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f316,f270])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f341,f270])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f342,f270])).
fof(f909,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,X2,X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f357,f270])).
fof(f1488,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X12),X13)))) )),
  inference(superposition,[],[f338,f270])).
fof(f2427,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f472,f338])).
fof(f2473,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2426,f309])).
fof(f2426,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f472,f357])).
fof(f2474,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2428,f309])).
fof(f2428,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f472,f316])).
fof(f2475,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f2429,f309])).
fof(f2429,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30)))))) )),
  inference(superposition,[],[f472,f337])).
fof(f2476,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2430,f309])).
fof(f2430,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f472,f341])).
fof(f2477,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(forward_demodulation,[],[f2431,f309])).
fof(f2431,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X35))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35)))))) )),
  inference(superposition,[],[f472,f336])).
fof(f2478,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(forward_demodulation,[],[f2432,f309])).
fof(f2432,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f472,f342])).
fof(f2480,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2479,f309])).
fof(f2479,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2446,f262])).
fof(f2446,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f333,f472])).
fof(f2482,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2481,f309])).
fof(f2481,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2447,f262])).
fof(f2447,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f418,f472])).
fof(f2483,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2450,f262])).
fof(f2450,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f447,f472])).
fof(f2520,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2498,f262])).
fof(f2498,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2480,f472])).
fof(f2521,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2499,f262])).
fof(f2499,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2480,f440])).
fof(f2522,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2500,f262])).
fof(f2500,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f2480,f2199])).
fof(f2792,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f472,f324])).
fof(f2975,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2856,f2199])).
fof(f2856,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2522,f309])).
fof(f2977,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2858,f2199])).
fof(f2858,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2522,f338])).
fof(f2979,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2860,f2199])).
fof(f2860,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(superposition,[],[f2522,f324])).
fof(f3013,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3012,f262])).
fof(f3012,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2897,f2199])).
fof(f2897,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f317,f2522])).
fof(f3420,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2865,f309])).
fof(f3421,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK18(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2865,f357])).
fof(f3422,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2865,f338])).
fof(f3423,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2865,f316])).
fof(f3424,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl10),
  inference(superposition,[],[f2865,f324])).
fof(f3425,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X76,bool),'const.sum.INR_1'(s(bool,X77)))))) ) | $spl10),
  inference(superposition,[],[f2865,f337])).
fof(f3426,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2865,f341])).
fof(f3427,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X83),'const.sum.INL_1'(s(bool,X84)))))) ) | $spl10),
  inference(superposition,[],[f2865,f336])).
fof(f3428,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2865,f342])).
fof(f3492,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK18(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2865,f357])).
fof(f3493,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2865,f338])).
fof(f3494,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2865,f316])).
fof(f3495,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61))) | s(bool,X62) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl10),
  inference(superposition,[],[f2865,f324])).
fof(f3496,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X63,bool),'const.sum.INR_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2865,f337])).
fof(f3497,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2865,f341])).
fof(f3498,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') = s(bool,X71) | s(bool,X72) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X70),'const.sum.INL_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f2865,f336])).
fof(f3499,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2865,f342])).
fof(f3674,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f275,f2865])).
fof(f4109,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3710,f264])).
fof(f3710,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f330,f2865])).
fof(f4116,plain,(
  ( ! [X97,X95,X98,X96] : (s(X95,X97) = s(X95,sK4(X97,X96,X95)) | s(bool,X98) = s(bool,'|dollar|exists_1'(s(fun(X95,bool),X96))) | s(bool,'F_0') = s(bool,X98)) ) | $spl10),
  inference(subsumption_resolution,[],[f3745,f264])).
fof(f3745,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,'F_0')) | s(X95,X97) = s(X95,sK4(X97,X96,X95)) | s(bool,X98) = s(bool,'|dollar|exists_1'(s(fun(X95,bool),X96))) | s(bool,'F_0') = s(bool,X98)) ) | $spl10),
  inference(superposition,[],[f318,f2865])).
fof(f4165,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),'const.list.NIL_0')))) | s(bool,X23) = s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),sK2(X22,X21)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f4147,f264])).
fof(f4147,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),'const.list.NIL_0')))) | s(bool,X23) = s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),sK2(X22,X21)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f420,f2865])).
fof(f4211,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3675,f309])).
fof(f4212,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK18(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3675,f357])).
fof(f4213,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3675,f338])).
fof(f4214,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3675,f316])).
fof(f4215,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X75),s('type.num.num',X76))) | s(bool,X77) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl10),
  inference(superposition,[],[f3675,f324])).
fof(f4216,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X80) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3675,f337])).
fof(f4217,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3675,f341])).
fof(f4218,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,X87)) ) | $spl10),
  inference(superposition,[],[f3675,f336])).
fof(f4219,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3675,f342])).
fof(f4274,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK18(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f513,f357])).
fof(f4275,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f513,f338])).
fof(f4276,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f513,f316])).
fof(f4277,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f513,f324])).
fof(f4278,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,X80)))))) )),
  inference(superposition,[],[f513,f337])).
fof(f4279,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f513,f341])).
fof(f4280,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X86),'const.sum.INL_1'(s(bool,X87)))))) )),
  inference(superposition,[],[f513,f336])).
fof(f4281,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f513,f342])).
fof(f4360,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4304,f262])).
fof(f4304,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f513])).
fof(f4363,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,X2,bool))) )),
  inference(superposition,[],[f351,f270])).
fof(f4371,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK14(X29,X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f351,f2865])).
fof(f4372,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,f(s(bool,X36))) != s(X35,g(s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK14(X33,X34,X35,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2865])).
fof(f4375,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK14(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK14(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f351,f270])).
fof(f4383,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') != s(bool,f(s(X29,sK14(X30,X31,bool,X29)))) | s(fun(X29,bool),X30) = s(fun(X29,bool),X31) | s(bool,X32) = s(bool,g(s(X29,sK14(X30,X31,bool,X29)))) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f351,f2865])).
fof(f4384,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,f(s(X33,sK14(X34,X35,bool,X33)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(X33,sK14(X34,X35,bool,X33))))) ) | $spl10),
  inference(superposition,[],[f351,f2865])).
fof(f4388,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4369,f2199])).
fof(f4369,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2522])).
fof(f4391,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f4381,f2199])).
fof(f4381,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f351,f2522])).
fof(f4422,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK18(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f540,f357])).
fof(f4423,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f540,f338])).
fof(f4424,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f540,f316])).
fof(f4425,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f540,f324])).
fof(f4426,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83)))))) )),
  inference(superposition,[],[f540,f337])).
fof(f4427,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X88,bool),'const.sum.INR_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f540,f341])).
fof(f4428,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90)))))) )),
  inference(superposition,[],[f540,f336])).
fof(f4429,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94)))))) )),
  inference(superposition,[],[f540,f342])).
fof(f4509,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4452,f262])).
fof(f4452,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f540])).
fof(f4515,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f836,f270])).
fof(f4523,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f836,f2865])).
fof(f4524,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2865])).
fof(f4566,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f275,f836])).
fof(f4567,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f276,f836])).
fof(f4577,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f472,f836])).
fof(f4579,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f513,f836])).
fof(f4580,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f540,f836])).
fof(f4584,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | s(bool,X110) = s(bool,sK4(X108,X109,bool)) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X109)))) )),
  inference(superposition,[],[f2480,f836])).
fof(f4585,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | s(bool,X113) != s(bool,sK4(X111,X112,bool)) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) )),
  inference(superposition,[],[f2482,f836])).
fof(f4595,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3675,f836])).
fof(f4680,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4521,f2199])).
fof(f4521,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2522])).
fof(f4836,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4835,f262])).
fof(f4835,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f4805,f2199])).
fof(f4805,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f4681,f2522])).
fof(f4961,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4926,f264])).
fof(f4926,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4840,f2865])).
fof(f5096,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5095,f262])).
fof(f5095,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f5067,f2199])).
fof(f5067,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f354,f2522])).
fof(f5259,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5258,f262])).
fof(f5258,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f5230,f2199])).
fof(f5230,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f355,f2522])).
fof(f5383,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f5382,f262])).
fof(f5382,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f5346,f2199])).
fof(f5346,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f4854,f2522])).
fof(f5390,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5357,f262])).
fof(f5357,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4854,f472])).
fof(f5392,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f5359,f262])).
fof(f5359,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f4854,f540])).
fof(f5404,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5367,f262])).
fof(f5367,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4854,f513])).
fof(f6631,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X86),bool),sK18(X87,X88,X86,bool)),s('type.list.list'(X86),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2475,f357])).
fof(f6632,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f2475,f338])).
fof(f6633,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X93))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2475,f316])).
fof(f6634,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X97),s('type.num.num',X98))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.SUC_1'(s('type.num.num',X98)))))) )),
  inference(superposition,[],[f2475,f324])).
fof(f6636,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X106,bool),'const.sum.INR_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f2475,f341])).
fof(f6638,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f2475,f342])).
fof(f6639,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f2475,f836])).
fof(f6732,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f6674,f264])).
fof(f6674,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f2482,f2475])).
fof(f6774,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK18(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2477,f357])).
fof(f6775,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f2477,f338])).
fof(f6776,plain,(
  ( ! [X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2477,f316])).
fof(f6777,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X102),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f2477,f324])).
fof(f6779,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f2477,f341])).
fof(f6781,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X115)))))) )),
  inference(superposition,[],[f2477,f342])).
fof(f6782,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f2477,f836])).
fof(f6876,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6818,f264])).
fof(f6818,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'F_0')) | s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2482,f2477])).
fof(f6979,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X28,X26),X29))),s(X28,X30))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30)))))) )),
  inference(superposition,[],[f369,f270])).
fof(f6987,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X71,bool),X72),s(fun(X73,X71),X74))),s(X73,X75))) | s(bool,X76) = s(bool,happ(s(fun(X71,bool),X72),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f369,f2865])).
fof(f6988,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,X82) = s(bool,happ(s(fun(X79,bool),'_dst_x0x1_2'(s(fun(X77,bool),X78),s(fun(X79,X77),X80))),s(X79,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,happ(s(fun(X77,bool),X78),s(X77,happ(s(fun(X79,X77),X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f369,f2865])).
fof(f7103,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'F_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) )),
  inference(superposition,[],[f275,f369])).
fof(f7104,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X123,bool),X124),s(fun(X125,X123),X126))),s(X125,X127)))) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,happ(s(fun(X125,X123),X126),s(X125,X127)))))) )),
  inference(superposition,[],[f276,f369])).
fof(f7114,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X184,X182),X185))),s(X184,X186))))) | s(bool,'F_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X184,X182),X185),s(X184,X186)))))) )),
  inference(superposition,[],[f472,f369])).
fof(f7116,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X198),'const.sum.INL_1'(s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X193,bool),X194),s(fun(X195,X193),X196))),s(X195,X197))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,happ(s(fun(X195,X193),X196),s(X195,X197)))))) )),
  inference(superposition,[],[f513,f369])).
fof(f7117,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),X200),s(fun(X201,X199),X202))),s(X201,X203))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X199,bool),X200),s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))))) )),
  inference(superposition,[],[f540,f369])).
fof(f7121,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X226,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X223,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X223,X221),X224))),s(X223,X225))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X223,X221),X224),s(X223,X225)))))) )),
  inference(superposition,[],[f2475,f369])).
fof(f7122,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X232),'const.sum.INL_1'(s(bool,happ(s(fun(X229,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X229,X227),X230))),s(X229,X231))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X229,X227),X230),s(X229,X231)))))) )),
  inference(superposition,[],[f2477,f369])).
fof(f7123,plain,(
  ( ! [X237,X235,X233,X238,X236,X234] : (~p(s(bool,happ(s(fun(X235,bool),'_dst_x0x1_2'(s(fun(X233,bool),X234),s(fun(X235,X233),X236))),s(X235,X237)))) | s(bool,X238) = s(bool,happ(s(fun(X233,bool),X234),s(X233,happ(s(fun(X235,X233),X236),s(X235,X237))))) | s(bool,'F_0') = s(bool,X238)) )),
  inference(superposition,[],[f2480,f369])).
fof(f7124,plain,(
  ( ! [X243,X241,X239,X244,X242,X240] : (p(s(bool,happ(s(fun(X241,bool),'_dst_x0x1_2'(s(fun(X239,bool),X240),s(fun(X241,X239),X242))),s(X241,X243)))) | s(bool,X244) != s(bool,happ(s(fun(X239,bool),X240),s(X239,happ(s(fun(X241,X239),X242),s(X241,X243))))) | s(bool,'F_0') = s(bool,X244)) )),
  inference(superposition,[],[f2482,f369])).
fof(f7134,plain,(
  ( ! [X300,X302,X298,X301,X297,X299] : (s(bool,'T_0') != s(bool,happ(s(fun(X299,bool),'_dst_x0x1_2'(s(fun(X297,bool),X298),s(fun(X299,X297),X300))),s(X299,X301))) | s(bool,X302) = s(bool,happ(s(fun(X297,bool),X298),s(X297,happ(s(fun(X299,X297),X300),s(X299,X301))))) | s(bool,'F_0') = s(bool,X302)) ) | $spl10),
  inference(superposition,[],[f3675,f369])).
fof(f7197,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) )),
  inference(superposition,[],[f275,f369])).
fof(f7198,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f276,f369])).
fof(f7208,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183)))) )),
  inference(superposition,[],[f472,f369])).
fof(f7210,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X195),'const.sum.INL_1'(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f513,f369])).
fof(f7211,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X197,bool),X198),s(X197,happ(s(fun(X196,X197),X199),s(X196,X200))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X197,bool),X198),s(fun(X196,X197),X199))),s(X196,X200)))) )),
  inference(superposition,[],[f540,f369])).
fof(f7215,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X223,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X218,X219),X221),s(X218,X222))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X218,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X218,X219),X221))),s(X218,X222)))) )),
  inference(superposition,[],[f2475,f369])).
fof(f7216,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X229),'const.sum.INL_1'(s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X224,X225),X227),s(X224,X228))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X224,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X224,X225),X227))),s(X224,X228)))) )),
  inference(superposition,[],[f2477,f369])).
fof(f7217,plain,(
  ( ! [X231,X235,X233,X230,X234,X232] : (~p(s(bool,happ(s(fun(X231,bool),X232),s(X231,happ(s(fun(X230,X231),X233),s(X230,X234)))))) | s(bool,X235) = s(bool,happ(s(fun(X230,bool),'_dst_x0x1_2'(s(fun(X231,bool),X232),s(fun(X230,X231),X233))),s(X230,X234))) | s(bool,'F_0') = s(bool,X235)) )),
  inference(superposition,[],[f2480,f369])).
fof(f7218,plain,(
  ( ! [X241,X239,X237,X240,X238,X236] : (p(s(bool,happ(s(fun(X237,bool),X238),s(X237,happ(s(fun(X236,X237),X239),s(X236,X240)))))) | s(bool,X241) != s(bool,happ(s(fun(X236,bool),'_dst_x0x1_2'(s(fun(X237,bool),X238),s(fun(X236,X237),X239))),s(X236,X240))) | s(bool,'F_0') = s(bool,X241)) )),
  inference(superposition,[],[f2482,f369])).
fof(f7228,plain,(
  ( ! [X294,X296,X298,X295,X297,X299] : (s(bool,'T_0') != s(bool,happ(s(fun(X295,bool),X296),s(X295,happ(s(fun(X294,X295),X297),s(X294,X298))))) | s(bool,X299) = s(bool,happ(s(fun(X294,bool),'_dst_x0x1_2'(s(fun(X295,bool),X296),s(fun(X294,X295),X297))),s(X294,X298))) | s(bool,'F_0') = s(bool,X299)) ) | $spl10),
  inference(superposition,[],[f3675,f369])).
fof(f7357,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(forward_demodulation,[],[f6985,f2199])).
fof(f6985,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(superposition,[],[f369,f2522])).
fof(f7457,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f7398,f432])).
fof(f7398,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f6722,f2865])).
fof(f7691,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(subsumption_resolution,[],[f7690,f262])).
fof(f7690,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(forward_demodulation,[],[f7654,f2199])).
fof(f7654,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(superposition,[],[f320,f2522])).
fof(f7888,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(subsumption_resolution,[],[f7887,f262])).
fof(f7887,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(forward_demodulation,[],[f7859,f2199])).
fof(f7859,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(superposition,[],[f321,f2522])).
fof(f8063,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15(X11,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15(X11,'_i_0_0',bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f8064,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X14,X13,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X14,X13,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f8072,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X43,X42,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X43,X42,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8073,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X47,X46,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X47,X46,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8095,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,X2))) )),
  inference(superposition,[],[f353,f270])).
fof(f8103,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,happ(s(fun(bool,X31),X29),s(bool,'F_0'))) != s(X31,happ(s(fun(bool,X31),X30),s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK15(X29,X30,bool,X31)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8104,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,happ(s(fun(bool,X35),X33),s(bool,X36))) != s(X35,happ(s(fun(bool,X35),X34),s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK15(X33,X34,bool,X35))) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8111,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15('_i_0_0',X11,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15('_i_0_0',X11,bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f8112,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X13,X14,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,X14,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f8120,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X42,X43,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X42,X43,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8121,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X46,X47,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X46,X47,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2865])).
fof(f8142,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f8070,f2199])).
fof(f8070,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2522])).
fof(f8156,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'T_0'))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(forward_demodulation,[],[f8101,f2199])).
fof(f8101,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(superposition,[],[f353,f2522])).
fof(f8160,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f8118,f2199])).
fof(f8118,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2522])).
fof(f9535,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK8(X0,X1,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool))) )),
  inference(superposition,[],[f339,f270])).
fof(f9543,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,sK8(X29,X30,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),X31)))) | s(bool,X32) = s(bool,sK7(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2865])).
fof(f9544,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,sK8(X33,X34,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK7(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2865])).
fof(f9550,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK7(X0,bool,X1)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK8(X0,bool,X1))) )),
  inference(superposition,[],[f339,f270])).
fof(f9558,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,sK7(X29,bool,X30)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),X31)))) | s(bool,X32) = s(bool,sK8(X29,bool,X30)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2865])).
fof(f9559,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,sK7(X33,bool,X34)),s(bool,X35)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK8(X33,bool,X34))) ) | $spl10),
  inference(superposition,[],[f339,f2865])).
fof(f9593,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'T_0'),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f9541,f2199])).
fof(f9541,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2522])).
fof(f9599,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(forward_demodulation,[],[f9556,f2199])).
fof(f9556,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(superposition,[],[f339,f2522])).
fof(f9607,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f9606,f262])).
fof(f9606,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(forward_demodulation,[],[f9571,f2199])).
fof(f9571,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(superposition,[],[f339,f2522])).
fof(f9740,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9739,f262])).
fof(f9739,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9707,f2199])).
fof(f9707,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f349,f2522])).
fof(f10021,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f10020,f262])).
fof(f10020,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9989,f2199])).
fof(f9989,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f350,f2522])).
fof(f10798,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f10724,f432])).
fof(f10724,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6866,f2865])).
fof(f11040,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),sK19(X8,X7,X6,bool)),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f367,f270])).
fof(f11048,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,X54))),s(X52,X55))) | s(bool,X56) = s(bool,happ(s(fun('type.pair.prod'(X51,X52),bool),sK19(X53,X52,X51,bool)),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl10),
  inference(superposition,[],[f367,f2865])).
fof(f11049,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X58,bool),happ(s(fun(X57,fun(X58,bool)),X59),s(X57,X60))),s(X58,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X57,X58),bool),sK19(X59,X58,X57,bool)),s('type.pair.prod'(X57,X58),'const.pair.,_2'(s(X57,X60),s(X58,X61)))))) ) | $spl10),
  inference(superposition,[],[f367,f2865])).
fof(f11088,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X207,bool),happ(s(fun(X206,fun(X207,bool)),X208),s(X206,X209))),s(X207,X210))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X206,X207),bool),sK19(X208,X207,X206,bool)),s('type.pair.prod'(X206,X207),'const.pair.,_2'(s(X206,X209),s(X207,X210)))))) )),
  inference(superposition,[],[f472,f367])).
fof(f11090,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X222),'const.sum.INL_1'(s(bool,happ(s(fun(X218,bool),happ(s(fun(X217,fun(X218,bool)),X219),s(X217,X220))),s(X218,X221))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X217,X218),bool),sK19(X219,X218,X217,bool)),s('type.pair.prod'(X217,X218),'const.pair.,_2'(s(X217,X220),s(X218,X221)))))) )),
  inference(superposition,[],[f513,f367])).
fof(f11091,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X228,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X224,bool),happ(s(fun(X223,fun(X224,bool)),X225),s(X223,X226))),s(X224,X227))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X223,X224),bool),sK19(X225,X224,X223,bool)),s('type.pair.prod'(X223,X224),'const.pair.,_2'(s(X223,X226),s(X224,X227)))))) )),
  inference(superposition,[],[f540,f367])).
fof(f11095,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X250,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X246,bool),happ(s(fun(X245,fun(X246,bool)),X247),s(X245,X248))),s(X246,X249))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X245,X246),bool),sK19(X247,X246,X245,bool)),s('type.pair.prod'(X245,X246),'const.pair.,_2'(s(X245,X248),s(X246,X249)))))) )),
  inference(superposition,[],[f2475,f367])).
fof(f11096,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X256),'const.sum.INL_1'(s(bool,happ(s(fun(X252,bool),happ(s(fun(X251,fun(X252,bool)),X253),s(X251,X254))),s(X252,X255))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X251,X252),bool),sK19(X253,X252,X251,bool)),s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X254),s(X252,X255)))))) )),
  inference(superposition,[],[f2477,f367])).
fof(f11108,plain,(
  ( ! [X325,X321,X323,X324,X326,X322] : (s(bool,'T_0') != s(bool,happ(s(fun(X322,bool),happ(s(fun(X321,fun(X322,bool)),X323),s(X321,X324))),s(X322,X325))) | s(bool,X326) = s(bool,happ(s(fun('type.pair.prod'(X321,X322),bool),sK19(X323,X322,X321,bool)),s('type.pair.prod'(X321,X322),'const.pair.,_2'(s(X321,X324),s(X322,X325))))) | s(bool,'F_0') = s(bool,X326)) ) | $spl10),
  inference(superposition,[],[f3675,f367])).
fof(f11303,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(forward_demodulation,[],[f11046,f2199])).
fof(f11046,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(superposition,[],[f367,f2522])).
fof(f11549,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X70,bool),'_dst_x0x1_2'(s(fun(X68,bool),X69),s(fun(X70,X68),X71))),s(X70,X72))))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,happ(s(fun(X70,X68),X71),s(X70,X72)))))) ) | $spl10),
  inference(superposition,[],[f2975,f369])).
fof(f11550,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X73,X74),X76),s(X73,X77))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X73,X74),X76))),s(X73,X77)))) ) | $spl10),
  inference(superposition,[],[f2975,f369])).
fof(f11552,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),X82))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X81),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X81),X82)))) ) | $spl10),
  inference(superposition,[],[f2975,f338])).
fof(f11553,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X84,bool),happ(s(fun(X83,fun(X84,bool)),X85),s(X83,X86))),s(X84,X87))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X83,X84),bool),sK19(X85,X84,X83,bool)),s('type.pair.prod'(X83,X84),'const.pair.,_2'(s(X83,X86),s(X84,X87)))))) ) | $spl10),
  inference(superposition,[],[f2975,f367])).
fof(f11555,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) ) | $spl10),
  inference(superposition,[],[f2975,f324])).
fof(f11560,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103))) | s(bool,'T_0') = s(bool,sK4(X103,X104,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X104)))) ) | $spl10),
  inference(superposition,[],[f2975,f836])).
fof(f11627,plain,(
  ( ! [X92] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X92),s(bool,sK15(X92,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X92) | s(bool,'T_0') = s(bool,sK15(X92,'_i_0_0',bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2975])).
fof(f11628,plain,(
  ( ! [X93] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X93),s(bool,sK15('_i_0_0',X93,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X93) | s(bool,'T_0') = s(bool,sK15('_i_0_0',X93,bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2975])).
fof(f11688,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) ) | $spl10),
  inference(subsumption_resolution,[],[f11606,f264])).
fof(f11606,plain,(
  ( ! [X48,X49] : (p(s(bool,'F_0')) | s(bool,X49) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) ) | $spl10),
  inference(superposition,[],[f2482,f2975])).
fof(f11868,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f11786,f432])).
fof(f11786,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5395,f2865])).
fof(f12039,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f12038,f262])).
fof(f12038,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f12013,f2199])).
fof(f12013,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f345,f2522])).
fof(f18671,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31))) | s(bool,X31) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f18588,f432])).
fof(f18588,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f11678,f2865])).
fof(f19303,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f19302,f262])).
fof(f19302,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f19276,f2199])).
fof(f19276,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f346,f2522])).
fof(f19694,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f19693,f262])).
fof(f19693,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f19667,f2199])).
fof(f19667,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f347,f2522])).
fof(f21180,plain,(
  ~p(s(bool,'F_0')) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f21179,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1497,f1374,f1376,f1378,f2515,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4364,f4365,f4516,f4517,f4605,f4611,f4617,f4686,f4813,f4840,f4916,f4915,f4909,f5038,f5122,f5121,f5380,f5391,f5396,f5400,f5787,f6722,f6866,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6980,f6981,f6992,f7142,f7147,f7159,f7163,f7165,f7172,f7174,f7235,f7240,f7252,f7256,f7258,f7265,f7267,f7420,f7422,f7424,f7425,f7426,f7427,f7428,f7685,f7694,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8060,f8061,f8062,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8096,f8097,f8108,f8109,f8110,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9611,f9736,f9737,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10754,f10756,f10758,f10759,f10760,f10761,f10762,f11056,f11347,f11348,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f5378,f5377,f5376,f5375,f5373,f5371,f5370,f4812,f4854,f5379,f5372,f7242,f7244,f7149,f7151,f7243,f7245,f7150,f7152,f5222,f5212,f5059,f5049,f356,f355,f354,f7135,f11109,f7229,f5374,f7893,f7841,f7840,f11096,f11095,f11091,f11090,f11088,f11040,f9550,f9535,f8112,f8111,f8095,f8064,f8063,f7884,f7839,f7218,f7217,f7216,f7215,f7211,f7210,f7208,f7198,f7197,f7124,f7123,f7122,f7121,f7117,f7116,f7114,f7104,f7103,f6979,f6876,f6782,f6781,f6779,f6777,f6776,f6775,f6774,f6732,f6639,f6638,f6636,f6634,f6633,f6632,f6631,f5404,f5392,f5390,f4585,f4584,f4580,f4579,f4577,f4567,f4566,f4515,f4509,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4375,f4363,f4360,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f2792,f2521,f2520,f2483,f2482,f2480,f2478,f2477,f2476,f2475,f2474,f2473,f2427,f1488,f909,f743,f686,f633,f540,f513,f472,f276,f275,f264,f20179,f20181,f11707,f5776,f4609,f8094,f8059,f4715,f4578,f4576,f4571,f4570,f4568,f4565,f4564,f1515,f4529,f6563,f6566,f2803,f4531,f11124,f11128,f7145,f7238,f4685,f4684,f4600,f4536,f4535,f4534,f4533,f4532,f4530,f4528,f4527,f930,f834,f833,f832,f831,f830,f829,f318,f20180,f10755,f9627,f7421,f2037,f1486,f1484,f9958,f9676,f9959,f9677,f9971,f9689,f9972,f9690,f9981,f9968,f9699,f9686,f5256,f5093,f398,f350,f349,f2017,f9585,f4153,f1533,f9582,f6588,f10752,f7418,f7098,f9587,f4155,f2841,f9524,f9525,f11353,f7141,f11070,f7143,f7234,f7236,f11113,f11112,f11073,f11072,f11042,f11041,f7146,f7239,f11340,f9534,f9591,f9590,f9589,f9588,f9586,f9583,f9584,f9604,f339,f11071,f7664,f4601,f4160,f4159,f4158,f4157,f4156,f4154,f4152,f4151,f420,f311,f7695,f7894,f9536,f9537,f9551,f9552,f7679,f7677,f11337,f7674,f7583,f7584,f7662,f4599,f1391,f1390,f1389,f1388,f1387,f1386,f1385,f319,f1392,f7099,f12005,f12026,f12023,f11994,f11995,f12027,f12029,f9581,f7193,f7190,f12036,f12033,f12032,f12031,f12030,f12028,f12025,f12024,f11991,f11990,f11989,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11981,f11980,f11979,f11978,f10753,f7663,f7419,f7209,f7207,f7202,f7201,f7199,f7196,f7195,f7192,f7191,f7115,f7113,f7108,f7107,f7105,f7102,f7101,f348,f345,f19691,f19690,f19688,f19687,f19686,f19685,f19684,f19683,f19682,f19681,f19680,f19679,f19678,f19659,f19649,f19648,f19645,f19644,f19643,f19642,f19641,f19640,f19639,f19638,f19637,f19636,f19635,f19634,f19633,f19632,f19300,f19299,f19297,f19296,f19295,f19294,f19293,f19292,f19291,f19290,f19289,f19288,f19287,f19268,f19258,f19257,f19254,f19253,f19252,f19251,f19250,f19249,f19248,f19247,f19246,f19245,f19244,f19243,f19242,f19241,f12035,f11338,f7885,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7867,f7838,f7829,f7828,f7688,f7683,f7682,f7681,f7680,f7678,f7676,f7675,f7593,f4833,f4754,f4681,f4598,f2032,f347,f346,f321,f320,f317,f10757,f7423,f2766,f2764,f4376,f4377])).
fof(f21179,plain,(
  ( ! [X76,X75] : (~p(s(bool,'F_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),X76))))) ) | ($spl10 | $spl26)),
  inference(backward_demodulation,[],[f21178,f20180])).
fof(f21178,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool))) ) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f20182,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1497,f1374,f1376,f1378,f2515,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4364,f4365,f4516,f4517,f4605,f4611,f4617,f4686,f4813,f4840,f4916,f4915,f4909,f5038,f5122,f5121,f5380,f5391,f5396,f5400,f5787,f6722,f6866,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6980,f6981,f6992,f7142,f7147,f7159,f7163,f7165,f7172,f7174,f7235,f7240,f7252,f7256,f7258,f7265,f7267,f7420,f7422,f7424,f7425,f7426,f7427,f7428,f7685,f7694,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8060,f8061,f8062,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8096,f8097,f8108,f8109,f8110,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9611,f9736,f9737,f9826,f9827,f10018,f10064,f10067,f10069,f10071,f10072,f10073,f10074,f10075,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10754,f10756,f10758,f10759,f10760,f10761,f10762,f11056,f11347,f11348,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f5378,f5377,f5376,f5375,f5373,f5371,f5370,f4812,f4854,f5379,f5372,f7242,f7244,f7149,f7151,f7243,f7245,f7150,f7152,f5222,f5212,f5059,f5049,f356,f355,f354,f7135,f11109,f7229,f5374,f7893,f7841,f7840,f11446,f11096,f11095,f11091,f11090,f11088,f11040,f10793,f10464,f10463,f10462,f10461,f10460,f10458,f10456,f10453,f9837,f9670,f9550,f9535,f8112,f8111,f8095,f8064,f8063,f7884,f7839,f7218,f7217,f7216,f7215,f7211,f7210,f7208,f7198,f7197,f7124,f7123,f7122,f7121,f7117,f7116,f7114,f7104,f7103,f6979,f6876,f6782,f6781,f6779,f6777,f6776,f6775,f6774,f6732,f6639,f6638,f6636,f6634,f6633,f6632,f6631,f5404,f5392,f5390,f4585,f4584,f4580,f4579,f4577,f4567,f4566,f4515,f4509,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4375,f4363,f4360,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f2792,f2521,f2520,f2483,f2482,f2480,f2478,f2477,f2476,f2475,f2474,f2473,f2427,f1488,f909,f743,f686,f633,f540,f513,f472,f276,f275,f264,f19750,f19753,f19756,f19758,f19759,f19760,f19761,f19762,f20179,f20181,f11707,f5776,f4609,f8094,f8059,f4715,f4578,f4576,f4571,f4570,f4568,f4565,f4564,f1515,f4529,f6563,f6566,f2803,f4531,f11124,f11128,f7145,f7238,f4685,f4684,f4600,f4536,f4535,f4534,f4533,f4532,f4530,f4528,f4527,f930,f834,f833,f832,f831,f830,f829,f318,f20180,f10755,f9627,f21177,f7421,f2037,f1486,f1484,f9958,f9676,f9959,f9677,f9971,f9689,f9972,f9690,f9981,f9968,f9699,f9686,f5256,f5093,f398,f350,f349,f2017,f9585,f4153,f1533,f9582,f6588,f10752,f7418,f7098,f9587,f4155,f2841,f9524,f9525,f11353,f7141,f11070,f7143,f7234,f7236,f11113,f11112,f11073,f11072,f11042,f11041,f7146,f7239,f11340,f9534,f9591,f9590,f9589,f9588,f9586,f9583,f9584,f9604,f339,f11071,f7664,f4601,f4160,f4159,f4158,f4157,f4156,f4154,f4152,f4151,f420,f311,f7695,f7894,f9536,f9537,f9551,f9552,f7679,f7677,f11337,f7674,f7583,f7584,f7662,f4599,f1391,f1390,f1389,f1388,f1387,f1386,f1385,f319,f1392,f7099,f12005,f12026,f12023,f11994,f11995,f12027,f12029,f9581,f7193,f7190,f12036,f12033,f12032,f12031,f12030,f12028,f12025,f12024,f11991,f11990,f11989,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11981,f11980,f11979,f11978,f10753,f7663,f7419,f7209,f7207,f7202,f7201,f7199,f7196,f7195,f7192,f7191,f7115,f7113,f7108,f7107,f7105,f7102,f7101,f348,f345,f19691,f19690,f19688,f19687,f19686,f19685,f19684,f19683,f19682,f19681,f19680,f19679,f19678,f19659,f19649,f19648,f19645,f19644,f19643,f19642,f19641,f19640,f19639,f19638,f19637,f19636,f19635,f19634,f19633,f19632,f19300,f19299,f19297,f19296,f19295,f19294,f19293,f19292,f19291,f19290,f19289,f19288,f19287,f19268,f19258,f19257,f19254,f19253,f19252,f19251,f19250,f19249,f19248,f19247,f19246,f19245,f19244,f19243,f19242,f19241,f12035,f11338,f7885,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7867,f7838,f7829,f7828,f7688,f7683,f7682,f7681,f7680,f7678,f7676,f7675,f7593,f4833,f4754,f4681,f4598,f2032,f347,f346,f321,f320,f317,f10757,f7423,f2766,f2764,f4376,f4377])).
fof(f20182,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22)))) ) | $spl10),
  inference(global_subsumption,[],[f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1497,f1374,f1376,f1378,f2036,f2199,f2180,f2297,f2515,f2804,f2810,f2865,f3011,f3027,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f3463,f3604,f3675,f3768,f4178,f4364,f4365,f4516,f4517,f4682,f4605,f4611,f4617,f4686,f4813,f4840,f4916,f4915,f4909,f5038,f5122,f5125,f5121,f5380,f5391,f5393,f5395,f5396,f5400,f5787,f6722,f6866,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6980,f6981,f6992,f7142,f7147,f7159,f7163,f7165,f7172,f7174,f7235,f7240,f7252,f7256,f7258,f7265,f7267,f7420,f7422,f7424,f7425,f7426,f7427,f7428,f7685,f7694,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8060,f8061,f8062,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8096,f8097,f8108,f8109,f8110,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9611,f9660,f9736,f9737,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10754,f10756,f10758,f10759,f10760,f10761,f10762,f11056,f11347,f11348,f11153,f11159,f11167,f11177,f11184,f11678,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11873,f11875,f5378,f5377,f5376,f5375,f5373,f5371,f5370,f4812,f4854,f11707,f5776,f4609,f8094,f8059,f4715,f5379,f4578,f4576,f4571,f4570,f4568,f4565,f4564,f1515,f4529,f6563,f6566,f11124,f11128,f7145,f7238,f2803,f4531,f4685,f4684,f4600,f4536,f4535,f4534,f4533,f4532,f4530,f4528,f4527,f930,f834,f833,f832,f831,f830,f829,f318,f10755,f7421,f5372,f2037,f1486,f1484,f7242,f7244,f9677,f9959,f7149,f7151,f9676,f9958,f7243,f7245,f9690,f9972,f7150,f7152,f9689,f9971,f9981,f9968,f9699,f9686,f5256,f5222,f5212,f5093,f5059,f5049,f356,f355,f354,f350,f349,f398,f2017,f9585,f1533,f4153,f9582,f10752,f7418,f7135,f7098,f6588,f9587,f2841,f4155,f11353,f9525,f9524,f11340,f7239,f7146,f11070,f7234,f7236,f7141,f7143,f11113,f11112,f11109,f11073,f11072,f11042,f11041,f11071,f9534,f9604,f9591,f9590,f9589,f9588,f9586,f9584,f9583,f339,f9552,f9551,f9537,f9536,f7894,f7695,f7664,f4601,f4160,f4159,f4158,f4157,f4156,f4154,f4152,f4151,f420,f311,f7679,f11337,f7677,f7674,f7584,f7583,f7662,f4599,f1391,f1390,f1389,f1388,f1387,f1385,f1386,f319,f1392,f7099,f12005,f12026,f12023,f12029,f12027,f11995,f11994,f9581,f7193,f7190,f12036,f12033,f12032,f12031,f12030,f12028,f12025,f12024,f11991,f11990,f11989,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11981,f11980,f11979,f11978,f10753,f7663,f7419,f7229,f7209,f7207,f7202,f7201,f7199,f7196,f7195,f7192,f7191,f7115,f7113,f7108,f7107,f7105,f7102,f7101,f348,f345,f12035,f11338,f7885,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7867,f7838,f7829,f7828,f7688,f7683,f7682,f7681,f7680,f7678,f7676,f7675,f7593,f4833,f4754,f4681,f4598,f2032,f347,f346,f321,f320,f317,f10757,f7423,f5374,f2766,f2764,f7893,f7841,f7840,f4376,f4377,f12039,f11868,f11688,f11628,f11627,f11560,f11555,f11553,f11552,f11550,f11549,f11303,f11108,f11096,f11095,f11091,f11090,f11088,f11049,f11048,f11040,f10798,f10021,f9740,f9607,f9599,f9593,f9559,f9558,f9550,f9544,f9543,f9535,f8160,f8156,f8142,f8121,f8120,f8112,f8111,f8104,f8103,f8095,f8073,f8072,f8064,f8063,f7888,f7884,f7880,f7848,f7847,f7839,f7691,f7457,f7357,f7228,f7218,f7217,f7216,f7215,f7211,f7210,f7208,f7198,f7197,f7134,f7124,f7123,f7122,f7121,f7117,f7116,f7114,f7104,f7103,f6988,f6987,f6979,f6876,f6782,f6781,f6779,f6777,f6776,f6775,f6774,f6732,f6639,f6638,f6636,f6634,f6633,f6632,f6631,f5404,f5392,f5390,f5383,f5259,f5096,f4961,f4859,f4836,f4680,f4595,f4585,f4584,f4580,f4579,f4577,f4567,f4566,f4524,f4523,f4515,f4509,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4391,f4388,f4384,f4383,f4375,f4372,f4371,f4363,f4360,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f4219,f4218,f4217,f4216,f4215,f4214,f4213,f4212,f4211,f4165,f4116,f4109,f3674,f3499,f3498,f3497,f3496,f3495,f3494,f3493,f3492,f3428,f3427,f3426,f3425,f3424,f3423,f3422,f3421,f3420,f3013,f2979,f2977,f2975,f2792,f2522,f2521,f2520,f2483,f2482,f2480,f2478,f2477,f2476,f2475,f2474,f2473,f2427,f1488,f909,f743,f686,f633,f540,f513,f472,f276,f275,f264,f18671,f18620,f18621,f18623,f18624,f18626,f18631,f19241,f19242,f19243,f19244,f19245,f19246,f19247,f19248,f19249,f19250,f19251,f19252,f19253,f19254,f19287,f19288,f19257,f19258,f19289,f19290,f19291,f19292,f19293,f19294,f19295,f19296,f19297,f19268,f19299,f19300,f19303,f19632,f19633,f19634,f19635,f19636,f19637,f19638,f19639,f19640,f19641,f19642,f19643,f19644,f19645,f19678,f19679,f19648,f19649,f19680,f19681,f19682,f19683,f19684,f19685,f19686,f19687,f19688,f19659,f19690,f19691,f19694,f20178,f20179,f20180,f20181,f9627,f9630,f9661,f9666])).
fof(f4377,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK14(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK14(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f351,f445])).
fof(f4376,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK14(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK14(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f2764,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f324,f445])).
fof(f2766,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f324,f445])).
fof(f7423,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X114,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X112),s('type.num.num',X113))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X111),s(fun('type.num.num',bool),X112),s('type.num.num','const.num.SUC_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f6722,f324])).
fof(f10757,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X138),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X136),s('type.num.num',X137))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X135),s(fun('type.num.num',bool),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))))) )),
  inference(superposition,[],[f6866,f324])).
fof(f18626,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X99),s('type.num.num',X100))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X98),s(fun('type.num.num',bool),X99),s('type.num.num','const.num.SUC_1'(s('type.num.num',X100)))))) ) | $spl10),
  inference(superposition,[],[f11678,f324])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9] : ! [X34,X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,'const.min.@_1'(s(fun(X9,bool),X34)))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.SELECT_AX')).
fof(f320,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X34] : (! [X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))) => (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => ? [X38] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.MONO_EXISTS')).
fof(f321,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f346,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X31] : ! [X37,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) => (! [X34] : (! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X34] : (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X6,X39] : (s(X9,X6) = s(X9,happ(s(fun(X31,X9),X37),s(X31,X39))) <=> s(X31,X39) = s(X31,happ(s(fun(X9,X31),X10),s(X9,X6))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.ISO_USAGE')).
fof(f347,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f2032,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2011,f262])).
fof(f2011,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f317,f445])).
fof(f4598,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f317,f836])).
fof(f4681,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4548,f271])).
fof(f4548,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f319,f836])).
fof(f4754,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4681,f836])).
fof(f4833,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4800,f262])).
fof(f4800,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4681,f445])).
fof(f7593,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,X73)))) | p(s(bool,happ(s(fun(X70,bool),X74),s(X70,sK6(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK5(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f320,f836])).
fof(f7675,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(forward_demodulation,[],[f7582,f309])).
fof(f7582,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f7676,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(forward_demodulation,[],[f7585,f357])).
fof(f7585,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(superposition,[],[f320,f357])).
fof(f7678,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(forward_demodulation,[],[f7587,f316])).
fof(f7587,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(superposition,[],[f320,f316])).
fof(f7680,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(forward_demodulation,[],[f7589,f337])).
fof(f7589,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(superposition,[],[f320,f337])).
fof(f7681,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(forward_demodulation,[],[f7590,f341])).
fof(f7590,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(superposition,[],[f320,f341])).
fof(f7682,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(forward_demodulation,[],[f7591,f336])).
fof(f7591,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(superposition,[],[f320,f336])).
fof(f7683,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(forward_demodulation,[],[f7592,f342])).
fof(f7592,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(superposition,[],[f320,f342])).
fof(f7688,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7649,f262])).
fof(f7649,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f320,f445])).
fof(f7828,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,sK5(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10))))) | ~p(s(bool,happ(s(fun(X10,bool),X16),s(X10,X17)))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK6(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f7829,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,sK5(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18))))) | ~p(s(bool,happ(s(fun(X18,bool),X24),s(X18,X25)))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK6(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f7838,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | ~p(s(bool,happ(s(fun(X70,bool),X73),s(X70,X74)))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK6(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f321,f836])).
fof(f7867,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,happ(s(fun(X66,X67),X69),s(X66,sK5(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66))))))) | ~p(s(bool,happ(s(fun(X66,bool),X70),s(X66,X71)))) | p(s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69))),s(X66,sK6(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f7869,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f7826,f368])).
fof(f7826,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0)))))) )),
  inference(superposition,[],[f321,f368])).
fof(f7870,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f7827,f309])).
fof(f7827,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f7871,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31))))) )),
  inference(forward_demodulation,[],[f7830,f357])).
fof(f7830,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31)))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26)))))) )),
  inference(superposition,[],[f321,f357])).
fof(f7872,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36))))) )),
  inference(forward_demodulation,[],[f7831,f338])).
fof(f7831,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36)))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32)))))) )),
  inference(superposition,[],[f321,f338])).
fof(f7873,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41))))) )),
  inference(forward_demodulation,[],[f7832,f316])).
fof(f7832,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37)))))) )),
  inference(superposition,[],[f321,f316])).
fof(f7874,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47))))) )),
  inference(forward_demodulation,[],[f7833,f324])).
fof(f7833,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47)))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)))))) )),
  inference(superposition,[],[f321,f324])).
fof(f7875,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52))))) )),
  inference(forward_demodulation,[],[f7834,f337])).
fof(f7834,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48)))))) )),
  inference(superposition,[],[f321,f337])).
fof(f7876,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58))))) )),
  inference(forward_demodulation,[],[f7835,f341])).
fof(f7835,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58)))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53)))))) )),
  inference(superposition,[],[f321,f341])).
fof(f7877,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63))))) )),
  inference(forward_demodulation,[],[f7836,f336])).
fof(f7836,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63)))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59)))))) )),
  inference(superposition,[],[f321,f336])).
fof(f7878,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69))))) )),
  inference(forward_demodulation,[],[f7837,f342])).
fof(f7837,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69)))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64)))))) )),
  inference(superposition,[],[f321,f342])).
fof(f7885,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7854,f262])).
fof(f7854,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f321,f445])).
fof(f11338,plain,(
  ( ! [X445,X441,X443,X444,X446,X440,X442,X439] : (p(s(bool,happ(s(fun(X439,bool),happ(s(fun(X441,fun(X439,bool)),happ(s(fun(X440,fun(X441,fun(X439,bool))),X442),s(X440,X443))),s(X441,X444))),s(X439,sK6(X445,happ(s(fun('type.pair.prod'(X440,X441),fun(X439,bool)),sK19(X442,X441,X440,fun(X439,bool))),s('type.pair.prod'(X440,X441),'const.pair.,_2'(s(X440,X443),s(X441,X444)))),X439))))) | ~p(s(bool,happ(s(fun(X439,bool),happ(s(fun(X441,fun(X439,bool)),happ(s(fun(X440,fun(X441,fun(X439,bool))),X442),s(X440,X443))),s(X441,X444))),s(X439,sK5(X445,happ(s(fun('type.pair.prod'(X440,X441),fun(X439,bool)),sK19(X442,X441,X440,fun(X439,bool))),s('type.pair.prod'(X440,X441),'const.pair.,_2'(s(X440,X443),s(X441,X444)))),X439))))) | ~p(s(bool,happ(s(fun(X439,bool),X445),s(X439,X446))))) )),
  inference(forward_demodulation,[],[f11127,f367])).
fof(f11127,plain,(
  ( ! [X445,X441,X443,X444,X446,X440,X442,X439] : (~p(s(bool,happ(s(fun(X439,bool),happ(s(fun(X441,fun(X439,bool)),happ(s(fun(X440,fun(X441,fun(X439,bool))),X442),s(X440,X443))),s(X441,X444))),s(X439,sK5(X445,happ(s(fun('type.pair.prod'(X440,X441),fun(X439,bool)),sK19(X442,X441,X440,fun(X439,bool))),s('type.pair.prod'(X440,X441),'const.pair.,_2'(s(X440,X443),s(X441,X444)))),X439))))) | ~p(s(bool,happ(s(fun(X439,bool),X445),s(X439,X446)))) | p(s(bool,happ(s(fun(X439,bool),happ(s(fun('type.pair.prod'(X440,X441),fun(X439,bool)),sK19(X442,X441,X440,fun(X439,bool))),s('type.pair.prod'(X440,X441),'const.pair.,_2'(s(X440,X443),s(X441,X444))))),s(X439,sK6(X445,happ(s(fun('type.pair.prod'(X440,X441),fun(X439,bool)),sK19(X442,X441,X440,fun(X439,bool))),s('type.pair.prod'(X440,X441),'const.pair.,_2'(s(X440,X443),s(X441,X444)))),X439)))))) )),
  inference(superposition,[],[f321,f367])).
fof(f12035,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f12034,f309])).
fof(f12034,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f12006,f262])).
fof(f12006,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f345,f440])).
fof(f19241,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK11(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f19242,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK11(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f19243,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK11(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f19244,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK11(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f19245,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK11(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41))))) )),
  inference(superposition,[],[f346,f357])).
fof(f19246,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK11(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48))))) )),
  inference(superposition,[],[f346,f338])).
fof(f19247,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK11(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58))))) )),
  inference(superposition,[],[f346,f367])).
fof(f19248,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK11(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f19249,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK11(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f19250,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK11(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f19251,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK11(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f19252,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK11(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f19253,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK11(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f19254,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK11(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f19257,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK11(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f19258,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK11(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f19268,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK11(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f19287,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f19255,f368])).
fof(f19255,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f19288,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f19256,f309])).
fof(f19256,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f19289,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f19259,f357])).
fof(f19259,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(superposition,[],[f346,f357])).
fof(f19290,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f19260,f338])).
fof(f19260,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(superposition,[],[f346,f338])).
fof(f19291,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f19261,f367])).
fof(f19261,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(superposition,[],[f346,f367])).
fof(f19292,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f19262,f316])).
fof(f19262,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f19293,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f19263,f324])).
fof(f19263,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f19294,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f19264,f337])).
fof(f19264,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f19295,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f19265,f341])).
fof(f19265,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f19296,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f19266,f336])).
fof(f19266,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f19297,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f19267,f342])).
fof(f19267,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f19299,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f19298,f309])).
fof(f19298,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(subsumption_resolution,[],[f19269,f262])).
fof(f19269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(superposition,[],[f346,f440])).
fof(f19300,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f19271,f262])).
fof(f19271,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f346,f445])).
fof(f19632,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f19633,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK12(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f19634,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23)))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK12(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f19635,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33)))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK12(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f19636,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41)))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK12(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f19637,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK12(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f19638,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58)))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK12(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f19639,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65)))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK12(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f19640,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73)))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK12(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f19641,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80)))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK12(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f19642,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88)))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK12(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f19643,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95)))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK12(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f19644,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103)))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK12(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f19645,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK12(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f19648,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23)))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK12(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f19649,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33)))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK12(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f19659,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK12(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f19678,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f19646,f368])).
fof(f19646,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f19679,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f19647,f309])).
fof(f19647,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f19680,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f19650,f357])).
fof(f19650,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f19681,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f19651,f338])).
fof(f19651,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f19682,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f19652,f367])).
fof(f19652,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58)))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f19683,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f19653,f316])).
fof(f19653,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f19684,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f19654,f324])).
fof(f19654,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73)))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f19685,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f19655,f337])).
fof(f19655,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f19686,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f19656,f341])).
fof(f19656,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f19687,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f19657,f336])).
fof(f19657,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f19688,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f19658,f342])).
fof(f19658,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f19690,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f19689,f309])).
fof(f19689,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(subsumption_resolution,[],[f19660,f262])).
fof(f19660,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(superposition,[],[f347,f440])).
fof(f19691,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f19662,f262])).
fof(f19662,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f347,f445])).
fof(f345,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f348,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f7101,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X108,bool),X109),s(fun(X110,X108),X111))),s(X110,X112)))) | s(bool,'T_0') != s(bool,happ(s(fun(X108,bool),X109),s(X108,happ(s(fun(X110,X108),X111),s(X110,X112)))))) )),
  inference(superposition,[],[f271,f369])).
fof(f7102,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117)))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f272,f369])).
fof(f7105,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X130,X128),X131))),s(X130,X132)))) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f330,f369])).
fof(f7107,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | ~p(s(bool,X145)) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) )),
  inference(superposition,[],[f333,f369])).
fof(f7108,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,X151) != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150))))) | p(s(bool,X151))) )),
  inference(superposition,[],[f418,f369])).
fof(f7113,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X178,bool),'_dst_x0x1_2'(s(fun(X176,bool),X177),s(fun(X178,X176),X179))),s(X178,X180)))) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),X177),s(X176,happ(s(fun(X178,X176),X179),s(X178,X180))))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f447,f369])).
fof(f7115,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (p(s(bool,happ(s(fun(X189,bool),'_dst_x0x1_2'(s(fun(X187,bool),X188),s(fun(X189,X187),X190))),s(X189,X191)))) | s(bool,X192) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f494,f369])).
fof(f7191,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))) )),
  inference(superposition,[],[f4681,f369])).
fof(f7192,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun(X91,X92),X94),s(X91,X95)))))) | p(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))),s(X91,'const.min.@_1'(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f7195,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f271,f369])).
fof(f7196,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114)))) )),
  inference(superposition,[],[f272,f369])).
fof(f7199,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129)))))) | p(s(bool,X130)) | s(bool,X130) = s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f330,f369])).
fof(f7201,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | ~p(s(bool,X142)) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f333,f369])).
fof(f7202,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X144,bool),X145),s(X144,happ(s(fun(X143,X144),X146),s(X143,X147)))))) | s(bool,X148) != s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X144,bool),X145),s(fun(X143,X144),X146))),s(X143,X147))) | p(s(bool,X148))) )),
  inference(superposition,[],[f418,f369])).
fof(f7207,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X173,X174),X176),s(X173,X177)))))) | s(bool,X178) != s(bool,happ(s(fun(X173,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X173,X174),X176))),s(X173,X177))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f447,f369])).
fof(f7209,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X184,X185),X187),s(X184,X188)))))) | s(bool,X189) = s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X184,X185),X187))),s(X184,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f494,f369])).
fof(f7419,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f6722,f369])).
fof(f7663,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | p(s(bool,happ(s(fun(X72,bool),X77),s(X72,sK6('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72))))) | p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,sK5('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f10753,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6866,f369])).
fof(f11978,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK10(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f11979,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK10(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f11980,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK10(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11981,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK10(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11982,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK10(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f11983,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK10(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f11984,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,sK10(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,X58))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f11985,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK10(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f11986,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK10(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f11987,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK10(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f11988,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK10(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f11989,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK10(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f11990,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK10(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f11991,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK10(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f12024,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f11993,f309])).
fof(f11993,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f12025,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f11996,f357])).
fof(f11996,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f12028,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f11999,f316])).
fof(f11999,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f12030,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f12001,f337])).
fof(f12001,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f12031,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f12002,f341])).
fof(f12002,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f12032,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f12003,f336])).
fof(f12003,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f12033,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f12004,f342])).
fof(f12004,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f12036,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f12008,f262])).
fof(f12008,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f345,f445])).
fof(f18621,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,X85))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))),s(X81,X85)))) ) | $spl10),
  inference(superposition,[],[f11678,f369])).
fof(f7190,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK4(X85,'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84)),X81))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f319,f369])).
fof(f7193,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,happ(s(fun('type.list.list'(X96),X97),X99),s('type.list.list'(X96),sK2('_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99)),X96))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X96),bool),'_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99))),s('type.list.list'(X96),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f9581,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun('type.pair.prod'(X72,X73),X74),X76),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,sK7('_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76)),X73,X72)),s(X73,sK8('_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76)),X73,X72))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X72,X73),bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun('type.pair.prod'(X72,X73),X74),X76))),s('type.pair.prod'(X72,X73),X77))))) )),
  inference(superposition,[],[f339,f369])).
fof(f12029,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f12000,f324])).
fof(f12000,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f12027,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f11998,f367])).
fof(f11998,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,X58))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f11995,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK10(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f11994,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK10(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f12023,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f11992,f368])).
fof(f11992,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f12026,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f11997,f338])).
fof(f11997,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f12005,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK10(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f7099,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun(X97,X95),X98))),s(X97,X99)))) | p(s(bool,happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f1392,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1379,f262])).
fof(f1379,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f319,f445])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : ! [X34,X24] : (! [X38] : (s(X9,X24) = s(X9,X38) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X34)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.JRH_INDUCT_UTIL')).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1367,f309])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f319,f309])).
fof(f1386,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1368,f357])).
fof(f1368,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f319,f357])).
fof(f1387,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1369,f316])).
fof(f1369,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f319,f316])).
fof(f1388,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12)))))) )),
  inference(forward_demodulation,[],[f1370,f337])).
fof(f1370,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14))))))))) )),
  inference(superposition,[],[f319,f337])).
fof(f1389,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1371,f341])).
fof(f1371,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f319,f341])).
fof(f1390,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | ~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21)))))) )),
  inference(forward_demodulation,[],[f1372,f336])).
fof(f1372,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23))))))))) )),
  inference(superposition,[],[f319,f336])).
fof(f1391,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X28)))) | ~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25)))))) )),
  inference(forward_demodulation,[],[f1373,f342])).
fof(f1373,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25))))) | p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28))))))))) )),
  inference(superposition,[],[f319,f342])).
fof(f4599,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f319,f836])).
fof(f7662,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70)))) | p(s(bool,happ(s(fun(X66,bool),X71),s(X66,sK6(X67,X71,X66))))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X67,X71,X66)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f7584,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,X24)))) | p(s(bool,happ(s(fun(X18,bool),X25),s(X18,sK6(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18))))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK5(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f7583,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,X16)))) | p(s(bool,happ(s(fun(X10,bool),X17),s(X10,sK6(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10))))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK5(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f7674,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f7581,f368])).
fof(f7581,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(superposition,[],[f320,f368])).
fof(f11337,plain,(
  ( ! [X436,X438,X432,X434,X431,X437,X433,X435] : (p(s(bool,happ(s(fun(X431,bool),happ(s(fun(X433,fun(X431,bool)),happ(s(fun(X432,fun(X433,fun(X431,bool))),X434),s(X432,X435))),s(X433,X436))),s(X431,sK5(happ(s(fun('type.pair.prod'(X432,X433),fun(X431,bool)),sK19(X434,X433,X432,fun(X431,bool))),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436)))),X438,X431))))) | ~p(s(bool,happ(s(fun(X431,bool),happ(s(fun(X433,fun(X431,bool)),happ(s(fun(X432,fun(X433,fun(X431,bool))),X434),s(X432,X435))),s(X433,X436))),s(X431,X437)))) | p(s(bool,happ(s(fun(X431,bool),X438),s(X431,sK6(happ(s(fun('type.pair.prod'(X432,X433),fun(X431,bool)),sK19(X434,X433,X432,fun(X431,bool))),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436)))),X438,X431)))))) )),
  inference(forward_demodulation,[],[f11126,f367])).
fof(f11126,plain,(
  ( ! [X436,X438,X432,X434,X431,X437,X433,X435] : (~p(s(bool,happ(s(fun(X431,bool),happ(s(fun(X433,fun(X431,bool)),happ(s(fun(X432,fun(X433,fun(X431,bool))),X434),s(X432,X435))),s(X433,X436))),s(X431,X437)))) | p(s(bool,happ(s(fun(X431,bool),X438),s(X431,sK6(happ(s(fun('type.pair.prod'(X432,X433),fun(X431,bool)),sK19(X434,X433,X432,fun(X431,bool))),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436)))),X438,X431))))) | p(s(bool,happ(s(fun(X431,bool),happ(s(fun('type.pair.prod'(X432,X433),fun(X431,bool)),sK19(X434,X433,X432,fun(X431,bool))),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436))))),s(X431,sK5(happ(s(fun('type.pair.prod'(X432,X433),fun(X431,bool)),sK19(X434,X433,X432,fun(X431,bool))),s('type.pair.prod'(X432,X433),'const.pair.,_2'(s(X432,X435),s(X433,X436)))),X438,X431)))))) )),
  inference(superposition,[],[f320,f367])).
fof(f7677,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(forward_demodulation,[],[f7586,f338])).
fof(f7586,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(superposition,[],[f320,f338])).
fof(f7679,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(forward_demodulation,[],[f7588,f324])).
fof(f7588,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(superposition,[],[f320,f324])).
fof(f9552,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,sK7(X7,bool,X8)),s(bool,X9)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK8(X7,bool,X8))) )),
  inference(superposition,[],[f339,f445])).
fof(f9551,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,sK7(X3,bool,X4)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),X5)))) | s(bool,X6) = s(bool,sK8(X3,bool,X4)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f9537,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X9),s(X8,sK8(X7,X8,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7(X7,X8,bool))) )),
  inference(superposition,[],[f339,f445])).
fof(f9536,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'T_0'),s(X4,sK8(X3,X4,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),X5)))) | s(bool,X6) = s(bool,sK7(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f7894,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74))))) )),
  inference(forward_demodulation,[],[f7868,f338])).
fof(f7868,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72))))))) )),
  inference(superposition,[],[f321,f338])).
fof(f7695,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(forward_demodulation,[],[f7665,f338])).
fof(f7665,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(superposition,[],[f320,f338])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X34] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X24] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X24)))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X35),s('type.list.list'(X9),X24)))))))) => ! [X36] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X36)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.list.list_INDUCT')).
fof(f420,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f4151,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f4130,f309])).
fof(f4130,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f309])).
fof(f4152,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f4131,f357])).
fof(f4131,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f357])).
fof(f4154,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f4133,f316])).
fof(f4133,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f316])).
fof(f4156,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16)))))) )),
  inference(forward_demodulation,[],[f4135,f337])).
fof(f4135,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f337])).
fof(f4157,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f4136,f341])).
fof(f4136,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f341])).
fof(f4158,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23)))))) )),
  inference(forward_demodulation,[],[f4137,f336])).
fof(f4137,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f336])).
fof(f4159,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26)))))) )),
  inference(forward_demodulation,[],[f4138,f342])).
fof(f4138,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26))))) | ~p(s(bool,happ(s(fun('type.list.list'(X26),bool),'const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29))))),s('type.list.list'(X26),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f342])).
fof(f4160,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f4139,f264])).
fof(f4139,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f420,f270])).
fof(f4601,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f420,f836])).
fof(f7664,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X80)) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),X81),s('type.list.list'(X78),sK6(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78)))))) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),sK18(X79,X80,X78,bool)),s('type.list.list'(X78),sK5(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78))))))) )),
  inference(superposition,[],[f320,f357])).
fof(f11071,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X113),s(X111,X114))),s(X112,X115)))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),X116),s('type.pair.prod'(X111,X112),sK6(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),sK19(X113,X112,X111,bool)),s('type.pair.prod'(X111,X112),sK5(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112))))))) )),
  inference(superposition,[],[f320,f367])).
fof(f339,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9,X31] : ! [X34] : (! [X59,X60] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X59),s(X31,X60)))))) => ! [X61] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),X61)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.pair.pair_induction')).
fof(f9604,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f9566,f262])).
fof(f9566,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f339,f445])).
fof(f9584,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),X31)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f9526,f357])).
fof(f9526,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0'))),s('type.pair.prod'(X26,X27),X31))))) )),
  inference(superposition,[],[f339,f357])).
fof(f9583,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),X9)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)))))))) )),
  inference(forward_demodulation,[],[f9523,f309])).
fof(f9523,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8))),s('type.pair.prod'(X6,X7),X9))))) )),
  inference(superposition,[],[f339,f309])).
fof(f9586,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),X41)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)))))))) )),
  inference(forward_demodulation,[],[f9528,f316])).
fof(f9528,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'(X37,X38),X41))))) )),
  inference(superposition,[],[f339,f316])).
fof(f9588,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),X52)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)))))))) )),
  inference(forward_demodulation,[],[f9530,f337])).
fof(f9530,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51))))),s('type.pair.prod'(X48,X49),X52))))) )),
  inference(superposition,[],[f339,f337])).
fof(f9589,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),X58)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)))))))) )),
  inference(forward_demodulation,[],[f9531,f341])).
fof(f9531,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57))))),s('type.pair.prod'(X53,X54),X58))))) )),
  inference(superposition,[],[f339,f341])).
fof(f9590,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),X63)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)))))))) )),
  inference(forward_demodulation,[],[f9532,f336])).
fof(f9532,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62))))),s('type.pair.prod'(X59,X60),X63))))) )),
  inference(superposition,[],[f339,f336])).
fof(f9591,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),X69)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)))))))) )),
  inference(forward_demodulation,[],[f9533,f342])).
fof(f9533,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),'const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68))))),s('type.pair.prod'(X64,X65),X69))))) )),
  inference(superposition,[],[f339,f342])).
fof(f9534,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),X72),s('type.pair.prod'(X70,X71),'const.pair.,_2'(s(X70,sK7(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70)),s(X71,sK8(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool))),s('type.pair.prod'(X70,X71),X74)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.pair.prod'(X70,X71),bool),bool),X73)))) )),
  inference(superposition,[],[f339,f836])).
fof(f11340,plain,(
  ( ! [X466,X461,X463,X465,X462,X464] : (~p(s(bool,happ(s(fun('type.list.list'(X461),bool),happ(s(fun(X463,fun('type.list.list'(X461),bool)),happ(s(fun(X462,fun(X463,fun('type.list.list'(X461),bool))),X464),s(X462,X465))),s(X463,X466))),s('type.list.list'(X461),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X461),bool),happ(s(fun(X463,fun('type.list.list'(X461),bool)),happ(s(fun(X462,fun(X463,fun('type.list.list'(X461),bool))),X464),s(X462,X465))),s(X463,X466))),s('type.list.list'(X461),sK2(happ(s(fun('type.pair.prod'(X462,X463),fun('type.list.list'(X461),bool)),sK19(X464,X463,X462,fun('type.list.list'(X461),bool))),s('type.pair.prod'(X462,X463),'const.pair.,_2'(s(X462,X465),s(X463,X466)))),X461)))))) )),
  inference(forward_demodulation,[],[f11130,f367])).
fof(f11130,plain,(
  ( ! [X466,X461,X463,X465,X462,X464] : (p(s(bool,happ(s(fun('type.list.list'(X461),bool),happ(s(fun(X463,fun('type.list.list'(X461),bool)),happ(s(fun(X462,fun(X463,fun('type.list.list'(X461),bool))),X464),s(X462,X465))),s(X463,X466))),s('type.list.list'(X461),sK2(happ(s(fun('type.pair.prod'(X462,X463),fun('type.list.list'(X461),bool)),sK19(X464,X463,X462,fun('type.list.list'(X461),bool))),s('type.pair.prod'(X462,X463),'const.pair.,_2'(s(X462,X465),s(X463,X466)))),X461))))) | ~p(s(bool,happ(s(fun('type.list.list'(X461),bool),happ(s(fun('type.pair.prod'(X462,X463),fun('type.list.list'(X461),bool)),sK19(X464,X463,X462,fun('type.list.list'(X461),bool))),s('type.pair.prod'(X462,X463),'const.pair.,_2'(s(X462,X465),s(X463,X466))))),s('type.list.list'(X461),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f367])).
fof(f7239,plain,(
  ( ! [X372,X374,X370,X373,X369,X371] : (p(s(bool,happ(s(fun('type.list.list'(X369),bool),happ(s(fun(X371,fun('type.list.list'(X369),bool)),X372),s(X371,happ(s(fun(X370,X371),X373),s(X370,X374))))),s('type.list.list'(X369),sK2(happ(s(fun(X370,fun('type.list.list'(X369),bool)),'_dst_x0x1_2'(s(fun(X371,fun('type.list.list'(X369),bool)),X372),s(fun(X370,X371),X373))),s(X370,X374)),X369))))) | ~p(s(bool,happ(s(fun('type.list.list'(X369),bool),happ(s(fun(X370,fun('type.list.list'(X369),bool)),'_dst_x0x1_2'(s(fun(X371,fun('type.list.list'(X369),bool)),X372),s(fun(X370,X371),X373))),s(X370,X374))),s('type.list.list'(X369),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f7146,plain,(
  ( ! [X383,X385,X387,X382,X384,X386] : (p(s(bool,happ(s(fun('type.list.list'(X382),bool),happ(s(fun(X385,fun('type.list.list'(X382),bool)),'_dst_x0x1_2'(s(fun(X383,fun('type.list.list'(X382),bool)),X384),s(fun(X385,X383),X386))),s(X385,X387))),s('type.list.list'(X382),sK2(happ(s(fun(X383,fun('type.list.list'(X382),bool)),X384),s(X383,happ(s(fun(X385,X383),X386),s(X385,X387)))),X382))))) | ~p(s(bool,happ(s(fun('type.list.list'(X382),bool),happ(s(fun(X383,fun('type.list.list'(X382),bool)),X384),s(X383,happ(s(fun(X385,X383),X386),s(X385,X387))))),s('type.list.list'(X382),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f11041,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,happ(s(fun('type.pair.prod'(X11,X12),bool),sK19(X13,X12,X11,bool)),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f367,f445])).
fof(f11042,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK19(X19,X18,X17,bool)),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f367,f445])).
fof(f11072,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X117,fun(X118,bool)),X119),s(X117,X120))),s(X118,X121)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X117,X118),bool),sK19(X119,X118,X117,bool)))))) )),
  inference(superposition,[],[f4681,f367])).
fof(f11073,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X122,fun(X123,bool)),X124),s(X122,X125))),s(X123,X126)))) | p(s(bool,happ(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)),s('type.pair.prod'(X122,X123),'const.min.@_1'(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)))))))) )),
  inference(superposition,[],[f317,f367])).
fof(f11112,plain,(
  ( ! [X350,X346,X349,X351,X347,X348] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X351,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X347,bool),happ(s(fun(X346,fun(X347,bool)),X348),s(X346,X349))),s(X347,X350))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X346,X347),bool),sK19(X348,X347,X346,bool)),s('type.pair.prod'(X346,X347),'const.pair.,_2'(s(X346,X349),s(X347,X350)))))) )),
  inference(superposition,[],[f6722,f367])).
fof(f11113,plain,(
  ( ! [X356,X352,X354,X357,X353,X355] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X357),'const.sum.INL_1'(s(bool,happ(s(fun(X353,bool),happ(s(fun(X352,fun(X353,bool)),X354),s(X352,X355))),s(X353,X356))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X352,X353),bool),sK19(X354,X353,X352,bool)),s('type.pair.prod'(X352,X353),'const.pair.,_2'(s(X352,X355),s(X353,X356)))))) )),
  inference(superposition,[],[f6866,f367])).
fof(f18624,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X92,bool),happ(s(fun(X91,fun(X92,bool)),X93),s(X91,X94))),s(X92,X95))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X91,X92),bool),sK19(X93,X92,X91,bool)),s('type.pair.prod'(X91,X92),'const.pair.,_2'(s(X91,X94),s(X92,X95)))))) ) | $spl10),
  inference(superposition,[],[f11678,f367])).
fof(f7236,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X348] : (~p(s(bool,happ(s(fun(X348,bool),happ(s(fun(X350,fun(X348,bool)),X351),s(X350,happ(s(fun(X349,X350),X352),s(X349,X353))))),s(X348,X354)))) | p(s(bool,happ(s(fun(X348,bool),happ(s(fun(X349,fun(X348,bool)),'_dst_x0x1_2'(s(fun(X350,fun(X348,bool)),X351),s(fun(X349,X350),X352))),s(X349,X353))),s(X348,'const.min.@_1'(s(fun(X348,bool),happ(s(fun(X349,fun(X348,bool)),'_dst_x0x1_2'(s(fun(X350,fun(X348,bool)),X351),s(fun(X349,X350),X352))),s(X349,X353))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f7234,plain,(
  ( ! [X335,X337,X339,X334,X340,X336,X338] : (~p(s(bool,happ(s(fun(X334,bool),happ(s(fun(X336,fun(X334,bool)),X337),s(X336,happ(s(fun(X335,X336),X338),s(X335,X339))))),s(X334,X340)))) | p(s(bool,'|dollar|exists_1'(s(fun(X334,bool),happ(s(fun(X335,fun(X334,bool)),'_dst_x0x1_2'(s(fun(X336,fun(X334,bool)),X337),s(fun(X335,X336),X338))),s(X335,X339))))))) )),
  inference(superposition,[],[f4681,f369])).
fof(f7143,plain,(
  ( ! [X364,X366,X362,X365,X367,X361,X363] : (~p(s(bool,happ(s(fun(X361,bool),happ(s(fun(X364,fun(X361,bool)),'_dst_x0x1_2'(s(fun(X362,fun(X361,bool)),X363),s(fun(X364,X362),X365))),s(X364,X366))),s(X361,X367)))) | p(s(bool,happ(s(fun(X361,bool),happ(s(fun(X362,fun(X361,bool)),X363),s(X362,happ(s(fun(X364,X362),X365),s(X364,X366))))),s(X361,'const.min.@_1'(s(fun(X361,bool),happ(s(fun(X362,fun(X361,bool)),X363),s(X362,happ(s(fun(X364,X362),X365),s(X364,X366))))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f11070,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(X108,bool),happ(s(fun(X107,fun(X108,bool)),X109),s(X107,sK7(sK19(X109,X108,X107,bool),X108,X107)))),s(X108,sK8(sK19(X109,X108,X107,bool),X108,X107))))) | p(s(bool,happ(s(fun('type.pair.prod'(X107,X108),bool),sK19(X109,X108,X107,bool)),s('type.pair.prod'(X107,X108),X110))))) )),
  inference(superposition,[],[f339,f367])).
fof(f7141,plain,(
  ( ! [X350,X352,X349,X351,X347,X353,X348] : (~p(s(bool,happ(s(fun(X347,bool),happ(s(fun(X350,fun(X347,bool)),'_dst_x0x1_2'(s(fun(X348,fun(X347,bool)),X349),s(fun(X350,X348),X351))),s(X350,X352))),s(X347,X353)))) | p(s(bool,'|dollar|exists_1'(s(fun(X347,bool),happ(s(fun(X348,fun(X347,bool)),X349),s(X348,happ(s(fun(X350,X348),X351),s(X350,X352))))))))) )),
  inference(superposition,[],[f4681,f369])).
fof(f11353,plain,(
  ( ! [X621,X617,X619,X615,X620,X622,X616,X618] : (p(s(bool,happ(s(fun('type.pair.prod'(X615,X616),bool),happ(s(fun(X618,fun('type.pair.prod'(X615,X616),bool)),happ(s(fun(X617,fun(X618,fun('type.pair.prod'(X615,X616),bool))),X619),s(X617,X620))),s(X618,X621))),s('type.pair.prod'(X615,X616),X622)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X615,X616),bool),happ(s(fun(X618,fun('type.pair.prod'(X615,X616),bool)),happ(s(fun(X617,fun(X618,fun('type.pair.prod'(X615,X616),bool))),X619),s(X617,X620))),s(X618,X621))),s('type.pair.prod'(X615,X616),'const.pair.,_2'(s(X615,sK7(happ(s(fun('type.pair.prod'(X617,X618),fun('type.pair.prod'(X615,X616),bool)),sK19(X619,X618,X617,fun('type.pair.prod'(X615,X616),bool))),s('type.pair.prod'(X617,X618),'const.pair.,_2'(s(X617,X620),s(X618,X621)))),X616,X615)),s(X616,sK8(happ(s(fun('type.pair.prod'(X617,X618),fun('type.pair.prod'(X615,X616),bool)),sK19(X619,X618,X617,fun('type.pair.prod'(X615,X616),bool))),s('type.pair.prod'(X617,X618),'const.pair.,_2'(s(X617,X620),s(X618,X621)))),X616,X615)))))))) )),
  inference(forward_demodulation,[],[f11147,f367])).
fof(f11147,plain,(
  ( ! [X621,X617,X619,X615,X620,X622,X616,X618] : (~p(s(bool,happ(s(fun('type.pair.prod'(X615,X616),bool),happ(s(fun(X618,fun('type.pair.prod'(X615,X616),bool)),happ(s(fun(X617,fun(X618,fun('type.pair.prod'(X615,X616),bool))),X619),s(X617,X620))),s(X618,X621))),s('type.pair.prod'(X615,X616),'const.pair.,_2'(s(X615,sK7(happ(s(fun('type.pair.prod'(X617,X618),fun('type.pair.prod'(X615,X616),bool)),sK19(X619,X618,X617,fun('type.pair.prod'(X615,X616),bool))),s('type.pair.prod'(X617,X618),'const.pair.,_2'(s(X617,X620),s(X618,X621)))),X616,X615)),s(X616,sK8(happ(s(fun('type.pair.prod'(X617,X618),fun('type.pair.prod'(X615,X616),bool)),sK19(X619,X618,X617,fun('type.pair.prod'(X615,X616),bool))),s('type.pair.prod'(X617,X618),'const.pair.,_2'(s(X617,X620),s(X618,X621)))),X616,X615))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X615,X616),bool),happ(s(fun('type.pair.prod'(X617,X618),fun('type.pair.prod'(X615,X616),bool)),sK19(X619,X618,X617,fun('type.pair.prod'(X615,X616),bool))),s('type.pair.prod'(X617,X618),'const.pair.,_2'(s(X617,X620),s(X618,X621))))),s('type.pair.prod'(X615,X616),X622))))) )),
  inference(superposition,[],[f339,f367])).
fof(f9525,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,sK7(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18)),s(X19,sK8(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24))),s('type.pair.prod'(X18,X19),X25))))) )),
  inference(superposition,[],[f339,f369])).
fof(f9524,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X14,fun('type.pair.prod'(X10,X11),bool)),'_dst_x0x1_2'(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s('type.pair.prod'(X10,X11),'const.pair.,_2'(s(X10,sK7(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10)),s(X11,sK8(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16))))),s('type.pair.prod'(X10,X11),X17))))) )),
  inference(superposition,[],[f339,f369])).
fof(f2841,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,'|dollar|exists_1'(s(fun(X138,bool),happ(s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num',X141)))))) | ~p(s(bool,happ(s(fun(X138,bool),happ(s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num',X141))),s(X138,sK4(X142,'const.ind_type.FCONS_3'(s(fun(X138,bool),X139),s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num','const.num.SUC_1'(s('type.num.num',X141)))),X138)))))) )),
  inference(forward_demodulation,[],[f2802,f324])).
fof(f2802,plain,(
  ( ! [X142,X140,X138,X141,X139] : (~p(s(bool,happ(s(fun(X138,bool),happ(s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num',X141))),s(X138,sK4(X142,'const.ind_type.FCONS_3'(s(fun(X138,bool),X139),s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num','const.num.SUC_1'(s('type.num.num',X141)))),X138))))) | p(s(bool,'|dollar|exists_1'(s(fun(X138,bool),'const.ind_type.FCONS_3'(s(fun(X138,bool),X139),s(fun('type.num.num',fun(X138,bool)),X140),s('type.num.num','const.num.SUC_1'(s('type.num.num',X141))))))))) )),
  inference(superposition,[],[f319,f324])).
fof(f4155,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f4134,f324])).
fof(f4134,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f324])).
fof(f9587,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),X47)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)))))))) )),
  inference(forward_demodulation,[],[f9529,f324])).
fof(f9529,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))),s('type.pair.prod'(X42,X43),X47))))) )),
  inference(superposition,[],[f339,f324])).
fof(f7098,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f4681,f369])).
fof(f7418,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f6722,f369])).
fof(f10752,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6866,f369])).
fof(f18620,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X78,bool),'_dst_x0x1_2'(s(fun(X76,bool),X77),s(fun(X78,X76),X79))),s(X78,X80))))) | s(bool,'T_0') = s(bool,happ(s(fun(X76,bool),X77),s(X76,happ(s(fun(X78,X76),X79),s(X78,X80)))))) ) | $spl10),
  inference(superposition,[],[f11678,f369])).
fof(f6588,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36)))))) )),
  inference(forward_demodulation,[],[f6567,f368])).
fof(f6567,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36))))) | ~p(s(bool,happ(s(fun('type.list.list'(X36),bool),happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39))),s('type.list.list'(X36),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f368])).
fof(f9582,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),X5)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f9522,f368])).
fof(f9522,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4))),s('type.pair.prod'(X0,X1),X5))))) )),
  inference(superposition,[],[f339,f368])).
fof(f1533,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1514,f338])).
fof(f1514,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f319,f338])).
fof(f4153,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f4132,f338])).
fof(f4132,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f338])).
fof(f9585,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)))))))) )),
  inference(forward_demodulation,[],[f9527,f338])).
fof(f9527,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36))))) )),
  inference(superposition,[],[f339,f338])).
fof(f2017,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.min.@_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)))))))) )),
  inference(superposition,[],[f317,f357])).
fof(f349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f398,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X31,X44,X9,X45] : ! [X46,X37,X47,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X45),X10),s(fun(X45,X9),X47)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X31,X44),X37),s(fun(X44,X31),X46))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X31),fun(X45,X44)),'_dst_x00x11_2'(s(fun(X31,X44),X37),s(fun(fun(X9,X31),fun(X45,X31)),'_dstsel_0_x01_1'(s(fun(X45,X9),X47))))),s(fun(fun(X45,X44),fun(X9,X31)),'_dst_x00x11_2'(s(fun(X44,X31),X46),s(fun(fun(X45,X44),fun(X9,X44)),'_dstsel_0_x01_1'(s(fun(X9,X45),X10)))))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.ISO_FUN')).
fof(f5093,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5062,f262])).
fof(f5062,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f354,f445])).
fof(f5256,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5225,f262])).
fof(f5225,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f355,f445])).
fof(f9686,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f9699,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f9968,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f9981,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f9690,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f9972,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f9689,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f9971,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f9677,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f9959,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f9676,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f9958,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f1484,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f338,f445])).
fof(f1486,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f338,f445])).
fof(f2037,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28))))) )),
  inference(forward_demodulation,[],[f2018,f338])).
fof(f2018,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f317,f338])).
fof(f7421,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f6722,f338])).
fof(f21177,plain,(
  ( ! [X4,X2] : (s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(global_subsumption,[],[f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1497,f1374,f1376,f1378,f2515,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4364,f4365,f4516,f4517,f4605,f4611,f4617,f4686,f4813,f4840,f4916,f4915,f4909,f5038,f5122,f5121,f5380,f5391,f5396,f5400,f5787,f6722,f6866,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6980,f6981,f6992,f7142,f7147,f7159,f7163,f7165,f7172,f7174,f7235,f7240,f7252,f7256,f7258,f7265,f7267,f7420,f7422,f7424,f7425,f7426,f7427,f7428,f7685,f7694,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8060,f8061,f8062,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8096,f8097,f8108,f8109,f8110,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9611,f9736,f9737,f9826,f9827,f10018,f10064,f10067,f10069,f10071,f10072,f10073,f10074,f10075,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10754,f10756,f10758,f10759,f10760,f10761,f10762,f11056,f11347,f11348,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f5378,f5377,f5376,f5375,f5373,f5371,f5370,f4812,f4854,f5379,f7421,f5372,f2037,f1486,f1484,f7242,f7244,f9677,f9959,f7149,f7151,f9676,f9958,f7243,f7245,f9690,f9972,f7150,f7152,f9689,f9971,f9981,f9968,f9699,f9686,f5256,f5222,f5212,f5093,f5059,f5049,f356,f355,f354,f350,f349,f398,f2017,f9585,f1533,f4153,f9582,f10752,f7418,f7135,f7098,f6588,f9587,f2841,f4155,f11353,f9525,f9524,f11340,f7239,f7146,f11070,f7234,f7236,f7141,f7143,f11113,f11112,f11109,f11073,f11072,f11042,f11041,f11071,f9534,f9604,f9591,f9590,f9589,f9588,f9586,f9584,f9583,f339,f9552,f9551,f9537,f9536,f7894,f7695,f7664,f4601,f4160,f4159,f4158,f4157,f4156,f4154,f4152,f4151,f420,f311,f7679,f11337,f7677,f7674,f7584,f7583,f7662,f4599,f1391,f1390,f1389,f1388,f1387,f1385,f1386,f319,f1392,f7099,f12005,f12026,f12023,f12029,f12027,f11995,f11994,f9581,f7193,f7190,f12036,f12033,f12032,f12031,f12030,f12028,f12025,f12024,f11991,f11990,f11989,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11981,f11980,f11979,f11978,f10753,f7663,f7419,f7229,f7209,f7207,f7202,f7201,f7199,f7196,f7195,f7192,f7191,f7115,f7113,f7108,f7107,f7105,f7102,f7101,f348,f345,f12035,f11338,f7885,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7867,f7838,f7829,f7828,f7688,f7683,f7682,f7681,f7680,f7678,f7676,f7675,f7593,f4833,f4754,f4681,f4598,f2032,f347,f346,f321,f320,f317,f10757,f7423,f5374,f2766,f2764,f7893,f7841,f7840,f4376,f4377,f11446,f11096,f11095,f11091,f11090,f11088,f11040,f10793,f10464,f10463,f10462,f10461,f10460,f10458,f10456,f10453,f9837,f9670,f9550,f9535,f8112,f8111,f8095,f8064,f8063,f7884,f7839,f7218,f7217,f7216,f7215,f7211,f7210,f7208,f7198,f7197,f7124,f7123,f7122,f7121,f7117,f7116,f7114,f7104,f7103,f6979,f6876,f6782,f6781,f6779,f6777,f6776,f6775,f6774,f6732,f6639,f6638,f6636,f6634,f6633,f6632,f6631,f5404,f5392,f5390,f4585,f4584,f4580,f4579,f4577,f4567,f4566,f4515,f4509,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4375,f4363,f4360,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f2792,f2521,f2520,f2483,f2482,f2480,f2478,f2477,f2476,f2475,f2474,f2473,f2427,f1488,f909,f743,f686,f633,f540,f513,f472,f276,f275,f264,f19241,f19242,f19243,f19244,f19245,f19246,f19247,f19248,f19249,f19250,f19251,f19252,f19253,f19254,f19287,f19288,f19257,f19258,f19289,f19290,f19291,f19292,f19293,f19294,f19295,f19296,f19297,f19268,f19299,f19300,f19632,f19633,f19634,f19635,f19636,f19637,f19638,f19639,f19640,f19641,f19642,f19643,f19644,f19645,f19678,f19679,f19648,f19649,f19680,f19681,f19682,f19683,f19684,f19685,f19686,f19687,f19688,f19659,f19690,f19691,f19750,f19753,f19756,f19758,f19759,f19760,f19761,f19762,f20179,f20181,f11707,f5776,f4609,f8094,f8059,f4715,f4578,f4576,f4571,f4570,f4568,f4565,f4564,f1515,f4529,f6563,f6566,f2803,f4531,f11124,f11128,f7145,f7238,f4685,f4684,f4600,f4536,f4535,f4534,f4533,f4532,f4530,f4528,f4527,f930,f834,f833,f832,f831,f830,f829,f318,f20180,f10755,f9627])).
fof(f9627,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f9613,f262])).
fof(f9613,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f9611,f445])).
fof(f10755,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f6866,f338])).
fof(f18623,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X89),X90))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X89),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X89),X90)))) ) | $spl10),
  inference(superposition,[],[f11678,f338])).
fof(f20180,plain,(
  ( ! [X76,X75] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),X76)))) | ~p(s(bool,sK7('const.pair.FST_0',X75,bool)))) )),
  inference(forward_demodulation,[],[f17516,f338])).
fof(f17516,plain,(
  ( ! [X76,X75] : (~p(s(bool,sK7('const.pair.FST_0',X75,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X75),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X75),X76))))) )),
  inference(forward_demodulation,[],[f17480,f341])).
fof(f17480,plain,(
  ( ! [X76,X75] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),'const.pair.,_2'(s(bool,sK7('const.pair.FST_0',X75,bool)),s(X75,sK8('const.pair.FST_0',X75,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X75),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X75),X76))))) )),
  inference(superposition,[],[f339,f338])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f318,f309])).
fof(f830,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X6) = s(X3,sK4(X6,'const.ind_type.FCONS_3'(s(fun(X3,bool),X4),s(fun('type.num.num',fun(X3,bool)),X5),s('type.num.num','const.num.|48|_0')),X3))) )),
  inference(superposition,[],[f318,f316])).
fof(f831,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X9)))) | s(X7,X10) = s(X7,sK4(X10,'const.sum.OUTR_1'(s('type.sum.sum'(X8,fun(X7,bool)),'const.sum.INR_1'(s(fun(X7,bool),X9)))),X7))) )),
  inference(superposition,[],[f318,f337])).
fof(f832,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f318,f341])).
fof(f833,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | s(X16,X19) = s(X16,sK4(X19,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X16,bool),X17),'const.sum.INL_1'(s(fun(X16,bool),X18)))),X16))) )),
  inference(superposition,[],[f318,f336])).
fof(f834,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X23)))) | s(X20,X24) = s(X20,sK4(X24,'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))) )),
  inference(superposition,[],[f318,f342])).
fof(f930,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK18(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f318,f357])).
fof(f4527,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f836,f309])).
fof(f4528,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f836,f357])).
fof(f4530,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f836,f316])).
fof(f4532,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTR_1'(s('type.sum.sum'(X22,fun(X21,bool)),'const.sum.INR_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f836,f337])).
fof(f4533,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f836,f341])).
fof(f4534,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X32))) | s(X30,X33) = s(X30,sK4(X33,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X30,bool),X31),'const.sum.INL_1'(s(fun(X30,bool),X32)))),X30))) )),
  inference(superposition,[],[f836,f336])).
fof(f4535,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X37))) | s(X34,X38) = s(X34,sK4(X38,'const.pair.SND_1'(s('type.pair.prod'(X35,fun(X34,bool)),'const.pair.,_2'(s(X35,X36),s(fun(X34,bool),X37)))),X34))) )),
  inference(superposition,[],[f836,f342])).
fof(f4536,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f836,f836])).
fof(f4600,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f318,f836])).
fof(f4684,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4632,f262])).
fof(f4632,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f333,f836])).
fof(f4685,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4633,f262])).
fof(f4633,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f418,f836])).
fof(f7238,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X363] : (p(s(bool,'|dollar|exists_1'(s(fun(X362,bool),happ(s(fun(X364,fun(X362,bool)),X365),s(X364,happ(s(fun(X363,X364),X366),s(X363,X367)))))))) | s(X362,X368) = s(X362,sK4(X368,happ(s(fun(X363,fun(X362,bool)),'_dst_x0x1_2'(s(fun(X364,fun(X362,bool)),X365),s(fun(X363,X364),X366))),s(X363,X367)),X362))) )),
  inference(superposition,[],[f318,f369])).
fof(f7145,plain,(
  ( ! [X377,X379,X380,X376,X378,X375,X381] : (p(s(bool,'|dollar|exists_1'(s(fun(X375,bool),happ(s(fun(X378,fun(X375,bool)),'_dst_x0x1_2'(s(fun(X376,fun(X375,bool)),X377),s(fun(X378,X376),X379))),s(X378,X380)))))) | s(X375,X381) = s(X375,sK4(X381,happ(s(fun(X376,fun(X375,bool)),X377),s(X376,happ(s(fun(X378,X376),X379),s(X378,X380)))),X375))) )),
  inference(superposition,[],[f318,f369])).
fof(f11128,plain,(
  ( ! [X447,X453,X449,X451,X452,X448,X450] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X447,bool),happ(s(fun(X449,fun(X447,bool)),happ(s(fun(X448,fun(X449,fun(X447,bool))),X450),s(X448,X451))),s(X449,X452))))) | s(X447,X453) = s(X447,sK4(X453,happ(s(fun('type.pair.prod'(X448,X449),fun(X447,bool)),sK19(X450,X449,X448,fun(X447,bool))),s('type.pair.prod'(X448,X449),'const.pair.,_2'(s(X448,X451),s(X449,X452)))),X447))) )),
  inference(superposition,[],[f836,f367])).
fof(f11124,plain,(
  ( ! [X420,X422,X418,X421,X423,X417,X419] : (p(s(bool,'|dollar|exists_1'(s(fun(X417,bool),happ(s(fun(X419,fun(X417,bool)),happ(s(fun(X418,fun(X419,fun(X417,bool))),X420),s(X418,X421))),s(X419,X422)))))) | s(X417,X423) = s(X417,sK4(X423,happ(s(fun('type.pair.prod'(X418,X419),fun(X417,bool)),sK19(X420,X419,X418,fun(X417,bool))),s('type.pair.prod'(X418,X419),'const.pair.,_2'(s(X418,X421),s(X419,X422)))),X417))) )),
  inference(superposition,[],[f318,f367])).
fof(f4531,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))) | s(X16,X20) = s(X16,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))) )),
  inference(superposition,[],[f836,f324])).
fof(f2803,plain,(
  ( ! [X146,X144,X147,X145,X143] : (p(s(bool,'|dollar|exists_1'(s(fun(X143,bool),happ(s(fun('type.num.num',fun(X143,bool)),X145),s('type.num.num',X146)))))) | s(X143,X147) = s(X143,sK4(X147,'const.ind_type.FCONS_3'(s(fun(X143,bool),X144),s(fun('type.num.num',fun(X143,bool)),X145),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))),X143))) )),
  inference(superposition,[],[f318,f324])).
fof(f6566,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun(X31,X32),X33)))))) | s(X31,X35) = s(X31,sK4(X35,happ(s(fun(fun(X32,bool),fun(X31,bool)),'_dstsel_0_x01_1'(s(fun(X31,X32),X33))),s(fun(X32,bool),X34)),X31))) )),
  inference(superposition,[],[f318,f368])).
fof(f6563,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0x1_2'(s(fun(X17,bool),X19),s(fun(X16,X17),X18))))) | s(X16,X20) = s(X16,sK4(X20,happ(s(fun(fun(X17,bool),fun(X16,bool)),'_dstsel_0_x01_1'(s(fun(X16,X17),X18))),s(fun(X17,bool),X19)),X16))) )),
  inference(superposition,[],[f836,f368])).
fof(f4529,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f836,f338])).
fof(f1515,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f318,f338])).
fof(f4564,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f271,f836])).
fof(f4565,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f272,f836])).
fof(f4568,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f330,f836])).
fof(f4570,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f333,f836])).
fof(f4571,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f418,f836])).
fof(f4576,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f447,f836])).
fof(f4578,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f494,f836])).
fof(f4715,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(fun(X57,X58),X59) != s(fun(X57,X58),X61) | s(X58,happ(s(fun(X57,X58),X61),s(X57,X62))) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f352,f836])).
fof(f8059,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f8094,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f4609,plain,(
  ( ! [X196,X197,X195] : (s('type.ast.lit',X195) != s('type.ast.lit',X197) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X197))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK4(X195,X196,'type.ast.lit')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X196)))) )),
  inference(superposition,[],[f335,f836])).
fof(f5776,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.list.list'('type.exhSem.v'),X48) != s('type.list.list'('type.exhSem.v'),X50) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X52),s('type.list.list'('type.exhSem.v'),X50))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X51),s('type.list.list'('type.exhSem.v'),sK4(X48,X49,'type.list.list'('type.exhSem.v'))))) | s('type.num.num',X51) != s('type.num.num',X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.exhSem.v'),bool),X49)))) )),
  inference(superposition,[],[f396,f836])).
fof(f11707,plain,(
  ( ! [X101,X99,X97,X102,X100,X98,X96] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X96),s('type.list.list'('type.string.char'),X98),s('type.exhLang.exp',X99))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X100),s('type.list.list'('type.string.char'),X101),s('type.exhLang.exp',X102))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X100) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK4(X96,X97,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),X97)))) )),
  inference(superposition,[],[f399,f836])).
fof(f9666,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f9653,f264])).
fof(f9653,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f9630,f2865])).
fof(f9661,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9645,f264])).
fof(f9645,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(superposition,[],[f9630,f270])).
fof(f9630,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(subsumption_resolution,[],[f9629,f262])).
fof(f9629,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f9618,f2199])).
fof(f9618,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(superposition,[],[f9611,f2522])).
fof(f19763,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f11446,f2865])).
fof(f19762,plain,(
  ( ! [X180,X181,X179] : (s(bool,X179) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X181),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X179,X180,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180)))) ) | $spl26),
  inference(superposition,[],[f11446,f836])).
fof(f19761,plain,(
  ( ! [X177,X175,X178,X176] : (s(bool,X177) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X178),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X175,bool),'const.pair.,_2'(s(X175,X176),s(bool,X177)))))) ) | $spl26),
  inference(superposition,[],[f11446,f342])).
fof(f19760,plain,(
  ( ! [X173,X174,X172] : (s(bool,X173) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X174),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X172),'const.sum.INL_1'(s(bool,X173)))))) ) | $spl26),
  inference(superposition,[],[f11446,f336])).
fof(f19759,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,X169) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X171),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X168),'const.pair.,_2'(s(bool,X169),s(X168,X170)))))) ) | $spl26),
  inference(superposition,[],[f11446,f341])).
fof(f19758,plain,(
  ( ! [X167,X165,X166] : (s(bool,X166) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X167),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X165,bool),'const.sum.INR_1'(s(bool,X166)))))) ) | $spl26),
  inference(superposition,[],[f11446,f337])).
fof(f19756,plain,(
  ( ! [X158,X159,X160] : (s(bool,X158) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X160),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X158),s(fun('type.num.num',bool),X159),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f11446,f316])).
fof(f19753,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X147) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X148),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X145),bool),sK18(X146,X147,X145,bool)),s('type.list.list'(X145),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f11446,f357])).
fof(f19750,plain,(
  ( ! [X132,X131] : (s(bool,X131) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X132),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X131)))) ) | $spl26),
  inference(superposition,[],[f11446,f309])).
fof(f18775,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f18758])).
fof(f18758,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f18664,f445])).
fof(f18631,plain,(
  ( ! [X111,X112] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X111))) | s(bool,'T_0') = s(bool,sK4(X111,X112,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) ) | $spl10),
  inference(superposition,[],[f11678,f836])).
fof(f18664,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f18567,f9670])).
fof(f18567,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('const.pair.FST_0',X0,bool))))) ) | ($spl10 | $spl26)),
  inference(unit_resulting_resolution,[],[f9827,f11678])).
fof(f264,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.exhSem.v_|49||49||split|4',aHL_FALSITY)).
fof(f9837,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(subsumption_resolution,[],[f9783,f264])).
fof(f9783,plain,(
  ( ! [X41,X40] : (p(s(bool,'F_0')) | s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(superposition,[],[f2482,f9670])).
fof(f10453,plain,(
  ( ! [X107,X106] : (s(bool,X106) != s(bool,sK7('const.pair.FST_0',X107,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X106)))) ) | $spl26),
  inference(superposition,[],[f9837,f309])).
fof(f10456,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X122) != s(bool,sK7('const.pair.FST_0',X123,bool)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X120),bool),sK18(X121,X122,X120,bool)),s('type.list.list'(X120),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f9837,f357])).
fof(f10458,plain,(
  ( ! [X127,X128,X129] : (s(bool,X127) != s(bool,sK7('const.pair.FST_0',X129,bool)) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f9837,f316])).
fof(f10460,plain,(
  ( ! [X134,X136,X135] : (s(bool,X135) != s(bool,sK7('const.pair.FST_0',X136,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X134,bool),'const.sum.INR_1'(s(bool,X135)))))) ) | $spl26),
  inference(superposition,[],[f9837,f337])).
fof(f10461,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X138) != s(bool,sK7('const.pair.FST_0',X140,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X138),s(X137,X139)))))) ) | $spl26),
  inference(superposition,[],[f9837,f341])).
fof(f10462,plain,(
  ( ! [X142,X143,X141] : (s(bool,X142) != s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X141),'const.sum.INL_1'(s(bool,X142)))))) ) | $spl26),
  inference(superposition,[],[f9837,f336])).
fof(f10463,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X146) != s(bool,sK7('const.pair.FST_0',X147,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X144,bool),'const.pair.,_2'(s(X144,X145),s(bool,X146)))))) ) | $spl26),
  inference(superposition,[],[f9837,f342])).
fof(f10464,plain,(
  ( ! [X149,X150,X148] : (s(bool,X148) != s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'F_0') = s(bool,sK4(X148,X149,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X149)))) ) | $spl26),
  inference(superposition,[],[f9837,f836])).
fof(f10793,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f10704,f9670])).
fof(f10704,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f9827,f6866])).
fof(f11446,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f11429])).
fof(f11429,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f10793,f445])).
fof(f11875,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK18(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK18(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f11830,f262])).
fof(f11830,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK18(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK18(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(superposition,[],[f4854,f5395])).
fof(f11873,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f11807,f262])).
fof(f11807,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f447,f5395])).
fof(f11706,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.exhLang.exp',X92))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X93),s('type.list.list'('type.string.char'),X94),s('type.exhLang.exp',X95))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X93) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.pair.SND_1'(s('type.pair.prod'(X88,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.,_2'(s(X88,X89),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X90)))))) )),
  inference(superposition,[],[f399,f342])).
fof(f11705,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X82),s('type.list.list'('type.string.char'),X83),s('type.exhLang.exp',X84))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X85),s('type.list.list'('type.string.char'),X86),s('type.exhLang.exp',X87))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X85) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X81),'const.sum.INL_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X82)))))) )),
  inference(superposition,[],[f399,f336])).
fof(f11704,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75,X73] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X74),s('type.list.list'('type.string.char'),X76),s('type.exhLang.exp',X77))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X78),s('type.list.list'('type.string.char'),X79),s('type.exhLang.exp',X80))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X78) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X73),'const.pair.,_2'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X74),s(X73,X75)))))) )),
  inference(superposition,[],[f399,f341])).
fof(f11703,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X67),s('type.list.list'('type.string.char'),X68),s('type.exhLang.exp',X69))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X70),s('type.list.list'('type.string.char'),X71),s('type.exhLang.exp',X72))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X70) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.sum.OUTR_1'(s('type.sum.sum'(X66,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.sum.INR_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X67)))))) )),
  inference(superposition,[],[f399,f337])).
fof(f11701,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X51),s('type.list.list'('type.string.char'),X53),s('type.exhLang.exp',X54))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X55),s('type.list.list'('type.string.char'),X56),s('type.exhLang.exp',X57))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X55) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.ind_type.FCONS_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X51),s(fun('type.num.num','type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X52),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f399,f316])).
fof(f11699,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X38),s('type.list.list'('type.string.char'),X39),s('type.exhLang.exp',X40))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X34),X35))),s('type.list.list'('type.string.char'),X36),s('type.exhLang.exp',X37))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X38) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X34),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X34),X35)))) )),
  inference(superposition,[],[f399,f338])).
fof(f11698,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X28),s('type.list.list'('type.string.char'),X29),s('type.exhLang.exp',X30))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X31),s('type.list.list'('type.string.char'),X32),s('type.exhLang.exp',X33))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X31) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.list.list'(X26),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK18(X27,X28,X26,'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),s('type.list.list'(X26),'const.list.NIL_0')))) )),
  inference(superposition,[],[f399,f357])).
fof(f11695,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'_i_0_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0)))) )),
  inference(superposition,[],[f399,f309])).
fof(f11678,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f11583,f264])).
fof(f11583,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f271,f2975])).
fof(f11184,plain,(
  ( ! [X907,X905,X899,X903,X901,X906,X904,X898,X902,X900] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X903),s('type.list.list'('type.string.char'),X904),s('type.exhLang.exp',X905))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X906),s('type.list.list'('type.string.char'),X907),s('type.exhLang.exp',happ(s(fun(X899,'type.exhLang.exp'),happ(s(fun(X898,fun(X899,'type.exhLang.exp')),X900),s(X898,X901))),s(X899,X902))))) | s('type.exhLang.exp',X905) = s('type.exhLang.exp',happ(s(fun('type.pair.prod'(X898,X899),'type.exhLang.exp'),sK19(X900,X899,X898,'type.exhLang.exp')),s('type.pair.prod'(X898,X899),'const.pair.,_2'(s(X898,X901),s(X899,X902)))))) )),
  inference(superposition,[],[f401,f367])).
fof(f11159,plain,(
  ( ! [X697,X699,X700,X702,X698,X701] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X702))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X698,'type.ast.lit'),happ(s(fun(X697,fun(X698,'type.ast.lit')),X699),s(X697,X700))),s(X698,X701))))) | s('type.ast.lit',X702) = s('type.ast.lit',happ(s(fun('type.pair.prod'(X697,X698),'type.ast.lit'),sK19(X699,X698,X697,'type.ast.lit')),s('type.pair.prod'(X697,X698),'const.pair.,_2'(s(X697,X700),s(X698,X701)))))) )),
  inference(superposition,[],[f334,f367])).
fof(f11153,plain,(
  ( ! [X660,X656,X658,X654,X657,X659,X653,X655] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X659),s('type.list.list'('type.exhSem.v'),X660))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X654,'type.num.num'),happ(s(fun(X653,fun(X654,'type.num.num')),X655),s(X653,X656))),s(X654,X657))),s('type.list.list'('type.exhSem.v'),X658))) | s('type.num.num',X659) = s('type.num.num',happ(s(fun('type.pair.prod'(X653,X654),'type.num.num'),sK19(X655,X654,X653,'type.num.num')),s('type.pair.prod'(X653,X654),'const.pair.,_2'(s(X653,X656),s(X654,X657)))))) )),
  inference(superposition,[],[f394,f367])).
fof(f11348,plain,(
  ( ! [X540,X536,X538,X534,X537,X539,X533,X535] : (s(fun(X533,X534),X540) = s(fun(X533,X534),happ(s(fun(X536,fun(X533,X534)),happ(s(fun(X535,fun(X536,fun(X533,X534))),X537),s(X535,X538))),s(X536,X539))) | s(X534,happ(s(fun(X533,X534),X540),s(X533,sK15(happ(s(fun('type.pair.prod'(X535,X536),fun(X533,X534)),sK19(X537,X536,X535,fun(X533,X534))),s('type.pair.prod'(X535,X536),'const.pair.,_2'(s(X535,X538),s(X536,X539)))),X540,X533,X534)))) != s(X534,happ(s(fun(X533,X534),happ(s(fun(X536,fun(X533,X534)),happ(s(fun(X535,fun(X536,fun(X533,X534))),X537),s(X535,X538))),s(X536,X539))),s(X533,sK15(happ(s(fun('type.pair.prod'(X535,X536),fun(X533,X534)),sK19(X537,X536,X535,fun(X533,X534))),s('type.pair.prod'(X535,X536),'const.pair.,_2'(s(X535,X538),s(X536,X539)))),X540,X533,X534))))) )),
  inference(forward_demodulation,[],[f11138,f367])).
fof(f11138,plain,(
  ( ! [X540,X536,X538,X534,X537,X539,X533,X535] : (s(X534,happ(s(fun(X533,X534),X540),s(X533,sK15(happ(s(fun('type.pair.prod'(X535,X536),fun(X533,X534)),sK19(X537,X536,X535,fun(X533,X534))),s('type.pair.prod'(X535,X536),'const.pair.,_2'(s(X535,X538),s(X536,X539)))),X540,X533,X534)))) != s(X534,happ(s(fun(X533,X534),happ(s(fun(X536,fun(X533,X534)),happ(s(fun(X535,fun(X536,fun(X533,X534))),X537),s(X535,X538))),s(X536,X539))),s(X533,sK15(happ(s(fun('type.pair.prod'(X535,X536),fun(X533,X534)),sK19(X537,X536,X535,fun(X533,X534))),s('type.pair.prod'(X535,X536),'const.pair.,_2'(s(X535,X538),s(X536,X539)))),X540,X533,X534)))) | s(fun(X533,X534),X540) = s(fun(X533,X534),happ(s(fun('type.pair.prod'(X535,X536),fun(X533,X534)),sK19(X537,X536,X535,fun(X533,X534))),s('type.pair.prod'(X535,X536),'const.pair.,_2'(s(X535,X538),s(X536,X539)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f11347,plain,(
  ( ! [X532,X528,X530,X526,X529,X531,X525,X527] : (s(fun(X525,X526),X532) = s(fun(X525,X526),happ(s(fun(X528,fun(X525,X526)),happ(s(fun(X527,fun(X528,fun(X525,X526))),X529),s(X527,X530))),s(X528,X531))) | s(X526,happ(s(fun(X525,X526),X532),s(X525,sK15(X532,happ(s(fun('type.pair.prod'(X527,X528),fun(X525,X526)),sK19(X529,X528,X527,fun(X525,X526))),s('type.pair.prod'(X527,X528),'const.pair.,_2'(s(X527,X530),s(X528,X531)))),X525,X526)))) != s(X526,happ(s(fun(X525,X526),happ(s(fun(X528,fun(X525,X526)),happ(s(fun(X527,fun(X528,fun(X525,X526))),X529),s(X527,X530))),s(X528,X531))),s(X525,sK15(X532,happ(s(fun('type.pair.prod'(X527,X528),fun(X525,X526)),sK19(X529,X528,X527,fun(X525,X526))),s('type.pair.prod'(X527,X528),'const.pair.,_2'(s(X527,X530),s(X528,X531)))),X525,X526))))) )),
  inference(forward_demodulation,[],[f11137,f367])).
fof(f11137,plain,(
  ( ! [X532,X528,X530,X526,X529,X531,X525,X527] : (s(X526,happ(s(fun(X525,X526),X532),s(X525,sK15(X532,happ(s(fun('type.pair.prod'(X527,X528),fun(X525,X526)),sK19(X529,X528,X527,fun(X525,X526))),s('type.pair.prod'(X527,X528),'const.pair.,_2'(s(X527,X530),s(X528,X531)))),X525,X526)))) != s(X526,happ(s(fun(X525,X526),happ(s(fun(X528,fun(X525,X526)),happ(s(fun(X527,fun(X528,fun(X525,X526))),X529),s(X527,X530))),s(X528,X531))),s(X525,sK15(X532,happ(s(fun('type.pair.prod'(X527,X528),fun(X525,X526)),sK19(X529,X528,X527,fun(X525,X526))),s('type.pair.prod'(X527,X528),'const.pair.,_2'(s(X527,X530),s(X528,X531)))),X525,X526)))) | s(fun(X525,X526),X532) = s(fun(X525,X526),happ(s(fun('type.pair.prod'(X527,X528),fun(X525,X526)),sK19(X529,X528,X527,fun(X525,X526))),s('type.pair.prod'(X527,X528),'const.pair.,_2'(s(X527,X530),s(X528,X531)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f11056,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))) = s(X6,happ(s(fun('type.pair.prod'(X1,X2),X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(superposition,[],[f369,f367])).
fof(f10762,plain,(
  ( ! [X154,X155,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X155),'const.sum.INL_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,sK4(X153,X154,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154)))) )),
  inference(superposition,[],[f6866,f836])).
fof(f10761,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X151))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X149,bool),'const.pair.,_2'(s(X149,X150),s(bool,X151)))))) )),
  inference(superposition,[],[f6866,f342])).
fof(f10760,plain,(
  ( ! [X146,X147,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X148),'const.sum.INL_1'(s(bool,X147))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X146),'const.sum.INL_1'(s(bool,X147)))))) )),
  inference(superposition,[],[f6866,f336])).
fof(f10759,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X145),'const.sum.INL_1'(s(bool,X143))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,X143),s(X142,X144)))))) )),
  inference(superposition,[],[f6866,f341])).
fof(f10758,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X141),'const.sum.INL_1'(s(bool,X140))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X139,bool),'const.sum.INR_1'(s(bool,X140)))))) )),
  inference(superposition,[],[f6866,f337])).
fof(f10756,plain,(
  ( ! [X134,X132,X133] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X134),'const.sum.INL_1'(s(bool,X132))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X132),s(fun('type.num.num',bool),X133),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6866,f316])).
fof(f10754,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X125),bool),sK18(X126,X127,X125,bool)),s('type.list.list'(X125),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6866,f357])).
fof(f10551,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X88),s('type.list.list'('type.string.char'),X86),s('type.exhLang.exp',X89))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.exhLang.exp',X92))) | s('type.list.list'('type.string.char'),X91) = s('type.list.list'('type.string.char'),sK4(X86,X87,'type.list.list'('type.string.char'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.string.char'),bool),X87)))) )),
  inference(superposition,[],[f400,f836])).
fof(f10550,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X81),s('type.list.list'('type.string.char'),X80),s('type.exhLang.exp',X82))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X83),s('type.list.list'('type.string.char'),X84),s('type.exhLang.exp',X85))) | s('type.list.list'('type.string.char'),X84) = s('type.list.list'('type.string.char'),'const.pair.SND_1'(s('type.pair.prod'(X78,'type.list.list'('type.string.char')),'const.pair.,_2'(s(X78,X79),s('type.list.list'('type.string.char'),X80)))))) )),
  inference(superposition,[],[f400,f342])).
fof(f10549,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X73),s('type.list.list'('type.string.char'),X72),s('type.exhLang.exp',X74))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X75),s('type.list.list'('type.string.char'),X76),s('type.exhLang.exp',X77))) | s('type.list.list'('type.string.char'),X76) = s('type.list.list'('type.string.char'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.string.char'),X71),'const.sum.INL_1'(s('type.list.list'('type.string.char'),X72)))))) )),
  inference(superposition,[],[f400,f336])).
fof(f10548,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X66),s('type.list.list'('type.string.char'),X64),s('type.exhLang.exp',X67))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X68),s('type.list.list'('type.string.char'),X69),s('type.exhLang.exp',X70))) | s('type.list.list'('type.string.char'),X69) = s('type.list.list'('type.string.char'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.string.char'),X63),'const.pair.,_2'(s('type.list.list'('type.string.char'),X64),s(X63,X65)))))) )),
  inference(superposition,[],[f400,f341])).
fof(f10547,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X58),s('type.list.list'('type.string.char'),X57),s('type.exhLang.exp',X59))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X60),s('type.list.list'('type.string.char'),X61),s('type.exhLang.exp',X62))) | s('type.list.list'('type.string.char'),X61) = s('type.list.list'('type.string.char'),'const.sum.OUTR_1'(s('type.sum.sum'(X56,'type.list.list'('type.string.char')),'const.sum.INR_1'(s('type.list.list'('type.string.char'),X57)))))) )),
  inference(superposition,[],[f400,f337])).
fof(f10545,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X43),s('type.list.list'('type.string.char'),X41),s('type.exhLang.exp',X44))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X45),s('type.list.list'('type.string.char'),X46),s('type.exhLang.exp',X47))) | s('type.list.list'('type.string.char'),X46) = s('type.list.list'('type.string.char'),'const.ind_type.FCONS_3'(s('type.list.list'('type.string.char'),X41),s(fun('type.num.num','type.list.list'('type.string.char')),X42),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f400,f316])).
fof(f10544,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X38),s('type.list.list'('type.string.char'),X39),s('type.exhLang.exp',X40))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X36),s('type.list.list'('type.string.char'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.string.char'),X34),X35))),s('type.exhLang.exp',X37))) | s('type.list.list'('type.string.char'),X39) = s('type.list.list'('type.string.char'),happ(s(fun('type.pair.prod'('type.list.list'('type.string.char'),X34),'type.list.list'('type.string.char')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.string.char'),X34),X35)))) )),
  inference(superposition,[],[f400,f338])).
fof(f10543,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X29),s('type.list.list'('type.string.char'),X28),s('type.exhLang.exp',X30))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X31),s('type.list.list'('type.string.char'),X32),s('type.exhLang.exp',X33))) | s('type.list.list'('type.string.char'),X32) = s('type.list.list'('type.string.char'),happ(s(fun('type.list.list'(X26),'type.list.list'('type.string.char')),sK18(X27,X28,X26,'type.list.list'('type.string.char'))),s('type.list.list'(X26),'const.list.NIL_0')))) )),
  inference(superposition,[],[f400,f357])).
fof(f10075,plain,(
  ( ! [X149,X150,X148] : (s(bool,X148) = s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'T_0') = s(bool,sK4(X148,X149,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X149)))) ) | $spl26),
  inference(superposition,[],[f9826,f836])).
fof(f10074,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X146) = s(bool,sK7('const.pair.FST_0',X147,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X144,bool),'const.pair.,_2'(s(X144,X145),s(bool,X146)))))) ) | $spl26),
  inference(superposition,[],[f9826,f342])).
fof(f10073,plain,(
  ( ! [X142,X143,X141] : (s(bool,X142) = s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X141),'const.sum.INL_1'(s(bool,X142)))))) ) | $spl26),
  inference(superposition,[],[f9826,f336])).
fof(f10072,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X138) = s(bool,sK7('const.pair.FST_0',X140,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X138),s(X137,X139)))))) ) | $spl26),
  inference(superposition,[],[f9826,f341])).
fof(f10071,plain,(
  ( ! [X134,X136,X135] : (s(bool,X135) = s(bool,sK7('const.pair.FST_0',X136,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X134,bool),'const.sum.INR_1'(s(bool,X135)))))) ) | $spl26),
  inference(superposition,[],[f9826,f337])).
fof(f10069,plain,(
  ( ! [X127,X128,X129] : (s(bool,X127) = s(bool,sK7('const.pair.FST_0',X129,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f9826,f316])).
fof(f10067,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X122) = s(bool,sK7('const.pair.FST_0',X123,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X120),bool),sK18(X121,X122,X120,bool)),s('type.list.list'(X120),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f9826,f357])).
fof(f10064,plain,(
  ( ! [X107,X106] : (s(bool,X106) = s(bool,sK7('const.pair.FST_0',X107,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X106)))) ) | $spl26),
  inference(superposition,[],[f9826,f309])).
fof(f10018,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9984,f262])).
fof(f9984,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f350,f445])).
fof(f9827,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(subsumption_resolution,[],[f9760,f264])).
fof(f9760,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(superposition,[],[f271,f9670])).
fof(f9826,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(subsumption_resolution,[],[f9748,f432])).
fof(f9748,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f9670,f445])).
fof(f9737,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9702,f262])).
fof(f9702,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f349,f445])).
fof(f9736,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(subsumption_resolution,[],[f9700,f262])).
fof(f9700,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f349,f440])).
fof(f9611,plain,(
  ( ! [X70,X71] : (~p(s(bool,sK7('const.pair.FST_0',X70,bool))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),X71))))) )),
  inference(forward_demodulation,[],[f9610,f338])).
fof(f9610,plain,(
  ( ! [X70,X71] : (~p(s(bool,sK7('const.pair.FST_0',X70,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X70),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X70),X71))))) )),
  inference(forward_demodulation,[],[f9580,f341])).
fof(f9580,plain,(
  ( ! [X70,X71] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,sK7('const.pair.FST_0',X70,bool)),s(X70,sK8('const.pair.FST_0',X70,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X70),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X70),X71))))) )),
  inference(superposition,[],[f339,f338])).
fof(f9204,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.exhSem.v'),X18))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X14),s('type.list.list'('type.exhSem.v'),X15))))),s('type.list.list'('type.exhSem.v'),X16))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X14))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X15)))))))) )),
  inference(superposition,[],[f394,f329])).
fof(f9198,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X2),s('type.list.list'('type.exhSem.v'),X3))))) = s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X2))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X3)))))))))) )),
  inference(superposition,[],[f329,f269])).
fof(f9240,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X3,'type.list.list'('type.exhSem.v')),X4),s(fun(X2,X3),X5))),s(X2,X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X3,'type.list.list'('type.exhSem.v')),X4),s(X3,happ(s(fun(X2,X3),X5),s(X2,X6)))))))) )),
  inference(forward_demodulation,[],[f9158,f269])).
fof(f9158,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X3,'type.list.list'('type.exhSem.v')),X4),s(fun(X2,X3),X5))),s(X2,X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X3,'type.list.list'('type.exhSem.v')),X4),s(X3,happ(s(fun(X2,X3),X5),s(X2,X6)))))))) )),
  inference(unit_resulting_resolution,[],[f369,f329,f396])).
fof(f9227,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),X2))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))))),s('type.list.list'('type.exhSem.v'),X2)))) )),
  inference(forward_demodulation,[],[f9226,f269])).
fof(f9226,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),X2))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),X2)))) )),
  inference(forward_demodulation,[],[f9165,f342])).
fof(f9165,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),X2))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X3,'type.list.list'('type.exhSem.v')),'const.pair.,_2'(s(X3,X4),s('type.list.list'('type.exhSem.v'),X2)))))))) )),
  inference(unit_resulting_resolution,[],[f342,f329,f396])).
fof(f9225,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3)))))) )),
  inference(forward_demodulation,[],[f9224,f269])).
fof(f9224,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3)))))) )),
  inference(forward_demodulation,[],[f9166,f338])).
fof(f9166,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),'type.list.list'('type.exhSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.exhSem.v'),X2),X3)))))) )),
  inference(unit_resulting_resolution,[],[f338,f329,f396])).
fof(f9223,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num',X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f9222,f269])).
fof(f9222,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num',X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f9167,f324])).
fof(f9167,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num',X3))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.exhSem.v'),X4),s(fun('type.num.num','type.list.list'('type.exhSem.v')),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) )),
  inference(unit_resulting_resolution,[],[f324,f329,f396])).
fof(f9221,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f9168,f269])).
fof(f9168,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1))))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f369,f329,f396])).
fof(f8617,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X88),s('type.list.list'('type.string.char'),X89),s('type.exhLang.exp',X86))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X90),s('type.list.list'('type.string.char'),X91),s('type.exhLang.exp',X92))) | s('type.exhLang.exp',X92) = s('type.exhLang.exp',sK4(X86,X87,'type.exhLang.exp')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.exhLang.exp',bool),X87)))) )),
  inference(superposition,[],[f401,f836])).
fof(f8616,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X81),s('type.list.list'('type.string.char'),X82),s('type.exhLang.exp',X80))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X83),s('type.list.list'('type.string.char'),X84),s('type.exhLang.exp',X85))) | s('type.exhLang.exp',X85) = s('type.exhLang.exp','const.pair.SND_1'(s('type.pair.prod'(X78,'type.exhLang.exp'),'const.pair.,_2'(s(X78,X79),s('type.exhLang.exp',X80)))))) )),
  inference(superposition,[],[f401,f342])).
fof(f8615,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X73),s('type.list.list'('type.string.char'),X74),s('type.exhLang.exp',X72))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X75),s('type.list.list'('type.string.char'),X76),s('type.exhLang.exp',X77))) | s('type.exhLang.exp',X77) = s('type.exhLang.exp','const.sum.OUTL_1'(s('type.sum.sum'('type.exhLang.exp',X71),'const.sum.INL_1'(s('type.exhLang.exp',X72)))))) )),
  inference(superposition,[],[f401,f336])).
fof(f8614,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X66),s('type.list.list'('type.string.char'),X67),s('type.exhLang.exp',X64))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X68),s('type.list.list'('type.string.char'),X69),s('type.exhLang.exp',X70))) | s('type.exhLang.exp',X70) = s('type.exhLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.exhLang.exp',X63),'const.pair.,_2'(s('type.exhLang.exp',X64),s(X63,X65)))))) )),
  inference(superposition,[],[f401,f341])).
fof(f8613,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X58),s('type.list.list'('type.string.char'),X59),s('type.exhLang.exp',X57))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X60),s('type.list.list'('type.string.char'),X61),s('type.exhLang.exp',X62))) | s('type.exhLang.exp',X62) = s('type.exhLang.exp','const.sum.OUTR_1'(s('type.sum.sum'(X56,'type.exhLang.exp'),'const.sum.INR_1'(s('type.exhLang.exp',X57)))))) )),
  inference(superposition,[],[f401,f337])).
fof(f8612,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X53),s('type.list.list'('type.string.char'),X54),s('type.exhLang.exp',X55))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X51),s('type.list.list'('type.string.char'),X52),s('type.exhLang.exp',happ(s(fun('type.num.num','type.exhLang.exp'),X49),s('type.num.num',X50))))) | s('type.exhLang.exp',X55) = s('type.exhLang.exp','const.ind_type.FCONS_3'(s('type.exhLang.exp',X48),s(fun('type.num.num','type.exhLang.exp'),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f401,f324])).
fof(f8611,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X43),s('type.list.list'('type.string.char'),X44),s('type.exhLang.exp',X41))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X45),s('type.list.list'('type.string.char'),X46),s('type.exhLang.exp',X47))) | s('type.exhLang.exp',X47) = s('type.exhLang.exp','const.ind_type.FCONS_3'(s('type.exhLang.exp',X41),s(fun('type.num.num','type.exhLang.exp'),X42),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f401,f316])).
fof(f8610,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X38),s('type.list.list'('type.string.char'),X39),s('type.exhLang.exp',X40))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X36),s('type.list.list'('type.string.char'),X37),s('type.exhLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.exhLang.exp',X34),X35))))) | s('type.exhLang.exp',X40) = s('type.exhLang.exp',happ(s(fun('type.pair.prod'('type.exhLang.exp',X34),'type.exhLang.exp'),'const.pair.FST_0'),s('type.pair.prod'('type.exhLang.exp',X34),X35)))) )),
  inference(superposition,[],[f401,f338])).
fof(f8609,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X29),s('type.list.list'('type.string.char'),X30),s('type.exhLang.exp',X28))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X31),s('type.list.list'('type.string.char'),X32),s('type.exhLang.exp',X33))) | s('type.exhLang.exp',X33) = s('type.exhLang.exp',happ(s(fun('type.list.list'(X26),'type.exhLang.exp'),sK18(X27,X28,X26,'type.exhLang.exp')),s('type.list.list'(X26),'const.list.NIL_0')))) )),
  inference(superposition,[],[f401,f357])).
fof(f8608,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X24),s('type.exhLang.exp',X25))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X21),s('type.list.list'('type.string.char'),X22),s('type.exhLang.exp',happ(s(fun(X17,'type.exhLang.exp'),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))))) | s('type.exhLang.exp',X25) = s('type.exhLang.exp',happ(s(fun(X16,'type.exhLang.exp'),'_dst_x0x1_2'(s(fun(X17,'type.exhLang.exp'),X18),s(fun(X16,X17),X19))),s(X16,X20)))) )),
  inference(superposition,[],[f401,f369])).
fof(f8607,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13),s('type.list.list'('type.string.char'),X14),s('type.exhLang.exp',X15))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp',happ(s(fun(X8,'type.exhLang.exp'),'_dst_x0x1_2'(s(fun(X6,'type.exhLang.exp'),X7),s(fun(X8,X6),X9))),s(X8,X10))))) | s('type.exhLang.exp',X15) = s('type.exhLang.exp',happ(s(fun(X6,'type.exhLang.exp'),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f401,f369])).
fof(f8129,plain,(
  ( ! [X70,X69] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X69),s('type.list.list'('type.exhSem.v'),sK15(happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)),X70,'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X70),s('type.list.list'('type.exhSem.v'),sK15(happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)),X70,'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) | s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X70) = s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)))) )),
  inference(superposition,[],[f353,f328])).
fof(f8128,plain,(
  ( ! [X68] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK15('const.exhSem.Vectorv_0',X68,'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X68),s('type.list.list'('type.exhSem.v'),sK15('const.exhSem.Vectorv_0',X68,'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) | s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),'const.exhSem.Vectorv_0') = s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f8127,plain,(
  ( ! [X67] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK15('const.exhSem.Litv_0',X67,'type.ast.lit','type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.ast.lit','type.exhSem.v'),X67),s('type.ast.lit',sK15('const.exhSem.Litv_0',X67,'type.ast.lit','type.exhSem.v')))) | s(fun('type.ast.lit','type.exhSem.v'),'const.exhSem.Litv_0') = s(fun('type.ast.lit','type.exhSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f8126,plain,(
  ( ! [X66] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK15('const.exhSem.Loc_0',X66,'type.num.num','type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.num.num','type.exhSem.v'),X66),s('type.num.num',sK15('const.exhSem.Loc_0',X66,'type.num.num','type.exhSem.v')))) | s(fun('type.num.num','type.exhSem.v'),'const.exhSem.Loc_0') = s(fun('type.num.num','type.exhSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f8125,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f8114,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X20,X21,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,X21,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f8113,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f8110,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f8109,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f8108,plain,(
  ( ! [X0,X1] : (s(X0,sK15('_i_0_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15('_i_0_0',X1,X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f8097,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,X9))) )),
  inference(superposition,[],[f353,f445])).
fof(f8096,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK15(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f353,f445])).
fof(f8154,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65))))) )),
  inference(forward_demodulation,[],[f8093,f342])).
fof(f8093,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f8153,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60))))) )),
  inference(forward_demodulation,[],[f8092,f336])).
fof(f8092,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f8152,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54))))) )),
  inference(forward_demodulation,[],[f8091,f341])).
fof(f8091,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f8151,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49))))) )),
  inference(forward_demodulation,[],[f8090,f337])).
fof(f8090,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f8150,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43))))) )),
  inference(forward_demodulation,[],[f8089,f324])).
fof(f8089,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f8149,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38))))) )),
  inference(forward_demodulation,[],[f8088,f316])).
fof(f8088,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f8148,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33))))) )),
  inference(forward_demodulation,[],[f8087,f338])).
fof(f8087,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f353,f338])).
fof(f8147,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27))))) )),
  inference(forward_demodulation,[],[f8086,f357])).
fof(f8086,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f8085,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f8084,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f8146,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7))))) )),
  inference(forward_demodulation,[],[f8083,f309])).
fof(f8083,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f8145,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1))))) )),
  inference(forward_demodulation,[],[f8082,f368])).
fof(f8082,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f8081,plain,(
  ( ! [X70,X69] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X69),s('type.list.list'('type.exhSem.v'),sK15(X70,happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)),'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X70),s('type.list.list'('type.exhSem.v'),sK15(X70,happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)),'type.list.list'('type.exhSem.v'),'type.exhSem.v')))) | s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X70) = s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),'type.exhSem.v')),'const.exhSem.Conv_0'),s('type.num.num',X69)))) )),
  inference(superposition,[],[f353,f328])).
fof(f8080,plain,(
  ( ! [X68] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK15(X68,'const.exhSem.Vectorv_0','type.list.list'('type.exhSem.v'),'type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X68),s('type.list.list'('type.exhSem.v'),sK15(X68,'const.exhSem.Vectorv_0','type.list.list'('type.exhSem.v'),'type.exhSem.v')))) | s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),'const.exhSem.Vectorv_0') = s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f8079,plain,(
  ( ! [X67] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK15(X67,'const.exhSem.Litv_0','type.ast.lit','type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.ast.lit','type.exhSem.v'),X67),s('type.ast.lit',sK15(X67,'const.exhSem.Litv_0','type.ast.lit','type.exhSem.v')))) | s(fun('type.ast.lit','type.exhSem.v'),'const.exhSem.Litv_0') = s(fun('type.ast.lit','type.exhSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f8078,plain,(
  ( ! [X66] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK15(X66,'const.exhSem.Loc_0','type.num.num','type.exhSem.v')))) != s('type.exhSem.v',happ(s(fun('type.num.num','type.exhSem.v'),X66),s('type.num.num',sK15(X66,'const.exhSem.Loc_0','type.num.num','type.exhSem.v')))) | s(fun('type.num.num','type.exhSem.v'),'const.exhSem.Loc_0') = s(fun('type.num.num','type.exhSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f8077,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f8066,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X21,X20,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X21,X20,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f8065,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f8062,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f8061,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f8060,plain,(
  ( ! [X0,X1] : (s(X0,sK15(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f8140,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65))))) )),
  inference(forward_demodulation,[],[f8058,f342])).
fof(f8058,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f8139,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60))))) )),
  inference(forward_demodulation,[],[f8057,f336])).
fof(f8057,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f8138,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54))))) )),
  inference(forward_demodulation,[],[f8056,f341])).
fof(f8056,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f8137,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49))))) )),
  inference(forward_demodulation,[],[f8055,f337])).
fof(f8055,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f8136,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43))))) )),
  inference(forward_demodulation,[],[f8054,f324])).
fof(f8054,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f8135,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38))))) )),
  inference(forward_demodulation,[],[f8053,f316])).
fof(f8053,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f8134,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33))))) )),
  inference(forward_demodulation,[],[f8052,f338])).
fof(f8052,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f353,f338])).
fof(f8133,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27))))) )),
  inference(forward_demodulation,[],[f8051,f357])).
fof(f8051,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f8050,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f8049,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f8132,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7))))) )),
  inference(forward_demodulation,[],[f8048,f309])).
fof(f8048,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f8131,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) )),
  inference(forward_demodulation,[],[f8047,f368])).
fof(f8047,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f7428,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f6722,f836])).
fof(f7427,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X125,bool),'const.pair.,_2'(s(X125,X126),s(bool,X127)))))) )),
  inference(superposition,[],[f6722,f342])).
fof(f7426,plain,(
  ( ! [X123,X124,X122] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X124,bool),'const.sum.INR_1'(s(bool,X123))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X122),'const.sum.INL_1'(s(bool,X123)))))) )),
  inference(superposition,[],[f6722,f336])).
fof(f7425,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X121,bool),'const.sum.INR_1'(s(bool,X119))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120)))))) )),
  inference(superposition,[],[f6722,f341])).
fof(f7424,plain,(
  ( ! [X116,X117,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X117,bool),'const.sum.INR_1'(s(bool,X116))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X115,bool),'const.sum.INR_1'(s(bool,X116)))))) )),
  inference(superposition,[],[f6722,f337])).
fof(f7422,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6722,f316])).
fof(f7420,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK18(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6722,f357])).
fof(f7258,plain,(
  ( ! [X520,X516,X518,X517,X519,X515] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X520))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X516,'type.ast.lit'),X517),s(X516,happ(s(fun(X515,X516),X518),s(X515,X519))))))) | s('type.ast.lit',X520) = s('type.ast.lit',happ(s(fun(X515,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X516,'type.ast.lit'),X517),s(fun(X515,X516),X518))),s(X515,X519)))) )),
  inference(superposition,[],[f334,f369])).
fof(f7256,plain,(
  ( ! [X505,X507,X508,X504,X506,X503] : (s('type.ast.lit',X508) != s('type.ast.lit',happ(s(fun(X504,'type.ast.lit'),X505),s(X504,happ(s(fun(X503,X504),X506),s(X503,X507))))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X508))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X503,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X504,'type.ast.lit'),X505),s(fun(X503,X504),X506))),s(X503,X507)))))) )),
  inference(superposition,[],[f335,f369])).
fof(f7252,plain,(
  ( ! [X476,X478,X472,X474,X471,X475,X477,X473] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X477),s('type.list.list'('type.exhSem.v'),X478))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X472,'type.num.num'),X473),s(X472,happ(s(fun(X471,X472),X474),s(X471,X475))))),s('type.list.list'('type.exhSem.v'),X476))) | s('type.num.num',X477) = s('type.num.num',happ(s(fun(X471,'type.num.num'),'_dst_x0x1_2'(s(fun(X472,'type.num.num'),X473),s(fun(X471,X472),X474))),s(X471,X475)))) )),
  inference(superposition,[],[f394,f369])).
fof(f7240,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (s(fun(X375,X376),X382) != s(fun(X375,X376),happ(s(fun(X378,fun(X375,X376)),X379),s(X378,happ(s(fun(X377,X378),X380),s(X377,X381))))) | s(X376,happ(s(fun(X375,X376),X382),s(X375,X383))) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),'_dst_x0x1_2'(s(fun(X378,fun(X375,X376)),X379),s(fun(X377,X378),X380))),s(X377,X381))),s(X375,X383)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7235,plain,(
  ( ! [X344,X346,X341,X343,X345,X347,X342] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X341,bool),happ(s(fun(X343,fun(X341,bool)),X344),s(X343,happ(s(fun(X342,X343),X345),s(X342,X346))))))) | s(X341,X347) = s(X341,sK4(X347,happ(s(fun(X342,fun(X341,bool)),'_dst_x0x1_2'(s(fun(X343,fun(X341,bool)),X344),s(fun(X342,X343),X345))),s(X342,X346)),X341))) )),
  inference(superposition,[],[f836,f369])).
fof(f7165,plain,(
  ( ! [X546,X542,X545,X541,X543,X544] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X546))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X543,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X541,'type.ast.lit'),X542),s(fun(X543,X541),X544))),s(X543,X545))))) | s('type.ast.lit',X546) = s('type.ast.lit',happ(s(fun(X541,'type.ast.lit'),X542),s(X541,happ(s(fun(X543,X541),X544),s(X543,X545)))))) )),
  inference(superposition,[],[f334,f369])).
fof(f7163,plain,(
  ( ! [X532,X534,X530,X533,X529,X531] : (s('type.ast.lit',X534) != s('type.ast.lit',happ(s(fun(X531,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X529,'type.ast.lit'),X530),s(fun(X531,X529),X532))),s(X531,X533))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X534))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X529,'type.ast.lit'),X530),s(X529,happ(s(fun(X531,X529),X532),s(X531,X533)))))))) )),
  inference(superposition,[],[f335,f369])).
fof(f7159,plain,(
  ( ! [X499,X500,X502,X498,X504,X501,X503,X497] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X503),s('type.list.list'('type.exhSem.v'),X504))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X499,'type.num.num'),'_dst_x0x1_2'(s(fun(X497,'type.num.num'),X498),s(fun(X499,X497),X500))),s(X499,X501))),s('type.list.list'('type.exhSem.v'),X502))) | s('type.num.num',X503) = s('type.num.num',happ(s(fun(X497,'type.num.num'),X498),s(X497,happ(s(fun(X499,X497),X500),s(X499,X501)))))) )),
  inference(superposition,[],[f394,f369])).
fof(f7147,plain,(
  ( ! [X389,X391,X393,X395,X388,X390,X396,X392,X394] : (s(fun(X388,X389),X395) != s(fun(X388,X389),happ(s(fun(X392,fun(X388,X389)),'_dst_x0x1_2'(s(fun(X390,fun(X388,X389)),X391),s(fun(X392,X390),X393))),s(X392,X394))) | s(X389,happ(s(fun(X388,X389),X395),s(X388,X396))) = s(X389,happ(s(fun(X388,X389),happ(s(fun(X390,fun(X388,X389)),X391),s(X390,happ(s(fun(X392,X390),X393),s(X392,X394))))),s(X388,X396)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7142,plain,(
  ( ! [X356,X358,X354,X360,X357,X359,X355] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X354,bool),happ(s(fun(X357,fun(X354,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X354,bool)),X356),s(fun(X357,X355),X358))),s(X357,X359))))) | s(X354,X360) = s(X354,sK4(X360,happ(s(fun(X355,fun(X354,bool)),X356),s(X355,happ(s(fun(X357,X355),X358),s(X357,X359)))),X354))) )),
  inference(superposition,[],[f836,f369])).
fof(f6992,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (s(fun(X101,X102),'_dst_x0x1_2'(s(fun(X103,X102),happ(s(fun(X105,fun(X103,X102)),X106),s(X105,X107))),s(fun(X101,X103),X104))) = s(fun(X101,X102),happ(s(fun(X105,fun(X101,X102)),'_dst_x0x1_2'(s(fun(fun(X103,X102),fun(X101,X102)),'_dstsel_0_x01_1'(s(fun(X101,X103),X104))),s(fun(X105,fun(X103,X102)),X106))),s(X105,X107)))) )),
  inference(superposition,[],[f369,f368])).
fof(f6981,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X39,X37),X40))),s(X39,X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X39,X37),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f369,f445])).
fof(f6980,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X33,X31),X34))),s(X33,X35))) | s(bool,X36) = s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X33,X31),X34),s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f369,f445])).
fof(f6978,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f369,f338])).
fof(f6977,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f6976,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(X110,happ(s(fun(fun(X105,X106),X110),X111),s(fun(X105,X106),'_dst_x0x1_2'(s(fun(X107,X106),X109),s(fun(X105,X107),X108))))) = s(X110,happ(s(fun(fun(X107,X106),X110),'_dst_x0x1_2'(s(fun(fun(X105,X106),X110),X111),s(fun(fun(X107,X106),fun(X105,X106)),'_dstsel_0_x01_1'(s(fun(X105,X107),X108))))),s(fun(X107,X106),X109)))) )),
  inference(superposition,[],[f369,f368])).
fof(f6975,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK18(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f369,f357])).
fof(f6974,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f6973,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f369,f369])).
fof(f6927,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f369,f352])).
fof(f7302,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),X5))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),X5)))) )),
  inference(forward_demodulation,[],[f6944,f342])).
fof(f6944,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),X5))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X6,'type.list.list'('type.exhSem.v')),'const.pair.,_2'(s(X6,X7),s('type.list.list'('type.exhSem.v'),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f342,f369,f396])).
fof(f7301,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X5),X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X5),X6)))))) )),
  inference(forward_demodulation,[],[f6945,f338])).
fof(f6945,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X5),X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.exhSem.v'),X5),'type.list.list'('type.exhSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.exhSem.v'),X5),X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f369,f396])).
fof(f7300,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X5),s('type.num.num',X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X5),s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f6946,f324])).
fof(f6946,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.exhSem.v'),happ(s(fun('type.num.num','type.list.list'('type.exhSem.v')),X5),s('type.num.num',X6))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.exhSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.exhSem.v'),X7),s(fun('type.num.num','type.list.list'('type.exhSem.v')),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(unit_resulting_resolution,[],[f324,f369,f396])).
fof(f6948,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X0,'type.ast.lit'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun(X2,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X0,'type.ast.lit'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f369,f335])).
fof(f7281,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f6968,f338])).
fof(f6968,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X0),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f369,f396])).
fof(f7280,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f6969,f324])).
fof(f6969,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X7),s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.list.list'('type.exhSem.v'),happ(s(fun(X4,'type.list.list'('type.exhSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.exhSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f324,f369,f396])).
fof(f6866,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6796,f264])).
fof(f6796,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2477])).
fof(f6722,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6653,f264])).
fof(f6653,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2475])).
fof(f5787,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X1))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(equality_resolution,[],[f396])).
fof(f5400,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f5365,f262])).
fof(f5365,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f4854,f4915])).
fof(f5396,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f5362,f262])).
fof(f5362,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4854,f836])).
fof(f5395,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5361,f262])).
fof(f5361,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4854,f2199])).
fof(f5391,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(subsumption_resolution,[],[f5358,f262])).
fof(f5358,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(superposition,[],[f4854,f440])).
fof(f5380,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5341,f262])).
fof(f5341,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4854,f445])).
fof(f5125,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5124,f262])).
fof(f5124,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f5115,f2199])).
fof(f5115,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f4813,f2522])).
fof(f5122,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5110,f262])).
fof(f5110,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4813,f445])).
fof(f5038,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4991,f262])).
fof(f4991,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f447,f4915])).
fof(f4909,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f357,f4840,f447])).
fof(f4915,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f432,f4840,f2480])).
fof(f4916,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4840,f4812])).
fof(f4840,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f262,f4812])).
fof(f4813,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4681,f338])).
fof(f4686,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4636,f262])).
fof(f4636,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f447,f836])).
fof(f4617,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X223),s('type.list.list'('type.exhSem.v'),X221))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X224),s('type.list.list'('type.exhSem.v'),X225))) | s('type.list.list'('type.exhSem.v'),X225) = s('type.list.list'('type.exhSem.v'),sK4(X221,X222,'type.list.list'('type.exhSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.exhSem.v'),bool),X222)))) )),
  inference(superposition,[],[f395,f836])).
fof(f4611,plain,(
  ( ! [X202,X203,X201] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X201))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X203))) | s('type.ast.lit',X203) = s('type.ast.lit',sK4(X201,X202,'type.ast.lit')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X202)))) )),
  inference(superposition,[],[f334,f836])).
fof(f4605,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X175),s('type.list.list'('type.exhSem.v'),X177))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X178),s('type.list.list'('type.exhSem.v'),X179))) | s('type.num.num',X178) = s('type.num.num',sK4(X175,X176,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X176)))) )),
  inference(superposition,[],[f394,f836])).
fof(f4682,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(forward_demodulation,[],[f4583,f2199])).
fof(f4583,plain,(
  ( ! [X107,X106] : (s(bool,X106) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(superposition,[],[f2297,f836])).
fof(f4517,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f836,f445])).
fof(f4516,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f836,f445])).
fof(f4365,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK14(X7,X8,X9,bool))) )),
  inference(superposition,[],[f351,f445])).
fof(f4364,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK14(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f4178,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3675,f445])).
fof(f3768,plain,(
  ( ! [X173,X174,X172] : (~p(s(bool,happ(s(fun(bool,bool),X173),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173)))) | s(bool,X174) = s(bool,sK4(X172,X173,bool)) | s(bool,'F_0') = s(bool,X174)) ) | $spl10),
  inference(superposition,[],[f319,f2865])).
fof(f3675,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f432,f2865])).
fof(f3604,plain,(
  ( ! [X185,X186,X184] : (~p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X186)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X185)))) | s(bool,'F_0') = s(bool,X186) | s(bool,'F_0') = s(bool,sK4(X184,X185,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2865])).
fof(f3463,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2865,f2865])).
fof(f3257,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X45),s('type.list.list'('type.exhSem.v'),X44))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X46),s('type.list.list'('type.exhSem.v'),X47))) | s('type.list.list'('type.exhSem.v'),X47) = s('type.list.list'('type.exhSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X42,'type.list.list'('type.exhSem.v')),'const.pair.,_2'(s(X42,X43),s('type.list.list'('type.exhSem.v'),X44)))))) )),
  inference(superposition,[],[f395,f342])).
fof(f3256,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X39),s('type.list.list'('type.exhSem.v'),X38))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X40),s('type.list.list'('type.exhSem.v'),X41))) | s('type.list.list'('type.exhSem.v'),X41) = s('type.list.list'('type.exhSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.exhSem.v'),X37),'const.sum.INL_1'(s('type.list.list'('type.exhSem.v'),X38)))))) )),
  inference(superposition,[],[f395,f336])).
fof(f3255,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X34),s('type.list.list'('type.exhSem.v'),X32))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X35),s('type.list.list'('type.exhSem.v'),X36))) | s('type.list.list'('type.exhSem.v'),X36) = s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.exhSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f395,f341])).
fof(f3254,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X28),s('type.list.list'('type.exhSem.v'),X27))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X29),s('type.list.list'('type.exhSem.v'),X30))) | s('type.list.list'('type.exhSem.v'),X30) = s('type.list.list'('type.exhSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.list.list'('type.exhSem.v')),'const.sum.INR_1'(s('type.list.list'('type.exhSem.v'),X27)))))) )),
  inference(superposition,[],[f395,f337])).
fof(f3252,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.exhSem.v'),X15))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.exhSem.v'),X19))) | s('type.list.list'('type.exhSem.v'),X19) = s('type.list.list'('type.exhSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.exhSem.v'),X15),s(fun('type.num.num','type.list.list'('type.exhSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f395,f316])).
fof(f3251,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X13),s('type.list.list'('type.exhSem.v'),X14))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.exhSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.exhSem.v'),X10),X11))))) | s('type.list.list'('type.exhSem.v'),X14) = s('type.list.list'('type.exhSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.exhSem.v'),X10),'type.list.list'('type.exhSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.exhSem.v'),X10),X11)))) )),
  inference(superposition,[],[f395,f338])).
fof(f3250,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X7),s('type.list.list'('type.exhSem.v'),X6))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.exhSem.v'),X9))) | s('type.list.list'('type.exhSem.v'),X9) = s('type.list.list'('type.exhSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.exhSem.v')),sK18(X5,X6,X4,'type.list.list'('type.exhSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f395,f357])).
fof(f3249,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.exhSem.v'),X0))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | s('type.list.list'('type.exhSem.v'),X3) = s('type.list.list'('type.exhSem.v'),happ(s(fun('type.list.list'('type.exhSem.v'),'type.list.list'('type.exhSem.v')),'_i_0_0'),s('type.list.list'('type.exhSem.v'),X0)))) )),
  inference(superposition,[],[f395,f309])).
fof(f3027,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2928,f2199])).
fof(f2928,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2522])).
fof(f3011,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3010,f262])).
fof(f3010,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2896,f2199])).
fof(f2896,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f319,f2522])).
fof(f2865,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2522,f2522])).
fof(f2810,plain,(
  ( ! [X177,X175,X176,X174] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X177))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',happ(s(fun('type.num.num','type.ast.lit'),X175),s('type.num.num',X176))))) | s('type.ast.lit',X177) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X174),s(fun('type.num.num','type.ast.lit'),X175),s('type.num.num','const.num.SUC_1'(s('type.num.num',X176)))))) )),
  inference(superposition,[],[f334,f324])).
fof(f2515,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2490,f262])).
fof(f2490,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2480,f445])).
fof(f2297,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2272,f262])).
fof(f2272,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f447,f2199])).
fof(f2180,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f357,f2036,f447])).
fof(f2199,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2036,f272])).
fof(f1378,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f319,f270])).
fof(f1376,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f319,f445])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f319,f445])).
fof(f1497,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f319,f338])).
fof(f1858,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X42),s('type.list.list'('type.exhSem.v'),X43))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.exhSem.v'),X45))) | s('type.num.num',X44) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X40,'type.num.num'),'const.pair.,_2'(s(X40,X41),s('type.num.num',X42)))))) )),
  inference(superposition,[],[f394,f342])).
fof(f1857,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X36),s('type.list.list'('type.exhSem.v'),X37))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X38),s('type.list.list'('type.exhSem.v'),X39))) | s('type.num.num',X38) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X35),'const.sum.INL_1'(s('type.num.num',X36)))))) )),
  inference(superposition,[],[f394,f336])).
fof(f1856,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X30),s('type.list.list'('type.exhSem.v'),X32))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X33),s('type.list.list'('type.exhSem.v'),X34))) | s('type.num.num',X33) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X29),'const.pair.,_2'(s('type.num.num',X30),s(X29,X31)))))) )),
  inference(superposition,[],[f394,f341])).
fof(f1855,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X25),s('type.list.list'('type.exhSem.v'),X26))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X27),s('type.list.list'('type.exhSem.v'),X28))) | s('type.num.num',X27) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X24,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f394,f337])).
fof(f1854,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X19),s('type.list.list'('type.exhSem.v'),X21))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X22),s('type.list.list'('type.exhSem.v'),X23))) | s('type.num.num',X22) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X19),s(fun('type.num.num','type.num.num'),X20),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f394,f316])).
fof(f1853,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.exhSem.v'),X18))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X14),X15))),s('type.list.list'('type.exhSem.v'),X16))) | s('type.num.num',X17) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X14),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X14),X15)))) )),
  inference(superposition,[],[f394,f338])).
fof(f1852,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X10),s('type.list.list'('type.exhSem.v'),X11))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.exhSem.v'),X13))) | s('type.num.num',X12) = s('type.num.num',happ(s(fun('type.list.list'(X8),'type.num.num'),sK18(X9,X10,X8,'type.num.num')),s('type.list.list'(X8),'const.list.NIL_0')))) )),
  inference(superposition,[],[f394,f357])).
fof(f1851,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X4),s('type.list.list'('type.exhSem.v'),X5))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X6),s('type.list.list'('type.exhSem.v'),X7))) | s('type.num.num',X6) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f394,f309])).
fof(f1519,plain,(
  ( ! [X99,X97,X98] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X99))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X97),X98))))) | s('type.ast.lit',X99) = s('type.ast.lit',happ(s(fun('type.pair.prod'('type.ast.lit',X97),'type.ast.lit'),'const.pair.FST_0'),s('type.pair.prod'('type.ast.lit',X97),X98)))) )),
  inference(superposition,[],[f334,f338])).
fof(f1360,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1117,f262])).
fof(f1117,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f447,f445])).
fof(f1359,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1116,f262])).
fof(f1116,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f418,f445])).
fof(f1358,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1115,f262])).
fof(f1115,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f333,f445])).
fof(f1006,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f445,f342])).
fof(f1005,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f445,f336])).
fof(f1004,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f445,f341])).
fof(f1003,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f445,f337])).
fof(f1002,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f445,f316])).
fof(f1001,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK18(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f445,f357])).
fof(f1000,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f445,f309])).
fof(f990,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f445,f445])).
fof(f973,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) )),
  inference(superposition,[],[f445,f342])).
fof(f972,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X39)))))) )),
  inference(superposition,[],[f445,f336])).
fof(f971,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f445,f341])).
fof(f970,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,X32)))))) )),
  inference(superposition,[],[f445,f337])).
fof(f969,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f445,f316])).
fof(f968,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f445,f357])).
fof(f967,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f445,f309])).
fof(f934,plain,(
  ( ! [X103,X105,X106,X104] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X105))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X106))) | s('type.ast.lit',X106) = s('type.ast.lit',happ(s(fun('type.list.list'(X103),'type.ast.lit'),sK18(X104,X105,X103,'type.ast.lit')),s('type.list.list'(X103),'const.list.NIL_0')))) )),
  inference(superposition,[],[f334,f357])).
fof(f842,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X17))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X18))) | s('type.ast.lit',X18) = s('type.ast.lit','const.pair.SND_1'(s('type.pair.prod'(X15,'type.ast.lit'),'const.pair.,_2'(s(X15,X16),s('type.ast.lit',X17)))))) )),
  inference(superposition,[],[f334,f342])).
fof(f841,plain,(
  ( ! [X14,X12,X13] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X13))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X14))) | s('type.ast.lit',X14) = s('type.ast.lit','const.sum.OUTL_1'(s('type.sum.sum'('type.ast.lit',X12),'const.sum.INL_1'(s('type.ast.lit',X13)))))) )),
  inference(superposition,[],[f334,f336])).
fof(f840,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X9))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X11))) | s('type.ast.lit',X11) = s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X8),'const.pair.,_2'(s('type.ast.lit',X9),s(X8,X10)))))) )),
  inference(superposition,[],[f334,f341])).
fof(f839,plain,(
  ( ! [X6,X7,X5] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X6))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))) | s('type.ast.lit',X7) = s('type.ast.lit','const.sum.OUTR_1'(s('type.sum.sum'(X5,'type.ast.lit'),'const.sum.INR_1'(s('type.ast.lit',X6)))))) )),
  inference(superposition,[],[f334,f337])).
fof(f838,plain,(
  ( ! [X4,X2,X3] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4))) | s('type.ast.lit',X4) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X2),s(fun('type.num.num','type.ast.lit'),X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f334,f316])).
fof(f837,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X1) = s('type.ast.lit',happ(s(fun('type.ast.lit','type.ast.lit'),'_i_0_0'),s('type.ast.lit',X0)))) )),
  inference(superposition,[],[f334,f309])).
fof(f836,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f835,f264])).
fof(f835,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f318,f270])).
fof(f803,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f790,f262])).
fof(f790,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f447,f440])).
fof(f494,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f489,f264])).
fof(f489,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f330,f270])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f275,f270])).
fof(f440,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f265,f272])).
fof(f432,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f264,f271])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & ((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5)) | s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16,X17,X21,X18,X19,X22] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X16),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X17),s('type.list.list'('type.string.char'),X21))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X19),s('type.list.list'('type.string.char'),X22))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X16) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X18) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X17) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X19) & s('type.list.list'('type.string.char'),X21) = s('type.list.list'('type.string.char'),X22)))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_|49||49||split|3')).
fof(f404,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4)) )),
  inference(cnf_transformation,[],[f259])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)) )),
  inference(cnf_transformation,[],[f259])).
fof(f406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X4) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)) )),
  inference(cnf_transformation,[],[f259])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.exhLang.exp',X2) = s('type.exhLang.exp',X5))) & (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.exhLang.exp',X2) != s('type.exhLang.exp',X5) | s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5)))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.exhLang.exp',X2) = s('type.exhLang.exp',X5))) & ((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.exhLang.exp',X2) != s('type.exhLang.exp',X5)) | s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5)))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.exhLang.exp',X2) = s('type.exhLang.exp',X5)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X16,X17,X21,X18,X19,X22] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X16),s('type.list.list'('type.string.char'),X17),s('type.exhLang.exp',X21))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X18),s('type.list.list'('type.string.char'),X19),s('type.exhLang.exp',X22))) <=> (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X16) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X18) & s('type.list.list'('type.string.char'),X17) = s('type.list.list'('type.string.char'),X19) & s('type.exhLang.exp',X21) = s('type.exhLang.exp',X22)))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_|49||49||split|2')).
fof(f400,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4)) )),
  inference(cnf_transformation,[],[f257])).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | s('type.exhLang.exp',X2) = s('type.exhLang.exp',X5)) )),
  inference(cnf_transformation,[],[f257])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.exhLang.exp',X2) != s('type.exhLang.exp',X5) | s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp',X5))) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f397,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X25,X26,X52,X27] : ! [X28,X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X52,X25),X28),s(fun(X26,X52),happ(s(fun(X27,fun(X26,X52)),X29),s(X27,X30))))) = s(fun(X26,X25),happ(s(fun(X27,fun(X26,X25)),'_dst_x00x11_2'(s(fun(X52,X25),X28),s(fun(X27,fun(X26,X52)),X29))),s(X27,X30)))),
  file('thm.exhSem.v_|49||49||split|4','_dst_x00x11_ax')).
fof(f395,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f396,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.exhSem.v'),X1) != s('type.list.list'('type.exhSem.v'),X3) | s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.exhSem.v'),X1))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.exhSem.v'),X3))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f418,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X38,X55,X43,X56] : ((s(bool,X38) = s(bool,X55) & (p(s(bool,X55)) => s(bool,X43) = s(bool,X56))) => ((p(s(bool,X38)) => p(s(bool,X43))) <=> (p(s(bool,X55)) => p(s(bool,X56)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.IMP_CONG')).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.exhSem.v_|49||49||split|4','_dst_x0x1_ax')).
fof(f368,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X25,X26,X27] : ! [X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X30),s(fun(X26,X27),X29))) = s(fun(X26,X25),happ(s(fun(fun(X27,X25),fun(X26,X25)),'_dstsel_0_x01_1'(s(fun(X26,X27),X29))),s(fun(X27,X25),X30)))),
  file('thm.exhSem.v_|49||49||split|4','_dstsel_0_x01_ax')).
fof(f367,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ! [X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X45,X9,X31] : ! [X10] : ? [X33] : ! [X38,X43] : s(X45,happ(s(fun(X31,X45),happ(s(fun(X9,fun(X31,X45)),X10),s(X9,X38))),s(X31,X43))) = s(X45,happ(s(fun('type.pair.prod'(X9,X31),X45),X33),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.pair.pair_Axiom')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.exhSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.exhSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.exhSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.exhSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X20,X23] : s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X20),s('type.list.list'('type.string.char'),X23))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.exhSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.exhSem.Recclosure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X20))),s('type.list.list'('type.string.char'),X23)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_3_const.exhSem.Recclosure')).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v',happ(s(fun('type.exhLang.exp','type.exhSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v'))),'const.exhSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0))),s('type.list.list'('type.string.char'),X1))),s('type.exhLang.exp',X2)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X2))) = s('type.exhSem.v',happ(s(fun('type.exhLang.exp','type.exhSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v'))),'const.exhSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0))),s('type.list.list'('type.string.char'),X1))),s('type.exhLang.exp',X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8,X20,X23] : s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8),s('type.list.list'('type.string.char'),X20),s('type.exhLang.exp',X23))) = s('type.exhSem.v',happ(s(fun('type.exhLang.exp','type.exhSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v')),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp','type.exhSem.v'))),'const.exhSem.Closure_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))),s('type.list.list'('type.string.char'),X20))),s('type.exhLang.exp',X23)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_3_const.exhSem.Closure')).
fof(f357,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f248])).
fof(f352,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X31,X9] : ! [X10,X37] : (s(fun(X9,X31),X10) = s(fun(X9,X31),X37) <=> ! [X38] : s(X31,happ(s(fun(X9,X31),X10),s(X9,X38))) = s(X31,happ(s(fun(X9,X31),X37),s(X9,X38))))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.FUN_EQ_THM')).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f351,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.exhSem.v_|49||49||split|4',aHL_EXT)).
fof(f342,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X31] : ! [X38,X43] : s(X31,X43) = s(X31,'const.pair.SND_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.pair.SND')).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31,X9] : ! [X38,X43] : s(X9,X38) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.pair.FST')).
fof(f340,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X31,X9] : ! [X48] : ? [X10] : ! [X49,X50,X51] : s(X31,happ(s(fun('type.ind_type.recspace'(X9),X31),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X49),s(X9,X50),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))))) = s(X31,happ(s(fun(fun('type.num.num',X31),X31),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)))),X48),s('type.num.num',X49))),s(X9,X50))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))),s(fun('type.num.num',X31),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X31),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.CONSTR_REC')).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9,X31] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X31),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X31),X8)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_1_const.pair.FST')).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X31] : ! [X38] : s(X31,X38) = s(X31,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X31),'const.sum.INR_1'(s(X31,X38)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.sum.OUTR')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X31,X9] : ! [X38] : s(X9,X38) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X31),'const.sum.INL_1'(s(X9,X38)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.sum.OUTL')).
fof(f334,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X6))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))) <=> s('type.ast.lit',X6) = s('type.ast.lit',X7))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_|49||49||split|0')).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f233])).
fof(f330,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X40,X41] : ((p(s(bool,X40)) => p(s(bool,X41))) => ((p(s(bool,X41)) => p(s(bool,X40))) => s(bool,X40) = s(bool,X41)))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f329,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X0),s('type.list.list'('type.exhSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X0))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X1)))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X16,X17] : s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X16),s('type.list.list'('type.exhSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v',X16))),s('type.num.num','const.exhSem.v3_size_1'(s('type.list.list'('type.exhSem.v'),X17)))))))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_size_def|split|10')).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X20,X23] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X20),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X20))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_1_3_const.ind_type.CONSTR')).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X10,X42] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X42))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.FCONS|split|1')).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.exhSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f272,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f269,plain,(
  ( ! [X0] : (s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.exhSem.v_size_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_size_def|split|4')).
fof(f268,plain,(
  ( ! [X0] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),'const.exhSem.Vectorv_0'),s('type.list.list'('type.exhSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),'const.exhSem.Vectorv_0'),s('type.list.list'('type.exhSem.v'),X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X8))) = s('type.exhSem.v',happ(s(fun('type.list.list'('type.exhSem.v'),'type.exhSem.v'),'const.exhSem.Vectorv_0'),s('type.list.list'('type.exhSem.v'),X8)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_1_const.exhSem.Vectorv')).
fof(f267,plain,(
  ( ! [X0] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v',happ(s(fun('type.ast.lit','type.exhSem.v'),'const.exhSem.Litv_0'),s('type.ast.lit',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v',happ(s(fun('type.ast.lit','type.exhSem.v'),'const.exhSem.Litv_0'),s('type.ast.lit',X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X8))) = s('type.exhSem.v',happ(s(fun('type.ast.lit','type.exhSem.v'),'const.exhSem.Litv_0'),s('type.ast.lit',X8)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_1_const.exhSem.Litv')).
fof(f266,plain,(
  ( ! [X0] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) = s('type.exhSem.v',happ(s(fun('type.num.num','type.exhSem.v'),'const.exhSem.Loc_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) = s('type.exhSem.v',happ(s(fun('type.num.num','type.exhSem.v'),'const.exhSem.Loc_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X8))) = s('type.exhSem.v',happ(s(fun('type.num.num','type.exhSem.v'),'const.exhSem.Loc_0'),s('type.num.num',X8)))),
  file('thm.exhSem.v_|49||49||split|4','ARITY_0_1_const.exhSem.Loc')).
fof(f265,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.exhSem.v_|49||49||split|4','thm.ind_type.ISO_REFL')).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exhSem.v_|49||49||split|4',aHL_TRUTH)).
fof(f24999,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f5337,f5202])).
fof(f5202,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f5163,f264])).
fof(f5163,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f5128,f2865])).
fof(f5337,plain,(
  ( ! [X23,X22] : (s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl12),
  inference(subsumption_resolution,[],[f5290,f262])).
fof(f5290,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl12),
  inference(superposition,[],[f447,f5152])).
fof(f5152,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f432,f5128,f2480])).
fof(f1394,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1395,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2394,plain,(
  ( ! [X23] : (~p(s(bool,sK4(X23,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f28427,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,sK4(X1,'_i_0_0',bool))) ) | $spl44),
  inference(cnf_transformation,[],[f28427_D])).
fof(f28427_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,sK4(X1,'_i_0_0',bool))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f28486,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2438,f262])).
fof(f2438,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f319,f472])).
fof(f7686,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f7686_D])).
fof(f7686_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4831,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4798,f262])).
fof(f4798,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4681,f472])).
fof(f2034,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) | $spl8),
  inference(cnf_transformation,[],[f2034_D])).
fof(f2034_D,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4811,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4681,f309])).
fof(f14353,plain,(
  ( ! [X46] : (~p(s(bool,sK4(X46,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f15989,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4681,f309])).
fof(f28445,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f16693,f262])).
fof(f16693,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f319,f472])).
fof(f16791,plain,(
  ( ! [X3] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f16695,f262])).
fof(f16695,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f4681,f472])).
fof(f21950,plain,(
  ( ! [X39] : (~p(s(bool,sK4(X39,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f28421,plain,(
  ( ! [X80] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl42),
  inference(cnf_transformation,[],[f28421_D])).
fof(f28421_D,plain,(
  ( ! [X80] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f22354,plain,(
  ( ! [X81,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f22301,f21178])).
fof(f22301,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,sK7('const.pair.FST_0',X80,bool)))))) | s(bool,'T_0') = s(bool,X79)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f22181])).
fof(f22181,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,sK7('const.pair.FST_0',X80,bool)))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | $spl26),
  inference(superposition,[],[f6722,f21177])).
fof(f20177,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X112),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl36),
  inference(cnf_transformation,[],[f20177_D])).
fof(f20177_D,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X112),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f22355,plain,(
  ( ! [X84,X82] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X84),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f22300,f21178])).
fof(f22300,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X84),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X83,bool)))))) | s(bool,'T_0') = s(bool,X82)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f22182])).
fof(f22182,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X84),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X83,bool)))))) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) ) | $spl26),
  inference(superposition,[],[f6866,f21177])).
fof(f24644,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4681,f309])).
fof(f28426,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X1,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f25363,f262])).
fof(f25363,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X1,'_i_0_0',bool))) )),
  inference(superposition,[],[f319,f472])).
fof(f25468,plain,(
  ( ! [X2] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f25364,f262])).
fof(f25364,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f4681,f472])).
fof(f28425,plain,(
  ( ! [X111,X112] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X112,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f26036,f6722])).
fof(f26036,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X112,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X111))))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f6722,f2475])).
fof(f28424,plain,(
  ( ! [X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X115),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f26037,f6722])).
fof(f26037,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X115),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f6866,f2475])).
fof(f28423,plain,(
  ( ! [X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X115,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f26299,f6866])).
fof(f26299,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X115,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f6722,f2477])).
fof(f28422,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(subsumption_resolution,[],[f26300,f6866])).
fof(f26300,plain,(
  ( ! [X118,X116,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f6866,f2477])).
fof(f28420,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X79)) ) | $spl10),
  inference(subsumption_resolution,[],[f26631,f11678])).
fof(f26631,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X79))) | s(bool,'T_0') = s(bool,X79)) ) | $spl10),
  inference(superposition,[],[f6722,f2975])).
fof(f28419,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X82),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X81)) ) | $spl10),
  inference(subsumption_resolution,[],[f26632,f11678])).
fof(f26632,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X82),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X81))) | s(bool,'T_0') = s(bool,X81)) ) | $spl10),
  inference(superposition,[],[f6866,f2975])).
fof(f28261,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X49)) ) | ($spl6 | $spl10)),
  inference(duplicate_literal_removal,[],[f28059])).
fof(f28059,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f2650,f22047])).
fof(f22047,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl10)),
  inference(trivial_inequality_removal,[],[f22036])).
fof(f22036,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f20176,f445])).
fof(f20176,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f18611,f1410])).
fof(f1410,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f271])).
fof(f18611,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X67,'_i_0_0',bool))) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f11678,f1409])).
fof(f1409,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f276])).
fof(f2650,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2649,f264])).
fof(f2649,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2616,f1409])).
fof(f2616,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2482,f1433])).
fof(f1433,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1411,f262])).
fof(f1411,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1395,f445])).
fof(f22352,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(cnf_transformation,[],[f22352_D])).
fof(f22352_D,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f22303,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl26)),
  inference(duplicate_literal_removal,[],[f22173])).
fof(f22173,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl26)),
  inference(superposition,[],[f2650,f21177])).
fof(f21176,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20190,f262])).
fof(f20190,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f333])).
fof(f21175,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20192,f262])).
fof(f20192,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f7455,f9671,f333])).
fof(f7455,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f7381,f1409])).
fof(f7381,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1410,f6722])).
fof(f21174,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20194,f262])).
fof(f20194,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1410,f9671,f333])).
fof(f21173,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20196,f9671])).
fof(f20196,plain,(
  ~p(s(bool,'F_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f333])).
fof(f21172,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20198,f9671])).
fof(f20198,plain,(
  ( ! [X0] : (~p(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0'))))))) ) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f7455,f9671,f333])).
fof(f21171,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20200,f9671])).
fof(f20200,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1410,f9671,f333])).
fof(f21170,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20204,f262])).
fof(f20204,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f333])).
fof(f21169,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20206,f262])).
fof(f20206,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f7455,f9671,f333])).
fof(f21168,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20208,f262])).
fof(f20208,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1410,f9671,f333])).
fof(f21167,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20210,f9671])).
fof(f20210,plain,(
  ~p(s(bool,'F_0')) | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f333])).
fof(f21166,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20212,f9671])).
fof(f20212,plain,(
  ( ! [X0] : (~p(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0'))))))) ) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f7455,f9671,f333])).
fof(f21165,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20214,f9671])).
fof(f20214,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1410,f9671,f333])).
fof(f21159,plain,(
  $false | ($spl6 | $spl10 | $spl28)),
  inference(subsumption_resolution,[],[f20823,f20329])).
fof(f20329,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f2480])).
fof(f20823,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl6 | $spl10 | $spl28)),
  inference(backward_demodulation,[],[f20329,f20176])).
fof(f21155,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20810,f20329])).
fof(f20810,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl28)),
  inference(backward_demodulation,[],[f20329,f13740])).
fof(f13740,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f13693,f337])).
fof(f13693,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f7455,f6866])).
fof(f21148,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f20737,f20329])).
fof(f20737,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl28)),
  inference(backward_demodulation,[],[f20329,f7455])).
fof(f21145,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f21144,f9671])).
fof(f21144,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl28),
  inference(subsumption_resolution,[],[f20722,f20329])).
fof(f20722,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl28),
  inference(backward_demodulation,[],[f20329,f7197])).
fof(f21143,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f21142,f9671])).
fof(f21142,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f20710,f20329])).
fof(f20710,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl28),
  inference(backward_demodulation,[],[f20329,f7103])).
fof(f21135,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f21134,f9671])).
fof(f21134,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f20542,f20329])).
fof(f20542,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl28),
  inference(backward_demodulation,[],[f20329,f275])).
fof(f21133,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20541,f262])).
fof(f20541,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(backward_demodulation,[],[f20329,f264])).
fof(f20330,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f432,f432,f9671,f2480])).
fof(f20322,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1409,f1410,f9671,f447])).
fof(f20320,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f337,f7455,f9671,f447])).
fof(f20318,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1409,f432,f9671,f447])).
fof(f20266,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f342,f9671,f418])).
fof(f20264,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f336,f9671,f418])).
fof(f20262,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f341,f9671,f418])).
fof(f20260,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f337,f9671,f418])).
fof(f20258,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f316,f9671,f418])).
fof(f20256,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f357,f9671,f418])).
fof(f20254,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f309,f9671,f418])).
fof(f20252,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f264,f1409,f9671,f418])).
fof(f20215,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f9671,f1410,f9671,f333])).
fof(f20213,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f9671,f7455,f9671,f333])).
fof(f20211,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f9671,f432,f9671,f333])).
fof(f20209,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f262,f1410,f9671,f333])).
fof(f20207,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f262,f7455,f9671,f333])).
fof(f20205,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f262,f432,f9671,f333])).
fof(f20201,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f9671,f1410,f9671,f333])).
fof(f20199,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f9671,f7455,f9671,f333])).
fof(f20197,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f9671,f432,f9671,f333])).
fof(f20195,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f262,f1410,f9671,f333])).
fof(f20193,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f262,f7455,f9671,f333])).
fof(f20191,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f262,f432,f9671,f333])).
fof(f20187,plain,(
  $false | ($spl6 | $spl28)),
  inference(unit_resulting_resolution,[],[f1409,f9671,f275])).
fof(f20185,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f432,f9671,f272])).
fof(f20183,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f264,f9671])).
fof(f20181,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f17348,f6722])).
fof(f17348,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,X106))))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f6866,f2475])).
fof(f20179,plain,(
  ( ! [X109,X110] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X110),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f17636,f6866])).
fof(f17636,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X110),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X108),'const.sum.INL_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f6866,f2477])).
fof(f20178,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X75)) ) | $spl10),
  inference(subsumption_resolution,[],[f17969,f11678])).
fof(f17969,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | $spl10),
  inference(superposition,[],[f6866,f2975])).
fof(f18432,plain,(
  ( ! [X112,X110] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X112),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X110)) ) | $spl6),
  inference(forward_demodulation,[],[f18337,f337])).
fof(f18337,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X112),'const.sum.INL_1'(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,'F_0'))))))))) | s(bool,'T_0') = s(bool,X110)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f18175])).
fof(f18175,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X112),'const.sum.INL_1'(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,'F_0'))))))))) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) ) | $spl6),
  inference(superposition,[],[f6866,f7540])).
fof(f7540,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f7526])).
fof(f7526,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f7455,f445])).
fof(f20099,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl34),
  inference(cnf_transformation,[],[f20099_D])).
fof(f20099_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f20024,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(duplicate_literal_removal,[],[f19827])).
fof(f19827,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(superposition,[],[f2650,f11446])).
fof(f13289,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278,f13280,f13282,f13284,f13286,f13288])).
fof(f13288,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (s(X394,X401) = s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X396,fun(X393,X394)),X397),s(fun(X395,X396),X398))),s(X395,X399))),s(X393,happ(s(fun(X394,X393),X400),s(X394,X401)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13015,f262])).
fof(f13015,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (~p(s(bool,'T_0')) | s(X394,X401) = s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X396,fun(X393,X394)),X397),s(fun(X395,X396),X398))),s(X395,X399))),s(X393,happ(s(fun(X394,X393),X400),s(X394,X401)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7242])).
fof(f13286,plain,(
  ( ! [X412,X414,X416,X418,X413,X415,X411,X417,X419] : (s(X411,X419) = s(X411,happ(s(fun(X412,X411),X418),s(X412,happ(s(fun(X411,X412),happ(s(fun(X413,fun(X411,X412)),'_dst_x0x1_2'(s(fun(X414,fun(X411,X412)),X415),s(fun(X413,X414),X416))),s(X413,X417))),s(X411,X419)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13014,f262])).
fof(f13014,plain,(
  ( ! [X412,X414,X416,X418,X413,X415,X411,X417,X419] : (~p(s(bool,'T_0')) | s(X411,X419) = s(X411,happ(s(fun(X412,X411),X418),s(X412,happ(s(fun(X411,X412),happ(s(fun(X413,fun(X411,X412)),'_dst_x0x1_2'(s(fun(X414,fun(X411,X412)),X415),s(fun(X413,X414),X416))),s(X413,X417))),s(X411,X419)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7244])).
fof(f13284,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13013,f262])).
fof(f13013,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9677])).
fof(f13282,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13012,f262])).
fof(f13012,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9959])).
fof(f13280,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X409,X411,X406] : (s(X407,X414) = s(X407,happ(s(fun(X406,X407),happ(s(fun(X408,fun(X406,X407)),X409),s(X408,happ(s(fun(X410,X408),X411),s(X410,X412))))),s(X406,happ(s(fun(X407,X406),X413),s(X407,X414)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13011,f262])).
fof(f13011,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X409,X411,X406] : (~p(s(bool,'T_0')) | s(X407,X414) = s(X407,happ(s(fun(X406,X407),happ(s(fun(X408,fun(X406,X407)),X409),s(X408,happ(s(fun(X410,X408),X411),s(X410,X412))))),s(X406,happ(s(fun(X407,X406),X413),s(X407,X414)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7149])).
fof(f13278,plain,(
  ( ! [X428,X430,X424,X426,X432,X429,X431,X425,X427] : (s(X424,X432) = s(X424,happ(s(fun(X425,X424),X431),s(X425,happ(s(fun(X424,X425),happ(s(fun(X426,fun(X424,X425)),X427),s(X426,happ(s(fun(X428,X426),X429),s(X428,X430))))),s(X424,X432)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13010,f262])).
fof(f13010,plain,(
  ( ! [X428,X430,X424,X426,X432,X429,X431,X425,X427] : (~p(s(bool,'T_0')) | s(X424,X432) = s(X424,happ(s(fun(X425,X424),X431),s(X425,happ(s(fun(X424,X425),happ(s(fun(X426,fun(X424,X425)),X427),s(X426,happ(s(fun(X428,X426),X429),s(X428,X430))))),s(X424,X432)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7151])).
fof(f13276,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13009,f262])).
fof(f13009,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9676])).
fof(f13274,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13008,f262])).
fof(f13008,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9958])).
fof(f13272,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (s(X402,X410) = s(X402,happ(s(fun(X403,X402),X409),s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),'_dst_x0x1_2'(s(fun(X405,fun(X402,X403)),X406),s(fun(X404,X405),X407))),s(X404,X408))),s(X402,X410)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13007,f262])).
fof(f13007,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (~p(s(bool,'T_0')) | s(X402,X410) = s(X402,happ(s(fun(X403,X402),X409),s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),'_dst_x0x1_2'(s(fun(X405,fun(X402,X403)),X406),s(fun(X404,X405),X407))),s(X404,X408))),s(X402,X410)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7243])).
fof(f13270,plain,(
  ( ! [X420,X422,X428,X424,X426,X421,X423,X425,X427] : (s(X421,X428) = s(X421,happ(s(fun(X420,X421),happ(s(fun(X422,fun(X420,X421)),'_dst_x0x1_2'(s(fun(X423,fun(X420,X421)),X424),s(fun(X422,X423),X425))),s(X422,X426))),s(X420,happ(s(fun(X421,X420),X427),s(X421,X428)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13006,f262])).
fof(f13006,plain,(
  ( ! [X420,X422,X428,X424,X426,X421,X423,X425,X427] : (~p(s(bool,'T_0')) | s(X421,X428) = s(X421,happ(s(fun(X420,X421),happ(s(fun(X422,fun(X420,X421)),'_dst_x0x1_2'(s(fun(X423,fun(X420,X421)),X424),s(fun(X422,X423),X425))),s(X422,X426))),s(X420,happ(s(fun(X421,X420),X427),s(X421,X428)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7245])).
fof(f13268,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13005,f262])).
fof(f13005,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9690])).
fof(f13266,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13004,f262])).
fof(f13004,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9972])).
fof(f13264,plain,(
  ( ! [X420,X422,X416,X418,X415,X421,X423,X417,X419] : (s(X415,X423) = s(X415,happ(s(fun(X416,X415),X422),s(X416,happ(s(fun(X415,X416),happ(s(fun(X417,fun(X415,X416)),X418),s(X417,happ(s(fun(X419,X417),X420),s(X419,X421))))),s(X415,X423)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13003,f262])).
fof(f13003,plain,(
  ( ! [X420,X422,X416,X418,X415,X421,X423,X417,X419] : (~p(s(bool,'T_0')) | s(X415,X423) = s(X415,happ(s(fun(X416,X415),X422),s(X416,happ(s(fun(X415,X416),happ(s(fun(X417,fun(X415,X416)),X418),s(X417,happ(s(fun(X419,X417),X420),s(X419,X421))))),s(X415,X423)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7150])).
fof(f13262,plain,(
  ( ! [X441,X436,X438,X434,X440,X437,X439,X433,X435] : (s(X434,X441) = s(X434,happ(s(fun(X433,X434),happ(s(fun(X435,fun(X433,X434)),X436),s(X435,happ(s(fun(X437,X435),X438),s(X437,X439))))),s(X433,happ(s(fun(X434,X433),X440),s(X434,X441)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13002,f262])).
fof(f13002,plain,(
  ( ! [X441,X436,X438,X434,X440,X437,X439,X433,X435] : (~p(s(bool,'T_0')) | s(X434,X441) = s(X434,happ(s(fun(X433,X434),happ(s(fun(X435,fun(X433,X434)),X436),s(X435,happ(s(fun(X437,X435),X438),s(X437,X439))))),s(X433,happ(s(fun(X434,X433),X440),s(X434,X441)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7152])).
fof(f13260,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13001,f262])).
fof(f13001,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9689])).
fof(f13258,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl30),
  inference(subsumption_resolution,[],[f13000,f262])).
fof(f13000,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f9971])).
fof(f13256,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12988,f262])).
fof(f12988,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f355])).
fof(f13254,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12987,f262])).
fof(f12987,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f354])).
fof(f13252,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl30),
  inference(subsumption_resolution,[],[f12986,f262])).
fof(f12986,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f350])).
fof(f13250,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl30),
  inference(subsumption_resolution,[],[f12985,f262])).
fof(f12985,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f349])).
fof(f13066,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f12491,f12417])).
fof(f12491,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f275])).
fof(f13287,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278,f13280,f13282,f13284,f13286])).
fof(f13285,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278,f13280,f13282,f13284])).
fof(f13283,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278,f13280,f13282])).
fof(f13281,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278,f13280])).
fof(f13279,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276,f13278])).
fof(f13277,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274,f13276])).
fof(f13275,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272,f13274])).
fof(f13273,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270,f13272])).
fof(f13271,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268,f13270])).
fof(f13269,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266,f13268])).
fof(f13267,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264,f13266])).
fof(f13265,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262,f13264])).
fof(f13263,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260,f13262])).
fof(f13261,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258,f13260])).
fof(f13259,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256,f13258])).
fof(f13257,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254,f13256])).
fof(f13255,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252,f13254])).
fof(f13253,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250,f13252])).
fof(f13251,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066,f13250])).
fof(f13103,plain,(
  $false | ($spl26 | $spl30)),
  inference(subsumption_resolution,[],[f12748,f12417])).
fof(f12748,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl26 | $spl30)),
  inference(backward_demodulation,[],[f12417,f10793])).
fof(f13085,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f12692,f12417])).
fof(f12692,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f12417,f7455])).
fof(f13082,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13081,f262])).
fof(f13081,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(forward_demodulation,[],[f13080,f12417])).
fof(f13080,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12677,f12417])).
fof(f12677,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7197])).
fof(f13079,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13078,f262])).
fof(f13078,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(forward_demodulation,[],[f13077,f12417])).
fof(f13077,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12665,f12417])).
fof(f12665,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(backward_demodulation,[],[f12417,f7103])).
fof(f13067,plain,(
  $false | $spl30),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1851,f1852,f1853,f1854,f1855,f1856,f1857,f1858,f2033,f1484,f1486,f1488,f2037,f398,f356,f355,f354,f350,f349,f2017,f339,f348,f345,f2032,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2473,f2427,f2474,f2475,f2476,f2477,f2478,f2480,f2482,f2483,f2515,f2520,f2521,f2764,f2766,f2792,f2841,f2803,f2804,f2810,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3256,f3257,f4151,f4152,f4153,f4154,f4155,f4156,f4157,f4158,f4159,f4160,f4274,f4275,f4276,f4277,f4278,f4279,f4280,f4281,f4360,f4363,f4364,f4365,f4375,f4376,f4377,f4422,f4423,f4424,f4425,f4426,f4427,f4428,f4429,f4509,f4515,f4516,f4517,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4681,f4564,f4565,f4566,f4567,f4568,f4570,f4571,f4576,f4577,f4578,f4579,f4580,f4584,f4585,f4598,f4599,f4600,f4601,f4605,f4609,f4611,f4617,f4684,f4685,f4686,f4715,f4754,f4833,f4812,f4813,f4840,f4854,f4916,f4915,f4909,f5038,f5049,f5059,f5093,f5122,f5121,f5212,f5222,f5256,f5380,f5390,f5391,f5392,f5396,f5400,f5404,f5370,f5371,f5372,f5373,f5374,f5375,f5376,f5377,f5378,f5379,f5776,f5787,f6563,f6566,f6588,f6631,f6632,f6633,f6634,f6636,f6638,f6639,f6722,f6732,f6774,f6775,f6776,f6777,f6779,f6781,f6782,f6866,f6876,f6970,f7280,f7281,f7282,f6959,f6948,f7300,f7301,f7302,f6927,f6973,f6974,f6975,f6976,f6977,f6978,f6979,f6980,f6981,f6992,f7098,f7099,f7101,f7102,f7103,f7104,f7105,f7107,f7108,f7113,f7114,f7115,f7116,f7117,f7121,f7122,f7123,f7124,f7135,f7141,f7142,f7143,f7145,f7146,f7147,f7149,f7150,f7151,f7152,f7159,f7163,f7165,f7172,f7174,f7190,f7191,f7192,f7193,f7195,f7196,f7197,f7198,f7199,f7201,f7202,f7207,f7208,f7209,f7210,f7211,f7215,f7216,f7217,f7218,f7229,f7234,f7235,f7236,f7238,f7239,f7240,f7242,f7243,f7244,f7245,f7252,f7256,f7258,f7265,f7267,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7674,f7675,f7583,f7584,f7676,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7593,f7685,f7688,f7694,f7662,f7663,f7664,f7695,f7869,f7870,f7828,f7829,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7838,f7839,f7840,f7841,f7884,f7885,f7893,f7867,f7894,f8131,f8132,f8049,f8050,f8133,f8134,f8135,f8136,f8137,f8138,f8139,f8140,f8059,f8060,f8061,f8062,f8063,f8064,f8065,f8066,f8077,f8078,f8079,f8080,f8081,f8145,f8146,f8084,f8085,f8147,f8148,f8149,f8150,f8151,f8152,f8153,f8154,f8094,f8095,f8096,f8097,f8108,f8109,f8110,f8111,f8112,f8113,f8114,f8125,f8126,f8127,f8128,f8129,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f9221,f9223,f9225,f9227,f9240,f9198,f9204,f9582,f9583,f9524,f9525,f9584,f9585,f9586,f9587,f9588,f9589,f9590,f9591,f9534,f9535,f9536,f9537,f9550,f9551,f9552,f9604,f9611,f9581,f9627,f9676,f9677,f9686,f9689,f9690,f9699,f9736,f9737,f9958,f9959,f9968,f9971,f9972,f9981,f10018,f10541,f10542,f10543,f10544,f10545,f10546,f10547,f10548,f10549,f10550,f10551,f10752,f10753,f10754,f10755,f10756,f10757,f10758,f10759,f10760,f10761,f10762,f11040,f11041,f11042,f11056,f11070,f11071,f11072,f11073,f11088,f11090,f11091,f11095,f11096,f11109,f11112,f11113,f11124,f11337,f11338,f11128,f11340,f11347,f11348,f11353,f11153,f11159,f11167,f11177,f11184,f11695,f11696,f11697,f11698,f11699,f11700,f11701,f11702,f11703,f11704,f11705,f11706,f11707,f11978,f11979,f11980,f11981,f11982,f11983,f11984,f11985,f11986,f11987,f11988,f11989,f11990,f11991,f12023,f12024,f11994,f11995,f12025,f12026,f12027,f12028,f12029,f12030,f12031,f12032,f12033,f12005,f12035,f12036,f12417,f13066])).
fof(f13065,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12490,f262])).
fof(f12490,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(backward_demodulation,[],[f12417,f264])).
fof(f12418,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl32),
  inference(cnf_transformation,[],[f12418_D])).
fof(f12418_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f12345,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f12164])).
fof(f12164,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2650,f7540])).
fof(f9670,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl26),
  inference(cnf_transformation,[],[f9670_D])).
fof(f9670_D,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9671,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl28),
  inference(cnf_transformation,[],[f9671_D])).
fof(f9671_D,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9660,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl10),
  inference(superposition,[],[f9630,f341])).
fof(f7752,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(cnf_transformation,[],[f7752_D])).
fof(f7752_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7754,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl24),
  inference(cnf_transformation,[],[f7754_D])).
fof(f7754_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f7736,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f7687,f309])).
fof(f7687,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl20),
  inference(cnf_transformation,[],[f7687_D])).
fof(f7687_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7694,plain,(
  ( ! [X64,X65] : (p(s(bool,sK5('_i_0_0',X65,bool))) | ~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool)))))) )),
  inference(forward_demodulation,[],[f7661,f309])).
fof(f7661,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X65,bool)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f7685,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f7684,f309])).
fof(f7684,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7647,f262])).
fof(f7647,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f320,f472])).
fof(f5398,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5398_D])).
fof(f5398_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5397,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5363,f275])).
fof(f5363,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4854,f2523])).
fof(f2523,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2501,f262])).
fof(f2501,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f2480,f1610])).
fof(f1610,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1557,f272])).
fof(f1557,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1410,f318])).
fof(f5394,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl14),
  inference(cnf_transformation,[],[f5394_D])).
fof(f5394_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5393,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5360,f275])).
fof(f5360,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4854,f2522])).
fof(f5128,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl12),
  inference(cnf_transformation,[],[f5128_D])).
fof(f5128_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5121,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4813,f341])).
fof(f2176,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2085,f264])).
fof(f2085,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f2034,f330])).
fof(f2175,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2087,f2034])).
fof(f2087,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f2034,f330])).
fof(f2174,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2095,f264])).
fof(f2095,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f2034,f330])).
fof(f2173,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2097,f2034])).
fof(f2097,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f2034,f330])).
fof(f2105,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f2034,f318])).
fof(f2102,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f440,f2034,f271])).
fof(f2100,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f2034,f276])).
fof(f2098,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f2034,f1410,f2034,f330])).
fof(f2096,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f2034,f330])).
fof(f2094,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f2034,f1410,f2034,f330])).
fof(f2092,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2034,f432,f2034,f330])).
fof(f2088,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f2034,f1410,f2034,f330])).
fof(f2086,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f2034,f330])).
fof(f2084,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f2034,f1410,f2034,f330])).
fof(f2082,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2034,f432,f2034,f330])).
fof(f2066,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f342,f2034,f418])).
fof(f2064,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f336,f2034,f418])).
fof(f2062,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f341,f2034,f418])).
fof(f2060,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f337,f2034,f418])).
fof(f2058,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f316,f2034,f418])).
fof(f2056,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f357,f2034,f418])).
fof(f2054,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f309,f2034,f418])).
fof(f2051,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f262,f1610,f2034,f418])).
fof(f2049,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f440,f2034,f418])).
fof(f2046,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1410,f2034,f494])).
fof(f2044,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f432,f2034,f494])).
fof(f2038,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f2034])).
fof(f2036,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2036_D])).
fof(f2036_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2033,plain,(
  ( ! [X23] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f2016,f309])).
fof(f2016,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f317,f309])).
fof(f1384,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f431,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f430])).
fof(f430,plain,(
  s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f427,f428])).
fof(f426,plain,(
  s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f428,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f198])).
fof(f198,plain,(
  ? [X0,X1] : ((s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ? [X0,X1] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.exhSem.v_|49||49||split|4','thm.exhSem.v_|49||49||split|4')).
fof(f427,plain,(
  s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f429,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f261,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK0))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
% SZS output end Proof for thm.exhSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 258
% Passive clauses: 3171
% Generated clauses: 122166
% Final active clauses: 137
% Final passive clauses: 1463
% Input formulas: 69
% Initial clauses: 147
% Discarded non-redundant clauses: 5
% 
% Duplicate literals: 1360
% Trivial inequalities: 299
% Fw subsumption resolutions: 1870
% Fw demodulations: 62261
% Bw demodulations: 1233
% Condensations: 1643
% Global subsumptions: 27
% 
% Simple tautologies: 432
% Equational tautologies: 251
% Forward subsumptions: 26902
% Fw demodulations to eq. taut.: 23807
% Bw demodulations to eq. taut.: 115
% 
% Unit resulting resolution: 31179
% Factoring: 1
% Forward superposition: 11005
% Backward superposition: 9098
% Self superposition: 201
% Equality factoring: 110
% Equality resolution: 46
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 3383
% SAT solver unit clauses: 279
% SAT solver binary clauses: 1137
% 
% Sat splits: 40
% Sat splitting refutations: 103
% 
% Memory used [KB]: 56289
% Time elapsed: 4.781 s
% ------------------------------
% ------------------------------
% Success in time 6.13 s
