Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.source_to_modProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.source_to_modProof
% SZS output start Proof for thm.source_to_modProof
fof(f812,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f810,f146,f49,f77,f73,f50,f75,f74])).
fof(f74,plain,(
  s(sK1,sK2) != s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) | $spl1),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  s(sK1,sK2) != s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f75,plain,(
  ( ! [X4] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X4),s(sK1,sK2)))) ) | $spl2),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  ( ! [X4] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X4),s(sK1,sK2)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f50,plain,(
  ( ! [X4] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X4),s(sK1,sK2))) | s(sK1,sK2) != s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  (s(sK1,sK2) = s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK4),s(sK1,sK2)))) & (s(sK1,sK2) != s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) | ! [X4] : s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X4),s(sK1,sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f38])).
fof(f38,plain,(
  ? [X0,X1,X2,X3] : ((s(X1,X2) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | ? [X5] : s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X2)))) & (s(X1,X2) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | ! [X4] : s('type.pair.prod'(X0,X1),X3) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2,X3] : ((s(X1,X2) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | ? [X4] : s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2)))) & (s(X1,X2) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | ! [X4] : s('type.pair.prod'(X0,X1),X3) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2)))))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ? [X0,X1,X2,X3] : (s(X1,X2) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) <~> ? [X4] : s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1,X2,X3] : (s(X1,X2) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) <=> ? [X4] : s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1] : ! [X2,X3] : (s(X1,X2) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) <=> ? [X4] : s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (s(X7,X8) = s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),X9))) <=> ? [X10] : s('type.pair.prod'(X6,X7),X9) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X10),s(X7,X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (s(X7,X8) = s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),X9))) <=> ? [X10] : s('type.pair.prod'(X6,X7),X9) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X10),s(X7,X8))))),
  file('thm.source_to_modProof.source_to_modProof_|49|87_prove_ltb','thm.source_to_modProof.source_to_modProof_|49|87_prove_ltb')).
fof(f73,plain,(
  s(sK1,sK2) = s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) | $spl0),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  s(sK1,sK2) = s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f77,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK4),s(sK1,sK2))) | $spl4),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK4),s(sK1,sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f49,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK4),s(sK1,sK2))) | s(sK1,sK2) = s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3)))),
  inference(cnf_transformation,[],[f39])).
fof(f146,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f138,f74])).
fof(f138,plain,(
  s(sK1,sK2) = s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3))) | $spl4),
  inference(superposition,[],[f67,f77])).
fof(f67,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7,X6] : ! [X9,X8] : s(X6,X8) = s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X8)))))),
  file('thm.source_to_modProof.source_to_modProof_|49|87_prove_ltb','thm.pair.SND')).
fof(f810,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f806,f75])).
fof(f806,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6(sK3,sK1,sK0)),s(sK1,sK2))) | $spl0),
  inference(superposition,[],[f66,f768])).
fof(f768,plain,(
  s(sK1,sK2) = s(sK1,sK7(sK3,sK1,sK0)) | $spl0),
  inference(superposition,[],[f750,f73])).
fof(f750,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK7(X2,X1,X0)) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))) )),
  inference(superposition,[],[f67,f66])).
fof(f66,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK6(X2,X1,X0)),s(X1,sK7(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK6(X2,X1,X0)),s(X1,sK7(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7,X6] : ! [X9] : ? [X15,X16] : s('type.pair.prod'(X7,X6),X9) = s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X15),s(X6,X16)))),
  file('thm.source_to_modProof.source_to_modProof_|49|87_prove_ltb','thm.pair.ABS_PAIR_THM')).
% SZS output end Proof for thm.source_to_modProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 25
% Passive clauses: 198
% Generated clauses: 729
% Final active clauses: 21
% Final passive clauses: 159
% Input formulas: 13
% Initial clauses: 23
% 
% Duplicate literals: 79
% Trivial inequalities: 8
% Fw subsumption resolutions: 57
% Fw demodulations: 24
% 
% Simple tautologies: 35
% Equational tautologies: 29
% Forward subsumptions: 294
% 
% Binary resolution: 29
% Factoring: 1
% Forward superposition: 166
% Backward superposition: 233
% Self superposition: 47
% Equality factoring: 50
% Equality resolution: 1
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 18
% SAT solver unit clauses: 4
% SAT solver binary clauses: 9
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 9
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 16
% TWLsolver calls for satisfiability: 11
% 
% Memory used [KB]: 895
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
% Success in time 0.03 s
