Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.PairingHeap
% (140155)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1014
% Passive clauses: 7474
% Generated clauses: 36915
% Final active clauses: 863
% Final passive clauses: 6046
% Input formulas: 261
% Initial clauses: 507
% 
% Duplicate literals: 181
% Trivial inequalities: 44
% Fw subsumption resolutions: 1734
% Fw demodulations: 3463
% Bw demodulations: 145
% 
% Simple tautologies: 319
% Equational tautologies: 25
% Forward subsumptions: 23342
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 6013
% Factoring: 454
% Forward superposition: 12290
% Backward superposition: 11843
% Self superposition: 54
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 374
% Split components: 750
% SAT solver clauses: 514
% SAT solver unit clauses: 15
% SAT solver binary clauses: 76
% 
% Sat splits: 374
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 514
% TWLsolver calls for satisfiability: 356
% 
% Memory used [KB]: 19573
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.PairingHeap
% (140180)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 578
% Passive clauses: 2760
% Generated clauses: 23041
% Final active clauses: 489
% Final passive clauses: 1976
% Input formulas: 261
% Initial clauses: 507
% 
% Duplicate literals: 97
% Trivial inequalities: 29
% Fw subsumption resolutions: 674
% Fw demodulations: 2693
% Bw demodulations: 262
% Condensations: 54
% Global subsumptions: 24
% 
% Simple tautologies: 278
% Equational tautologies: 9
% Forward subsumptions: 13107
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 1976
% Factoring: 123
% Forward superposition: 6763
% Backward superposition: 6846
% Self superposition: 76
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 3680
% SAT solver unit clauses: 646
% SAT solver binary clauses: 132
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 17
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 3584
% TWLsolver calls for satisfiability: 3106
% 
% Memory used [KB]: 28272
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.PairingHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.PairingHeap
% SZS output start Proof for thm.PairingHeap
fof(f54022,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f927,f1463,f1461,f926,f1462,f1460,f5883,f5903,f5884,f5906,f5887,f5908,f5888,f5911,f5909,f5890,f5912,f5891,f5913,f5892,f5905,f5914,f5900,f5919,f8094,f8169,f5917,f12005,f12025,f12006,f12026,f12028,f12029,f12031,f12033,f12037,f12039,f12041,f12043,f12045,f12047,f12051,f12053,f12055,f12057,f12059,f12061,f12073,f12075,f12087,f12089,f12097,f12105,f12112,f12114,f12121,f12123,f12130,f12132,f12139,f12141,f12179,f12318,f12320,f12322,f12324,f12327,f12348,f12350,f12358,f12360,f12368,f12370,f12385,f12387,f12588,f12593,f12595,f12597,f12600,f12603,f12605,f12611,f12613,f12619,f12621,f12632,f12634,f13199,f13201,f13203,f13205,f13208,f13229,f13231,f13239,f13241,f13249,f13251,f13266,f13268,f13469,f13474,f13476,f13478,f13481,f13484,f13486,f13492,f13494,f13500,f13502,f13513,f13515,f13977,f13979,f13981,f14001,f14003,f14005,f14035,f14037,f14051,f14053,f14083,f14085,f14099,f14101,f14131,f14133,f14147,f14149,f14186,f14188,f14190,f14192,f14194,f14196,f15254,f15255,f14241,f14245,f14248,f14252,f14257,f14259,f14268,f14276,f14278,f14282,f14285,f14289,f14294,f14296,f14305,f14313,f14327,f14341,f14361,f14363,f14365,f14373,f14375,f14383,f14385,f14393,f14395,f14403,f14405,f14421,f14440,f14459,f14543,f14547,f14562,f14564,f14569,f14571,f14596,f14601,f14604,f14615,f14617,f14619,f14621,f14623,f14625,f14627,f14629,f14631,f14634,f14655,f14670,f14689,f14704,f14706,f14721,f14723,f14732,f14734,f14749,f14751,f14760,f14762,f14779,f14781,f14790,f14792,f14812,f14831,f14868,f14870,f14884,f14886,f14900,f14902,f14916,f14966,f14969,f14987,f14989,f14991,f15005,f15007,f15016,f15024,f15032,f15049,f15052,f15054,f15056,f15058,f15060,f15073,f15075,f15077,f15079,f15081,f15083,f15085,f15087,f15089,f15091,f15093,f15095,f15097,f15142,f15144,f15146,f15148,f15178,f15180,f15208,f15210,f15218,f15220,f15222,f15256,f15257,f15258,f15259,f15260,f15261,f15262,f15263,f15264,f15265,f15266,f15267,f15268,f15269,f15270,f15271,f15272,f15273,f15274,f15275,f15276,f15277,f15278,f15279,f15280,f15281,f15282,f15283,f15284,f15285,f15286,f15287,f15288,f15289,f15290,f15291,f15292,f15293,f15294,f15295,f15296,f15297,f15298,f15299,f15300,f15301,f15302,f15303,f15304,f15305,f15306,f15307,f15308,f15309,f15310,f15311,f15312,f15313,f15314,f15315,f15316,f15317,f15318,f15319,f15320,f15321,f15322,f15323,f15324,f15325,f15326,f15327,f15328,f15329,f15330,f15331,f15332,f15333,f15334,f15335,f15336,f15337,f15338,f15339,f15340,f15341,f15342,f15343,f15344,f15345,f15346,f15347,f15348,f15349,f15350,f15351,f15352,f15353,f15354,f15355,f15356,f15357,f15358,f15359,f15360,f15361,f15362,f15363,f15364,f15365,f15366,f15367,f15368,f15369,f15370,f15371,f15372,f15373,f15374,f15375,f15376,f15377,f15378,f15379,f15380,f15381,f15382,f15383,f15384,f15385,f15386,f15387,f15388,f15389,f15390,f15391,f15392,f15393,f15394,f15395,f15396,f15397,f15398,f15399,f15400,f15401,f15402,f15403,f15404,f15405,f15406,f15407,f15408,f15409,f15410,f15411,f15412,f15413,f15414,f15415,f15416,f15417,f15418,f15419,f15420,f15421,f15422,f15423,f15424,f15425,f15426,f15427,f15428,f15429,f15430,f15431,f15432,f15433,f15434,f15435,f15436,f15437,f15438,f15439,f15440,f15441,f15442,f15443,f15444,f15445,f15446,f15447,f15448,f15455,f15457,f15458,f15459,f15460,f15461,f15463,f15465,f15467,f15469,f15471,f15472,f15474,f15475,f15476,f15477,f15478,f15479,f15480,f15481,f15488,f15489,f15490,f15491,f15492,f15493,f15494,f15495,f15496,f15497,f15498,f15499,f15500,f15501,f15502,f15503,f15504,f15505,f15506,f15507,f15508,f15509,f15510,f15511,f15512,f15513,f15514,f15515,f15516,f15517,f15518,f15519,f15520,f15521,f15522,f15523,f15524,f15525,f15526,f15527,f15528,f15529,f15530,f15531,f15532,f15533,f15534,f15535,f15536,f15537,f15538,f15539,f15540,f15541,f15542,f15543,f15544,f15545,f15546,f15547,f15548,f15549,f15550,f15551,f15552,f15553,f15554,f15555,f15556,f15557,f15558,f15559,f15560,f15561,f15562,f15563,f15564,f15565,f15566,f15567,f15568,f15569,f15570,f15571,f15572,f15573,f15574,f15575,f15576,f15577,f15578,f15579,f15580,f15581,f15582,f15583,f15584,f15585,f15586,f15587,f15588,f15589,f15590,f15591,f15592,f15593,f15594,f15595,f15596,f15597,f15598,f15599,f15600,f15601,f15602,f15603,f15604,f15605,f15606,f15607,f15608,f15609,f15610,f15611,f15612,f15613,f15614,f15615,f15616,f15617,f15618,f15619,f15620,f15621,f15622,f15623,f15624,f15625,f15626,f15627,f15628,f15629,f15630,f15631,f15632,f15633,f15634,f15635,f15636,f15637,f15638,f15639,f15640,f15641,f15642,f15643,f15644,f15645,f15646,f15647,f15648,f15649,f15650,f15651,f15652,f15653,f15654,f15655,f15656,f15657,f15658,f15659,f15660,f15661,f15662,f15663,f15664,f15665,f15666,f15667,f15668,f15669,f15670,f15671,f15672,f15673,f15674,f15675,f15676,f15677,f15678,f15679,f15680,f15681,f15682,f15683,f15684,f15685,f15686,f15687,f15688,f15689,f15690,f15691,f15692,f15693,f15694,f15695,f15696,f15697,f15698,f15699,f15700,f15701,f15702,f15703,f15704,f15705,f15706,f15707,f15708,f15709,f15710,f15711,f15712,f15713,f15714,f15715,f15716,f15717,f15718,f15719,f15720,f15721,f15722,f15723,f15724,f15725,f15726,f15727,f15728,f15729,f15730,f15731,f15732,f15733,f28784,f28785,f28787,f28788,f28791,f28793,f28795,f28800,f28802,f28804,f28806,f28808,f28810,f28814,f28816,f28818,f28820,f28822,f28824,f28826,f28828,f28842,f28844,f28858,f28860,f28870,f28880,f28889,f28897,f28899,f28908,f28916,f28918,f28926,f28928,f28936,f28938,f28967,f29035,f29147,f29149,f29303,f29305,f29307,f29309,f29312,f29327,f29329,f29343,f29345,f29353,f29355,f29363,f29365,f29380,f29382,f29524,f29526,f29557,f29559,f29680,f29684,f29687,f29689,f29691,f29694,f29698,f29700,f29704,f29708,f29710,f29717,f29719,f29732,f29734,f29921,f29927,f29934,f29936,f30388,f30418,f30485,f30487,f30489,f30511,f30528,f30565,f30567,f30569,f30571,f30573,f30575,f32237,f32238,f30620,f30624,f30627,f30631,f30636,f30638,f30640,f30663,f30672,f30676,f30680,f30683,f30687,f30692,f30694,f30696,f30719,f30728,f30798,f30800,f30802,f30804,f30806,f30808,f30825,f30842,f30884,f30886,f30891,f30908,f30910,f30912,f30914,f30920,f30922,f30964,f30966,f30971,f30988,f30990,f30992,f30994,f31000,f31002,f31024,f31041,f31078,f31080,f31082,f31104,f31106,f31118,f31120,f31144,f31146,f31158,f31160,f31182,f31184,f31196,f31198,f31217,f31234,f31253,f31270,f31293,f31316,f31354,f31356,f31358,f31398,f31402,f31405,f31408,f31432,f31434,f31436,f31448,f31450,f31462,f31464,f31476,f31478,f31490,f31492,f31518,f31536,f31573,f31575,f31577,f31595,f31597,f31599,f31601,f31603,f31605,f31661,f31663,f31665,f31667,f31669,f31671,f31673,f31675,f31677,f31679,f31715,f31720,f31756,f31761,f31796,f31798,f31830,f31832,f31834,f31846,f31857,f31868,f31870,f31881,f31883,f31904,f31922,f31941,f31959,f32014,f32016,f32018,f32020,f32022,f32024,f32026,f32028,f32030,f32032,f32049,f32051,f32053,f32055,f32057,f32059,f32061,f32063,f32065,f32067,f32069,f32071,f32073,f32075,f32077,f32090,f32092,f32094,f32096,f32098,f32100,f32102,f32104,f32106,f32108,f32123,f32125,f32127,f32129,f32131,f32133,f32135,f32137,f32139,f32141,f32143,f32146,f32150,f32156,f32158,f32160,f32167,f32169,f32177,f32179,f32181,f32183,f32930,f32944,f32956,f32959,f33015,f33039,f33054,f33058,f33075,f33084,f33086,f33088,f33093,f33097,f33099,f33171,f33304,f33313,f33318,f33327,f33329,f33332,f33333,f33334,f33335,f33336,f33337,f33338,f33339,f33340,f33341,f33342,f33343,f33344,f33345,f33346,f33347,f33348,f33349,f33350,f33351,f33352,f33353,f33354,f33355,f33356,f33357,f33358,f33359,f33360,f33361,f33362,f33363,f33364,f33365,f33366,f33367,f33368,f33369,f33370,f33371,f33372,f33373,f33374,f33375,f33376,f33377,f33378,f33379,f33380,f33381,f33382,f33383,f33384,f33385,f33386,f33387,f33388,f33389,f33390,f33391,f33392,f33393,f33394,f33395,f33396,f33397,f33398,f33399,f33400,f33401,f33402,f33403,f33404,f33405,f33406,f33407,f33408,f33409,f33410,f33411,f33412,f33413,f33414,f33415,f33416,f33417,f33418,f33419,f33420,f33421,f33422,f33423,f33424,f33425,f33426,f33427,f33428,f33429,f33430,f33431,f33432,f33433,f33434,f33435,f33436,f33437,f33438,f33439,f33440,f33441,f33442,f33443,f33444,f33445,f33446,f33447,f33448,f33449,f33450,f33451,f33452,f33453,f33454,f33455,f33456,f33457,f33458,f33459,f33460,f33461,f33462,f33463,f33464,f33465,f33466,f33467,f33468,f33469,f33470,f33471,f33472,f33473,f33474,f33475,f33476,f33477,f33478,f33479,f33480,f33481,f33482,f33483,f33484,f33485,f33486,f33487,f33488,f33489,f33490,f33491,f33492,f33493,f33494,f33495,f33496,f33497,f33498,f33499,f33500,f33501,f33502,f33503,f33504,f33505,f33506,f33507,f33508,f33509,f33510,f33511,f33512,f33513,f33514,f33515,f33516,f33517,f33518,f33519,f33520,f33521,f33522,f33523,f33524,f33525,f33526,f33527,f33528,f33529,f33530,f33531,f33532,f33533,f33534,f33535,f33536,f33537,f33538,f33539,f33540,f33541,f33542,f33543,f33544,f33545,f33546,f33547,f33548,f33549,f33550,f33551,f33552,f33553,f33554,f33555,f33556,f33557,f33558,f33559,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33572,f33573,f33574,f33575,f33576,f33577,f33578,f33579,f33580,f33581,f33582,f33583,f33584,f33585,f33586,f33587,f33588,f33589,f33590,f33591,f33592,f33593,f33594,f33595,f33596,f33597,f33598,f33599,f33600,f33601,f33602,f33603,f33604,f33605,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33621,f33622,f33623,f33624,f33625,f33626,f33627,f33628,f33629,f33630,f33631,f33632,f33633,f33634,f33635,f33636,f33637,f33638,f33639,f33640,f33641,f33642,f33643,f33644,f33645,f33646,f33647,f33648,f33649,f33650,f33651,f33652,f33653,f33654,f33655,f33702,f33703,f33704,f33705,f33706,f33707,f33708,f33709,f33710,f33711,f33712,f33713,f33714,f33715,f33716,f33717,f33718,f33719,f33720,f33721,f33722,f33723,f33724,f33725,f33726,f33727,f33728,f33729,f33730,f33731,f33732,f33733,f33734,f33735,f33736,f33737,f33738,f33739,f33740,f33741,f33742,f33743,f33744,f33745,f33746,f33747,f33748,f33749,f33750,f33751,f33752,f33753,f33754,f33755,f33756,f33757,f33758,f33759,f33760,f33761,f33762,f33763,f33764,f33765,f33766,f33767,f33768,f33769,f33770,f33771,f33772,f33773,f33774,f33775,f33776,f33777,f33778,f33779,f33780,f33781,f33782,f33783,f33784,f33785,f33786,f33787,f33788,f33789,f33790,f33791,f33792,f33793,f33794,f33795,f33796,f33797,f33798,f33799,f33800,f33801,f33802,f33803,f33804,f33805,f33806,f33807,f33808,f33809,f33810,f33811,f33812,f33813,f33814,f33815,f33816,f33817,f33818,f33819,f33820,f33821,f33822,f33823,f33824,f33825,f33826,f33827,f33828,f33829,f33830,f33831,f33832,f33833,f33834,f33835,f33836,f33837,f33838,f33839,f33840,f33841,f33842,f33843,f33844,f33845,f33846,f33847,f33848,f33849,f33850,f33851,f33852,f33853,f33854,f33855,f33856,f33857,f33858,f33859,f33860,f33861,f33862,f33863,f33864,f33865,f33866,f33867,f33868,f33869,f33870,f33871,f33872,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33882,f33883,f33884,f33885,f33886,f33887,f33888,f33889,f33890,f33891,f33892,f33893,f33894,f33895,f33896,f33897,f33898,f33899,f33900,f33901,f33902,f33903,f33904,f33905,f33906,f33907,f33908,f33909,f33910,f33911,f33912,f33913,f33914,f33915,f33916,f33917,f33918,f33919,f33920,f33921,f33922,f33923,f33924,f33925,f33926,f33927,f33928,f33929,f33930,f33931,f33932,f33933,f33934,f33935,f33936,f33937,f33938,f33939,f33940,f33941,f33942,f33943,f33944,f33945,f33946,f33947,f33948,f33949,f33950,f33951,f33952,f33953,f33954,f33955,f33956,f33957,f33958,f33959,f33960,f33961,f33962,f33963,f33964,f33965,f33966,f33967,f33968,f33969,f33970,f33971,f33972,f33973,f33974,f33975,f33976,f33977,f33978,f33979,f33980,f33981,f33982,f33983,f33984,f33985,f33986,f33987,f33988,f33989,f45772,f45773,f45774,f28786,f45775,f53976,f54017,f54021])).
fof(f54021,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f54020,f1464])).
fof(f1464,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f936,f928])).
fof(f928,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb',aHL_TRUTH)).
fof(f936,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f747])).
fof(f747,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f54020,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f54019,f6319])).
fof(f6319,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1186])).
fof(f1186,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f684])).
fof(f684,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f683])).
fof(f683,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X45,X46] : ((p(s(bool,X45)) => p(s(bool,X46))) => ((p(s(bool,X46)) => p(s(bool,X45))) => s(bool,X45) = s(bool,X46)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f6300,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) )),
  inference(forward_demodulation,[],[f6289,f1103])).
fof(f1103,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK31(X1,X0))) )),
  inference(cnf_transformation,[],[f814])).
fof(f814,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,sK31(X1,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f406])).
fof(f406,plain,(
  ! [X0,X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0] : ! [X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X6] : ! [X19] : ? [X9] : s(X6,X9) = s(X6,X19)),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.EXISTS_REFL')).
fof(f6289,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),sK31('const.PairingHeap.Empty_0','type.PairingHeap.heap'(X0))))))) )),
  inference(unit_resulting_resolution,[],[f1103,f1086])).
fof(f1086,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1) & s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK28(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK29(X1,X0)))))) & (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ! [X2,X3] : s('type.PairingHeap.heap'(X0),X1) != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))) | p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29])],[f801])).
fof(f801,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1) & ? [X4,X5] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X4),s('type.list.list'('type.PairingHeap.heap'(X0)),X5))))) & (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ! [X2,X3] : s('type.PairingHeap.heap'(X0),X1) != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))) | p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1))))))),
  inference(rectify,[],[f800])).
fof(f800,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1) & ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))) & (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ! [X2,X3] : s('type.PairingHeap.heap'(X0),X1) != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))) | p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1))))))),
  inference(flattening,[],[f799])).
fof(f799,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1) & ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))) & ((s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ! [X2,X3] : s('type.PairingHeap.heap'(X0),X1) != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3)))) | p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1))))))),
  inference(nnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) <=> (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1) & ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3)))))),
  inference(flattening,[],[f393])).
fof(f393,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) <=> (~s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) & ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X8] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X6),X8)))) <=> (~s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),X8) & ? [X9,X10] : s('type.PairingHeap.heap'(X6),X8) = s('type.PairingHeap.heap'(X6),'const.PairingHeap.Tree_2'(s(X6,X9),s('type.list.list'('type.PairingHeap.heap'(X6)),X10)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.find_min_side_def')).
fof(f54019,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0')))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1460,f1463])).
fof(f54017,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f53975,f1462])).
fof(f53975,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f48805,f1075])).
fof(f1075,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK15(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK16(X1,X0)))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f789])).
fof(f789,plain,(
  ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK15(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK16(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f382])).
fof(f382,plain,(
  ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0] : ! [X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X18] : (s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),X18) | ? [X19,X20] : s('type.PairingHeap.heap'(X6),X18) = s('type.PairingHeap.heap'(X6),'const.PairingHeap.Tree_2'(s(X6,X19),s('type.list.list'('type.PairingHeap.heap'(X6)),X20))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.heap_nchotomy')).
fof(f48805,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK13),sK14) != s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Tree_2'(s(sK13,X0),s('type.list.list'('type.PairingHeap.heap'(sK13)),X1)))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f48804,f1462])).
fof(f48804,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK13),sK14) != s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Tree_2'(s(sK13,X0),s('type.list.list'('type.PairingHeap.heap'(sK13)),X1))) | s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14)) ) | $spl1),
  inference(subsumption_resolution,[],[f48765,f1464])).
fof(f48765,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.PairingHeap.heap'(sK13),sK14) != s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Tree_2'(s(sK13,X0),s('type.list.list'('type.PairingHeap.heap'(sK13)),X1))) | s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14)) ) | $spl1),
  inference(superposition,[],[f1088,f1672])).
fof(f1672,plain,(
  s(bool,'F_0') = s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1000])).
fof(f1000,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f751])).
fof(f751,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X36] : (s(bool,'F_0') = s(bool,X36) <=> ~p(s(bool,X36)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f1088,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | s('type.PairingHeap.heap'(X0),X1) != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f802])).
fof(f53976,plain,(
  $false | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f1462,f48805,f1075])).
fof(f45775,plain,(
  ( ! [X9] : (~p(s(bool,sK37('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) ) | $spl42),
  inference(cnf_transformation,[],[f45775_D])).
fof(f45775_D,plain,(
  ( ! [X9] : (~p(s(bool,sK37('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f28786,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl36),
  inference(cnf_transformation,[],[f28786_D])).
fof(f28786_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f45774,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | ~p(s(bool,sK37('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f45748,f1470])).
fof(f1470,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f1066,f1059])).
fof(f1059,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f349])).
fof(f349,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X9] : s(X6,X9) = s(X6,'const.combin.I_1'(s(X6,X9)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.combin.I_THM')).
fof(f1066,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X34] : s(X6,'const.combin.I_1'(s(X6,X34))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X34)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','ARITY_0_1_const.combin.I')).
fof(f45748,plain,(
  ( ! [X10,X9] : (~p(s(bool,sK37('const.combin.I_0',X9,bool))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X10)))) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f1128,f1470])).
fof(f1128,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f826])).
fof(f826,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36,sK37])],[f825])).
fof(f825,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f824])).
fof(f824,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f823])).
fof(f823,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f437])).
fof(f437,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X6] : ! [X32,X56] : ((p(s(bool,X32)) & ! [X9] : p(s(bool,happ(s(fun(X6,bool),X56),s(X6,X9))))) <=> ! [X9] : (p(s(bool,X32)) & p(s(bool,happ(s(fun(X6,bool),X56),s(X6,X9))))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f45773,plain,(
  ( ! [X7] : (~p(s(bool,sK37('_i_0_0',X7,bool))) | ~p(s(bool,X7))) ) | $spl40),
  inference(cnf_transformation,[],[f45773_D])).
fof(f45773_D,plain,(
  ( ! [X7] : (~p(s(bool,sK37('_i_0_0',X7,bool))) | ~p(s(bool,X7))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f45772,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | ~p(s(bool,sK37('_i_0_0',X7,bool))) | ~p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f45747,f1061])).
fof(f1061,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X41] : ! [X43] : s(X41,X43) = s(X41,happ(s(fun(X41,X41),'_i_0_0'),s(X41,X43)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','_i0_ax')).
fof(f45747,plain,(
  ( ! [X8,X7] : (~p(s(bool,sK37('_i_0_0',X7,bool))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X8)))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f1128,f1061])).
fof(f33989,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28794,f32168])).
fof(f32168,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1304])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f881])).
fof(f881,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f256])).
fof(f256,axiom,(
  ! [X88,X89,X90] : (p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),'=_0'),s(X88,X89))),s(X88,X90)))) <=> s(X88,X89) = s(X88,X90))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb',arity_0_HL_EQ)).
fof(f28794,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f999])).
fof(f999,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f751])).
fof(f33988,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28799,f928])).
fof(f28799,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f1189])).
fof(f1189,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f684])).
fof(f1621,plain,(
  s(bool,'T_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f995])).
fof(f995,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f749])).
fof(f749,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X36] : (s(bool,'T_0') = s(bool,X36) <=> p(s(bool,X36)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f33987,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28801,f928])).
fof(f28801,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1622,f28786,f1189])).
fof(f1622,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1464,f995])).
fof(f33986,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28803,f928])).
fof(f28803,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1623,f28786,f1189])).
fof(f1623,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f940,f995])).
fof(f940,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f129])).
fof(f129,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|34')).
fof(f33985,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28805,f28786])).
fof(f28805,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f1189])).
fof(f33984,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28809,f28786])).
fof(f28809,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1623,f28786,f1189])).
fof(f33983,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28813,f928])).
fof(f28813,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f1189])).
fof(f33982,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28815,f928])).
fof(f28815,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1622,f28786,f1189])).
fof(f33981,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28817,f928])).
fof(f28817,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1623,f28786,f1189])).
fof(f33980,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28819,f28786])).
fof(f28819,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f1189])).
fof(f33979,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28823,f28786])).
fof(f28823,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1623,f28786,f1189])).
fof(f33978,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28830,f28786])).
fof(f28830,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1806,f28786,f1216])).
fof(f1216,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f851])).
fof(f851,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f850])).
fof(f850,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f849])).
fof(f849,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f727])).
fof(f727,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1806,plain,(
  ( ! [X0,X1] : (sP1('T_0','const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f1102,f1217])).
fof(f1217,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f851])).
fof(f1102,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  inference(flattening,[],[f403])).
fof(f403,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  inference(rectify,[],[f236])).
fof(f236,axiom,(
  ! [X6] : ! [X9] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X6,X9),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bag.NOT_IN_EMPTY_BAG')).
fof(f33977,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28831,f28786])).
fof(f28831,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6325,f28786,f1216])).
fof(f6325,plain,(
  ( ! [X0] : (sP1('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1217])).
fof(f33976,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28832,f928])).
fof(f28832,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1716,f28786,f1216])).
fof(f1716,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1464,f928,f1220])).
fof(f1220,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f851])).
fof(f33975,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28833,f28786])).
fof(f28833,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1894,f28786,f1216])).
fof(f1894,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f1464,f1721])).
fof(f1721,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,'const.bool.DATATYPE_1'(s(X0,X1))) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f1717,f928])).
fof(f1717,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1(X2,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(superposition,[],[f1220,f1060])).
fof(f1060,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f351])).
fof(f351,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X9] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X6,X9)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.DATATYPE_TAG_THM')).
fof(f33974,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28834,f28786])).
fof(f28834,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19863,f28786,f1216])).
fof(f19863,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1220])).
fof(f19033,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f17817,f1474])).
fof(f1474,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f1113,f1069])).
fof(f1069,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f369])).
fof(f369,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f243])).
fof(f243,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.EXISTS_DEF')).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f674])).
fof(f674,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f425])).
fof(f425,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X32,X9] : (p(s(bool,happ(s(fun(X6,bool),X32),s(X6,X9)))) => p(s(bool,happ(s(fun(X6,bool),X32),s(X6,'const.min.@_1'(s(fun(X6,bool),X32)))))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.SELECT_AX')).
fof(f17817,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'F_0')))) | $spl1),
  inference(forward_demodulation,[],[f17590,f5716])).
fof(f5716,plain,(
  s(bool,'F_0') = s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(forward_demodulation,[],[f5710,f1059])).
fof(f5710,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5390,f1343])).
fof(f1343,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f900])).
fof(f900,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f735])).
fof(f735,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f618,f734,f733,f732,f731])).
fof(f731,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f732,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f733,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f734,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f618,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f617])).
fof(f617,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X66,X67,X68] : ((p(s(bool,X68)) <=> s(bool,X66) = s(bool,X67)) <=> ((p(s(bool,X68)) | p(s(bool,X67)) | p(s(bool,X66))) & (p(s(bool,X68)) | ~p(s(bool,X66)) | ~p(s(bool,X67))) & (p(s(bool,X67)) | ~p(s(bool,X66)) | ~p(s(bool,X68))) & (p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,X68)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f5390,plain,(
  sP6('T_0','const.combin.I_1'(s(bool,'F_0')),sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(superposition,[],[f4102,f5343])).
fof(f5343,plain,(
  s(bool,'F_0') = s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(forward_demodulation,[],[f5337,f1059])).
fof(f5337,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5270,f1343])).
fof(f5270,plain,(
  sP6('T_0','const.combin.I_1'(s(bool,'F_0')),sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(superposition,[],[f4102,f1672])).
fof(f4102,plain,(
  ( ! [X21] : (sP6('T_0','const.combin.I_1'(s(bool,X21)),sK31(X21,bool))) )),
  inference(superposition,[],[f3486,f1103])).
fof(f3486,plain,(
  ( ! [X0] : (sP6('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1059,f1342])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f900])).
fof(f17590,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)))),s(bool,'F_0')))) | $spl1),
  inference(unit_resulting_resolution,[],[f5716,f1303])).
fof(f1303,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f880])).
fof(f880,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f613])).
fof(f613,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f257])).
fof(f257,axiom,(
  ! [X88,X89,X90] : (p(s(bool,happ(s(fun(X88,bool),'=_1'(s(X88,X89))),s(X88,X90)))) <=> s(X88,X89) = s(X88,X90))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb',arity_1_HL_EQ)).
fof(f33973,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28835,f28786])).
fof(f28835,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1925,f28786,f1216])).
fof(f1925,plain,(
  sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f1464,f1745])).
fof(f1745,plain,(
  ( ! [X3] : (sP1(X3,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f1732,f928])).
fof(f1732,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1220,f1639])).
fof(f1639,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f930,f996])).
fof(f996,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f749])).
fof(f930,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f127])).
fof(f127,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|35')).
fof(f33972,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28836,f928])).
fof(f28836,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6326,f28786,f1216])).
fof(f6326,plain,(
  ( ! [X0] : (sP1('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1220])).
fof(f33971,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28837,f928])).
fof(f28837,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1807,f28786,f1216])).
fof(f1807,plain,(
  ( ! [X0,X1] : (sP1('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1102,f1220])).
fof(f33970,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28845,f928])).
fof(f28845,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1704,f28786,f1216])).
fof(f1704,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1464,f928,f1217])).
fof(f33969,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28846,f928])).
fof(f28846,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1806,f28786,f1216])).
fof(f33968,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28847,f928])).
fof(f28847,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6325,f28786,f1216])).
fof(f33967,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28852,f28786])).
fof(f28852,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6326,f28786,f1216])).
fof(f33966,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28853,f28786])).
fof(f28853,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1807,f28786,f1216])).
fof(f33965,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28854,f28786])).
fof(f28854,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1870,f28786,f1216])).
fof(f1870,plain,(
  ( ! [X0,X1] : (sP1('const.bool.DATATYPE_1'(s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1709])).
fof(f1709,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f1705,f928])).
fof(f1705,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f1217,f1060])).
fof(f33964,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28855,f28786])).
fof(f28855,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19862,f28786,f1216])).
fof(f19862,plain,(
  sP1('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1217])).
fof(f33963,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28856,f28786])).
fof(f28856,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1920,f28786,f1216])).
fof(f1920,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1744])).
fof(f1744,plain,(
  ( ! [X2] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f1731,f928])).
fof(f1731,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X2)) )),
  inference(superposition,[],[f1217,f1639])).
fof(f33962,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28885,f28786])).
fof(f28885,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6327,f28786,f1309])).
fof(f1309,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f884])).
fof(f884,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f883])).
fof(f883,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f882])).
fof(f882,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f729])).
fof(f729,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f6327,plain,(
  ( ! [X0,X1] : (~sP2(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1306])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f884])).
fof(f33961,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28886,f28786])).
fof(f28886,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6342,f28786,f1309])).
fof(f6342,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1307])).
fof(f1307,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f884])).
fof(f33960,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28891,f28786])).
fof(f28891,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6330,f28786,f1309])).
fof(f6330,plain,(
  ( ! [X0,X1] : (~sP2(X0,'F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1306])).
fof(f33959,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28892,f28786])).
fof(f28892,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1905,f28786,f1309])).
fof(f1905,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1307])).
fof(f33958,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28893,f28786])).
fof(f28893,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6336,f28786,f1309])).
fof(f6336,plain,(
  ( ! [X0,X1] : (~sP2('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1307])).
fof(f33957,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28895,f28786])).
fof(f28895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6342,f28786,f1309])).
fof(f33956,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28904,f28786])).
fof(f28904,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6327,f28786,f1310])).
fof(f1310,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f884])).
fof(f33955,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28905,f28786])).
fof(f28905,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6342,f28786,f1310])).
fof(f33954,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28909,f28786])).
fof(f28909,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f1310])).
fof(f1891,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1306])).
fof(f33953,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28910,f28786])).
fof(f28910,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6330,f28786,f1310])).
fof(f33952,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28912,f28786])).
fof(f28912,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6336,f28786,f1310])).
fof(f33951,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28913,f28786])).
fof(f28913,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6327,f28786,f1310])).
fof(f33950,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28920,f28786])).
fof(f28920,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6330,f28786,f1311])).
fof(f1311,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f884])).
fof(f33949,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28921,f28786])).
fof(f28921,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1905,f28786,f1311])).
fof(f33948,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28922,f28786])).
fof(f28922,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6336,f28786,f1311])).
fof(f33947,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28924,f28786])).
fof(f28924,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6342,f28786,f1311])).
fof(f33946,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28929,f28786])).
fof(f28929,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f1311])).
fof(f33945,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28930,f28786])).
fof(f28930,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6330,f28786,f1311])).
fof(f33944,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28932,f28786])).
fof(f28932,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6336,f28786,f1311])).
fof(f33943,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28933,f28786])).
fof(f28933,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6327,f28786,f1311])).
fof(f33942,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28940,f28786])).
fof(f28940,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3481,f1464,f28786,f1325])).
fof(f1325,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f891])).
fof(f891,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f890])).
fof(f890,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f889])).
fof(f889,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f734])).
fof(f3481,plain,(
  sP6('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1672,f1342])).
fof(f33941,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28941,f28786])).
fof(f28941,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f3486,f1464,f28786,f1325])).
fof(f33940,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28942,f28786])).
fof(f28942,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f3482,f1464,f28786,f1325])).
fof(f3482,plain,(
  sP6('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f928,f1674,f1342])).
fof(f1674,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f940,f1000])).
fof(f33939,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28943,f28786])).
fof(f28943,plain,(
  ~p(s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5366,f1464,f28786,f1325])).
fof(f5366,plain,(
  sP6('T_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5343,f1342])).
fof(f33938,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28944,f28786])).
fof(f28944,plain,(
  ~p(s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5805,f1464,f28786,f1325])).
fof(f5805,plain,(
  sP6('T_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5716,f1342])).
fof(f33937,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28945,f28786])).
fof(f28945,plain,(
  ~p(s(bool,sK31('F_0',bool))) | $spl36),
  inference(unit_resulting_resolution,[],[f3490,f1464,f28786,f1325])).
fof(f3490,plain,(
  ( ! [X0] : (sP6('T_0',sK31(X0,bool),X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1103,f1342])).
fof(f33936,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28946,f28786])).
fof(f28946,plain,(
  ~p(s(bool,sK40('F_0','=_1'(s(bool,'F_0')),bool))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28255,f1464,f28786,f1325])).
fof(f28255,plain,(
  sP6('T_0',sK40('F_0','=_1'(s(bool,'F_0')),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f19032,f1342])).
fof(f19032,plain,(
  s(bool,'F_0') = s(bool,sK40('F_0','=_1'(s(bool,'F_0')),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f17817,f1134])).
fof(f1134,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK40(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f833])).
fof(f833,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK40(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40])],[f832])).
fof(f832,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f831])).
fof(f831,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f441])).
fof(f441,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f179])).
fof(f179,axiom,(
  ! [X6] : ! [X32,X19] : (? [X9] : (s(X6,X9) = s(X6,X19) & p(s(bool,happ(s(fun(X6,bool),X32),s(X6,X9))))) <=> p(s(bool,happ(s(fun(X6,bool),X32),s(X6,X19)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.UNWIND_THM2')).
fof(f33935,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28947,f928])).
fof(f28947,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f3468,f1464,f28786,f1325])).
fof(f3468,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1622,f1341])).
fof(f1341,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f900])).
fof(f33934,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f28948,f28786])).
fof(f28948,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f22412,f1464,f28786,f1325])).
fof(f22412,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1517,f1512,f19981,f1328])).
fof(f1328,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | p(s(bool,X0)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f891])).
fof(f19981,plain,(
  ( ! [X0,X1] : (sP3(X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1339])).
fof(f1339,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f899])).
fof(f899,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f898])).
fof(f898,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f731])).
fof(f1512,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1332])).
fof(f1332,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f894])).
fof(f894,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f893])).
fof(f893,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f892])).
fof(f892,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f733])).
fof(f1517,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1335])).
fof(f1335,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f897])).
fof(f897,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f896])).
fof(f896,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f895])).
fof(f895,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f732])).
fof(f33933,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28949,f28786])).
fof(f28949,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f11400,f1464,f28786,f1325])).
fof(f11400,plain,(
  sP6('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1517,f1509,f1525,f1464,f1328])).
fof(f1525,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f930,f1339])).
fof(f1509,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1331])).
fof(f1331,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f894])).
fof(f33932,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f28950,f928])).
fof(f28950,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6418,f1464,f28786,f1325])).
fof(f6418,plain,(
  ( ! [X0] : (sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1622,f6300,f1341])).
fof(f33931,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29132,f928])).
fof(f29132,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f11162,f1464,f28786,f1325])).
fof(f11162,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1520,f1509,f1522,f1464,f1327])).
fof(f1327,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | p(s(bool,X1)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f891])).
fof(f1522,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1338])).
fof(f1338,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f899])).
fof(f1520,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1336])).
fof(f1336,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f897])).
fof(f33930,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29133,f928])).
fof(f29133,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3481,f1464,f28786,f1325])).
fof(f33929,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29134,f928])).
fof(f29134,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f3486,f1464,f28786,f1325])).
fof(f33928,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29135,f928])).
fof(f29135,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f3482,f1464,f28786,f1325])).
fof(f33927,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29136,f928])).
fof(f29136,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5366,f1464,f28786,f1325])).
fof(f33926,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29137,f928])).
fof(f29137,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5805,f1464,f28786,f1325])).
fof(f33925,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29138,f928])).
fof(f29138,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f3490,f1464,f28786,f1325])).
fof(f33924,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29139,f928])).
fof(f29139,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28255,f1464,f28786,f1325])).
fof(f33923,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29143,f28786])).
fof(f29143,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6418,f1464,f28786,f1325])).
fof(f33922,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29144,f28786])).
fof(f29144,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f22179,f1464,f28786,f1325])).
fof(f22179,plain,(
  sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1520,f1509,f19980,f1327])).
fof(f19980,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1338])).
fof(f33921,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29145,f28786])).
fof(f29145,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f11171,f1464,f28786,f1325])).
fof(f11171,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1520,f1509,f1523,f1464,f1327])).
fof(f1523,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f930,f1338])).
fof(f33920,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29150,f30574])).
fof(f30574,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1340])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f899])).
fof(f29150,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f28786,f1326])).
fof(f1326,plain,(
  ( ! [X2,X0,X1] : (~sP3(X2,X1,X0) | ~p(s(bool,X2)) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f891])).
fof(f1506,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1330])).
fof(f1330,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f894])).
fof(f1514,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1334])).
fof(f1334,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f897])).
fof(f3518,plain,(
  ~sP6('T_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1621,f1343])).
fof(f33919,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29151,f30574])).
fof(f29151,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1506,f28786,f1326])).
fof(f6345,plain,(
  ( ! [X0] : (~sP6('T_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f928,f6300,f1325])).
fof(f33918,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29152,f30574])).
fof(f29152,plain,(
  ~sP3('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1506,f28786,f1326])).
fof(f3520,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f928,f1622,f1343])).
fof(f33917,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29153,f30574])).
fof(f29153,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1506,f28786,f1326])).
fof(f3521,plain,(
  ~sP6('T_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f928,f1623,f1343])).
fof(f33916,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29154,f30574])).
fof(f29154,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1506,f28786,f1326])).
fof(f19979,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1334])).
fof(f20060,plain,(
  ~sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1622,f19033,f1343])).
fof(f33915,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29155,f30574])).
fof(f29155,plain,(
  ~sP3('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f28786,f1326])).
fof(f33914,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29156,f30574])).
fof(f29156,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f1506,f28786,f1326])).
fof(f33913,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29157,f30574])).
fof(f29157,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f28786,f1326])).
fof(f1516,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1335])).
fof(f33912,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29158,f30574])).
fof(f29158,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f28786,f1326])).
fof(f6370,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1335])).
fof(f33911,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29159,f30574])).
fof(f29159,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f28786,f1326])).
fof(f1518,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1335])).
fof(f33910,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29160,f30574])).
fof(f29160,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f19954,f28786,f1326])).
fof(f19954,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1330])).
fof(f19890,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1464,f19033,f1325])).
fof(f33909,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29161,f30574])).
fof(f29161,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f19954,f28786,f1326])).
fof(f33908,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29162,f30574])).
fof(f29162,plain,(
  ~sP3('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f28786,f1326])).
fof(f33907,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29163,f30574])).
fof(f29163,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f1509,f28786,f1326])).
fof(f33906,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29172,f30574])).
fof(f29172,plain,(
  ( ! [X0] : (~sP3('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6425,f1517,f1509,f28786,f1326])).
fof(f6425,plain,(
  ( ! [X0] : (~sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1672,f6300,f1344])).
fof(f1344,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f900])).
fof(f33905,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29173,f30574])).
fof(f29173,plain,(
  ( ! [X0] : (~sP3('F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6431,f1517,f1509,f28786,f1326])).
fof(f6431,plain,(
  ( ! [X0] : (~sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5343,f6300,f1344])).
fof(f33904,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29174,f30574])).
fof(f29174,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f1509,f28786,f1326])).
fof(f33903,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29175,f30574])).
fof(f29175,plain,(
  ~sP3('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f28786,f1326])).
fof(f33902,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29176,f30574])).
fof(f29176,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f1509,f28786,f1326])).
fof(f33901,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29177,f30574])).
fof(f29177,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1509,f28786,f1326])).
fof(f33900,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29180,f30574])).
fof(f29180,plain,(
  ( ! [X0] : (~sP3('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6425,f1519,f1509,f28786,f1326])).
fof(f1519,plain,(
  ( ! [X0,X1] : (sP4('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1336])).
fof(f33899,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29182,f30574])).
fof(f29182,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f28786,f1326])).
fof(f1508,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1331])).
fof(f33898,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29185,f30574])).
fof(f29185,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3539,f1516,f1508,f28786,f1326])).
fof(f3539,plain,(
  ( ! [X0] : (~sP6('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1464,f1059,f1344])).
fof(f33897,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29188,f30574])).
fof(f29188,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f28786,f1326])).
fof(f33896,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29191,f30574])).
fof(f29191,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f28786,f1326])).
fof(f6359,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1331])).
fof(f33895,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29192,f30574])).
fof(f29192,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f6370,f6359,f28786,f1326])).
fof(f33894,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29194,f30574])).
fof(f29194,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f28786,f1326])).
fof(f33893,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29195,f30574])).
fof(f29195,plain,(
  ( ! [X0,X1] : (~sP3('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'const.combin.I_1'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1810,f1808,f28786,f1326])).
fof(f1808,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,'const.bag.BAG_IN_2'(s(X1,X2),s(fun(X1,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1331])).
fof(f1810,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bag.BAG_IN_2'(s(X1,X2),s(fun(X1,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1335])).
fof(f33892,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29197,f30574])).
fof(f29197,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f28786,f1326])).
fof(f1510,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1331])).
fof(f33891,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29199,f30574])).
fof(f29199,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1518,f1510,f28786,f1326])).
fof(f33890,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29201,f30574])).
fof(f29201,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f28786,f1326])).
fof(f33889,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29213,f30574])).
fof(f29213,plain,(
  ~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3539,f1516,f1512,f28786,f1326])).
fof(f33888,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29215,f30574])).
fof(f29215,plain,(
  ( ! [X0] : (~sP3('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f6370,f1512,f28786,f1326])).
fof(f33887,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29217,f30574])).
fof(f29217,plain,(
  ( ! [X0,X1] : (~sP3('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'const.combin.I_1'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1810,f1512,f28786,f1326])).
fof(f33886,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29220,f30574])).
fof(f29220,plain,(
  ~sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1518,f1512,f28786,f1326])).
fof(f33885,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29233,f30574])).
fof(f29233,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3555,f1519,f1512,f28786,f1326])).
fof(f3555,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1464,f1059,f1344])).
fof(f33884,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29235,f30574])).
fof(f29235,plain,(
  ( ! [X0] : (~sP3('const.combin.I_1'(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f6371,f1512,f28786,f1326])).
fof(f6371,plain,(
  ( ! [X2,X0,X1] : (sP4('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1336])).
fof(f33883,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29237,f30574])).
fof(f29237,plain,(
  ( ! [X0,X1] : (~sP3('const.combin.I_1'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))),'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1811,f1512,f28786,f1326])).
fof(f1811,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1336])).
fof(f33882,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29240,f30574])).
fof(f29240,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1521,f1512,f28786,f1326])).
fof(f1521,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1336])).
fof(f33881,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29241,f30574])).
fof(f29241,plain,(
  ( ! [X0] : (~sP3('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6425,f1517,f6360,f28786,f1326])).
fof(f6360,plain,(
  ( ! [X2,X0,X1] : (sP5('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1332])).
fof(f33880,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29242,f30574])).
fof(f29242,plain,(
  ( ! [X0] : (~sP3('F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6431,f1517,f6360,f28786,f1326])).
fof(f33879,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29246,f30574])).
fof(f29246,plain,(
  ( ! [X0] : (~sP3('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6425,f1519,f6360,f28786,f1326])).
fof(f33878,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29247,f30488])).
fof(f30488,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1330])).
fof(f29247,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1522,f28786,f1326])).
fof(f33877,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29248,f30488])).
fof(f29248,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1522,f28786,f1326])).
fof(f33876,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29249,f30488])).
fof(f29249,plain,(
  ~sP5('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1522,f28786,f1326])).
fof(f33875,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29250,f30488])).
fof(f29250,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1522,f28786,f1326])).
fof(f33874,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29254,f30488])).
fof(f29254,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f1522,f28786,f1326])).
fof(f33873,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29256,f30488])).
fof(f29256,plain,(
  ~sP5('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1522,f28786,f1326])).
fof(f33872,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29257,f30488])).
fof(f29257,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f1522,f28786,f1326])).
fof(f33871,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29258,f30488])).
fof(f29258,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1522,f28786,f1326])).
fof(f33870,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29259,f30488])).
fof(f29259,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1522,f28786,f1326])).
fof(f33869,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29260,f30488])).
fof(f29260,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1522,f28786,f1326])).
fof(f33868,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29298,f30568])).
fof(f30568,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1334])).
fof(f29298,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1522,f28786,f1326])).
fof(f33867,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29299,f30568])).
fof(f29299,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1522,f28786,f1326])).
fof(f33866,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29300,f30568])).
fof(f29300,plain,(
  ~sP4('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1522,f28786,f1326])).
fof(f33865,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29301,f30568])).
fof(f29301,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1522,f28786,f1326])).
fof(f33864,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29325,f30568])).
fof(f29325,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f19954,f1522,f28786,f1326])).
fof(f33863,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29340,f30568])).
fof(f29340,plain,(
  ~sP4('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1522,f28786,f1326])).
fof(f33862,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29341,f30568])).
fof(f29341,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1509,f1522,f28786,f1326])).
fof(f33861,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29351,f30568])).
fof(f29351,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1522,f28786,f1326])).
fof(f33860,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29361,f30568])).
fof(f29361,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1522,f28786,f1326])).
fof(f33859,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29378,f30568])).
fof(f29378,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1522,f28786,f1326])).
fof(f33858,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29518,f30488])).
fof(f29518,plain,(
  ~sP5('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f19980,f28786,f1326])).
fof(f33857,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29520,f30488])).
fof(f29520,plain,(
  ~sP5('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f19980,f28786,f1326])).
fof(f33856,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29522,f30568])).
fof(f29522,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1506,f19980,f28786,f1326])).
fof(f33855,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29555,f30568])).
fof(f29555,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1509,f19980,f28786,f1326])).
fof(f33854,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29658,f30488])).
fof(f29658,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1524,f28786,f1326])).
fof(f1524,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1339])).
fof(f33853,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29659,f30488])).
fof(f29659,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1524,f28786,f1326])).
fof(f33852,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29660,f30488])).
fof(f29660,plain,(
  ~sP5('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1524,f28786,f1326])).
fof(f33851,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29661,f30488])).
fof(f29661,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1524,f28786,f1326])).
fof(f33850,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29662,f30488])).
fof(f29662,plain,(
  ~sP5('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1524,f28786,f1326])).
fof(f33849,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29663,f30488])).
fof(f29663,plain,(
  ~sP5('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1524,f28786,f1326])).
fof(f33848,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29664,f30488])).
fof(f29664,plain,(
  ~sP5('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f1524,f28786,f1326])).
fof(f33847,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29665,f30488])).
fof(f29665,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1524,f28786,f1326])).
fof(f33846,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29666,f30488])).
fof(f29666,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1524,f28786,f1326])).
fof(f33845,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29667,f30488])).
fof(f29667,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1524,f28786,f1326])).
fof(f33844,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29668,f30568])).
fof(f29668,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1524,f28786,f1326])).
fof(f33843,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29669,f30568])).
fof(f29669,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1524,f28786,f1326])).
fof(f33842,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29670,f30568])).
fof(f29670,plain,(
  ~sP4('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1524,f28786,f1326])).
fof(f33841,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29671,f30568])).
fof(f29671,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1524,f28786,f1326])).
fof(f33840,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29678,f30568])).
fof(f29678,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1506,f1524,f28786,f1326])).
fof(f33839,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29686,f30799])).
fof(f30799,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1622,f28786,f1343])).
fof(f29686,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1506,f1524,f28786,f1326])).
fof(f33838,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29688,f30797])).
fof(f30797,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f1343])).
fof(f29688,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1506,f1524,f28786,f1326])).
fof(f33837,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29693,f30801])).
fof(f30801,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1623,f28786,f1343])).
fof(f29693,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1506,f1524,f28786,f1326])).
fof(f33836,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29695,f30568])).
fof(f29695,plain,(
  ~sP4('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1524,f28786,f1326])).
fof(f33835,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29696,f30568])).
fof(f29696,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1509,f1524,f28786,f1326])).
fof(f33834,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29697,f30799])).
fof(f29697,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1509,f1524,f28786,f1326])).
fof(f33833,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29699,f3520])).
fof(f29699,plain,(
  sP6('T_0','T_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1514,f1509,f1524,f28786,f1326])).
fof(f33832,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29701,f30799])).
fof(f29701,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'T_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1586,f1509,f1524,f28786,f1326])).
fof(f1586,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f932,f1334])).
fof(f932,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(cnf_transformation,[],[f217])).
fof(f217,axiom,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.prim_rec.WF_LESS')).
fof(f33831,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29702,f30799])).
fof(f29702,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1602,f1509,f1524,f28786,f1326])).
fof(f1602,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f943,f1334])).
fof(f943,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.relation.WF_EMPTY_REL')).
fof(f33830,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29703,f20060])).
fof(f29703,plain,(
  sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19979,f1509,f1524,f28786,f1326])).
fof(f33829,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29705,f30799])).
fof(f29705,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1515,f1509,f1524,f28786,f1326])).
fof(f1515,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f930,f1334])).
fof(f33828,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29706,f30568])).
fof(f29706,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1524,f28786,f1326])).
fof(f33827,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29707,f30797])).
fof(f29707,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1508,f1524,f28786,f1326])).
fof(f33826,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29709,f3518])).
fof(f29709,plain,(
  sP6('T_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1514,f1508,f1524,f28786,f1326])).
fof(f33825,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29711,f30797])).
fof(f29711,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1586,f1508,f1524,f28786,f1326])).
fof(f33824,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29712,f30797])).
fof(f29712,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1602,f1508,f1524,f28786,f1326])).
fof(f33823,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29713,f20059])).
fof(f20059,plain,(
  ~sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1621,f19033,f1343])).
fof(f29713,plain,(
  sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19979,f1508,f1524,f28786,f1326])).
fof(f33822,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29714,f30797])).
fof(f29714,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1515,f1508,f1524,f28786,f1326])).
fof(f33821,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29715,f30568])).
fof(f29715,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1524,f28786,f1326])).
fof(f33820,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29718,f6345])).
fof(f29718,plain,(
  ( ! [X0] : (sP6('T_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1514,f6359,f1524,f28786,f1326])).
fof(f33819,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29730,f30568])).
fof(f29730,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1524,f28786,f1326])).
fof(f33818,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29731,f30801])).
fof(f29731,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1510,f1524,f28786,f1326])).
fof(f33817,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29733,f3521])).
fof(f29733,plain,(
  sP6('T_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f1514,f1510,f1524,f28786,f1326])).
fof(f33816,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29735,f30801])).
fof(f29735,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f1586,f1510,f1524,f28786,f1326])).
fof(f33815,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29736,f30801])).
fof(f29736,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1602,f1510,f1524,f28786,f1326])).
fof(f33814,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29737,f20061])).
fof(f20061,plain,(
  ~sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(unit_resulting_resolution,[],[f1623,f19033,f1343])).
fof(f29737,plain,(
  sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19979,f1510,f1524,f28786,f1326])).
fof(f33813,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29738,f30801])).
fof(f29738,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f1515,f1510,f1524,f28786,f1326])).
fof(f33812,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29752,f30799])).
fof(f29752,plain,(
  sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1517,f1511,f1524,f28786,f1326])).
fof(f1511,plain,(
  ( ! [X0,X1] : (sP5('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1332])).
fof(f33811,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29753,f30797])).
fof(f29753,plain,(
  sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1511,f1524,f28786,f1326])).
fof(f33810,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29756,f30801])).
fof(f29756,plain,(
  sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1518,f1511,f1524,f28786,f1326])).
fof(f33809,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29762,f30799])).
fof(f29762,plain,(
  ( ! [X0,X1] : (sP6('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1809,f1524,f28786,f1326])).
fof(f1809,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1332])).
fof(f33808,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29763,f30797])).
fof(f29763,plain,(
  ( ! [X0,X1] : (sP6('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1809,f1524,f28786,f1326])).
fof(f33807,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29766,f30801])).
fof(f29766,plain,(
  ( ! [X0,X1] : (sP6('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1809,f1524,f28786,f1326])).
fof(f33806,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29767,f30799])).
fof(f29767,plain,(
  sP6('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1513,f1524,f28786,f1326])).
fof(f1513,plain,(
  ( ! [X0,X1] : (sP5('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1332])).
fof(f33805,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29768,f30797])).
fof(f29768,plain,(
  sP6('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1513,f1524,f28786,f1326])).
fof(f33804,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29771,f30801])).
fof(f29771,plain,(
  sP6('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1513,f1524,f28786,f1326])).
fof(f33803,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29773,f30568])).
fof(f29773,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1585,f1588,f28786,f1326])).
fof(f1588,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f932,f1339])).
fof(f1585,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f932,f1330])).
fof(f33802,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29817,f30568])).
fof(f29817,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1588,f28786,f1326])).
fof(f33801,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29844,f30568])).
fof(f29844,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1601,f1604,f28786,f1326])).
fof(f1604,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f943,f1339])).
fof(f1601,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f943,f1330])).
fof(f33800,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29888,f30568])).
fof(f29888,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1604,f28786,f1326])).
fof(f33799,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29914,f30488])).
fof(f29914,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f19981,f28786,f1326])).
fof(f33798,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29915,f30488])).
fof(f29915,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f19981,f28786,f1326])).
fof(f33797,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29916,f30568])).
fof(f29916,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f19954,f19981,f28786,f1326])).
fof(f33796,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29919,f30568])).
fof(f29919,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3555,f19954,f19981,f28786,f1326])).
fof(f33795,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29932,f30568])).
fof(f29932,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1509,f19981,f28786,f1326])).
fof(f33794,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29935,f19890])).
fof(f29935,plain,(
  sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1514,f1509,f19981,f28786,f1326])).
fof(f33793,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29939,f19891])).
fof(f19891,plain,(
  ~sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1464,f19033,f1325])).
fof(f29939,plain,(
  sP6('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19979,f1509,f19981,f28786,f1326])).
fof(f33792,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f29969,f30568])).
fof(f29969,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3555,f1512,f19981,f28786,f1326])).
fof(f33791,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f29997,f30568])).
fof(f29997,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1507,f1525,f28786,f1326])).
fof(f1507,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f930,f1330])).
fof(f33790,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30042,f30568])).
fof(f30042,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1525,f28786,f1326])).
fof(f33789,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30165,f30568])).
fof(f30165,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1589,f28786,f1326])).
fof(f1589,plain,(
  ( ! [X0,X1] : (sP3('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f932,f1340])).
fof(f33788,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30219,f30568])).
fof(f30219,plain,(
  ( ! [X0] : (~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1605,f28786,f1326])).
fof(f1605,plain,(
  ( ! [X2,X0,X1] : (sP3('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f943,f1340])).
fof(f33787,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30277,f30568])).
fof(f30277,plain,(
  ~sP4('const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))),'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3539,f1512,f19982,f28786,f1326])).
fof(f19982,plain,(
  ( ! [X0,X1] : (sP3('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1340])).
fof(f33786,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30333,f30568])).
fof(f30333,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1527,f28786,f1326])).
fof(f1527,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f930,f1340])).
fof(f33785,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30361,f28786])).
fof(f30361,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1508,f1464,f28786,f1329])).
fof(f1329,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f894])).
fof(f33784,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30362,f28786])).
fof(f30362,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6359,f1464,f28786,f1329])).
fof(f33783,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30363,f28786])).
fof(f30363,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1808,f1464,f28786,f1329])).
fof(f33782,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30364,f28786])).
fof(f30364,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1510,f1464,f28786,f1329])).
fof(f33781,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30365,f28786])).
fof(f30365,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f4299,f1464,f28786,f1329])).
fof(f4299,plain,(
  ( ! [X0] : (sP5('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f3501,f1324])).
fof(f1324,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f891])).
fof(f3501,plain,(
  ( ! [X0] : (sP6('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f928,f1059,f1342])).
fof(f33780,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30366,f28786])).
fof(f30366,plain,(
  ~p(s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5458,f1464,f28786,f1329])).
fof(f5458,plain,(
  sP5('T_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f5369,f1324])).
fof(f5369,plain,(
  sP6('T_0','F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5343,f1342])).
fof(f33779,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30367,f28786])).
fof(f30367,plain,(
  ~p(s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6009,f1464,f28786,f1329])).
fof(f6009,plain,(
  sP5('T_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f5808,f1324])).
fof(f5808,plain,(
  sP6('T_0','F_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5716,f1342])).
fof(f33778,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30368,f28786])).
fof(f30368,plain,(
  ~p(s(bool,sK31('F_0',bool))) | $spl36),
  inference(unit_resulting_resolution,[],[f3819,f1464,f28786,f1329])).
fof(f3819,plain,(
  ( ! [X0] : (sP5('T_0',sK31(X0,bool),X0)) )),
  inference(unit_resulting_resolution,[],[f3505,f1324])).
fof(f3505,plain,(
  ( ! [X0] : (sP6('T_0',X0,sK31(X0,bool))) )),
  inference(unit_resulting_resolution,[],[f928,f1103,f1342])).
fof(f33777,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30369,f28786])).
fof(f30369,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1512,f1464,f28786,f1329])).
fof(f33776,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30370,f28786])).
fof(f30370,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1511,f1464,f28786,f1329])).
fof(f33775,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30371,f28786])).
fof(f30371,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6360,f1464,f28786,f1329])).
fof(f33774,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30372,f28786])).
fof(f30372,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1809,f1464,f28786,f1329])).
fof(f33773,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30373,f28786])).
fof(f30373,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1513,f1464,f28786,f1329])).
fof(f33772,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30387,f30488])).
fof(f30387,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f928,f28786,f1329])).
fof(f33771,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30417,f30488])).
fof(f30417,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f28786,f1329])).
fof(f33770,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30470,f28786])).
fof(f30470,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1509,f1464,f28786,f1329])).
fof(f33769,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30471,f28786])).
fof(f30471,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1508,f1464,f28786,f1329])).
fof(f33768,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30472,f28786])).
fof(f30472,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6359,f1464,f28786,f1329])).
fof(f33767,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30473,f28786])).
fof(f30473,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1808,f1464,f28786,f1329])).
fof(f33766,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30474,f28786])).
fof(f30474,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1510,f1464,f28786,f1329])).
fof(f33765,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30475,f928])).
fof(f30475,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f4299,f1464,f28786,f1329])).
fof(f33764,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30476,f928])).
fof(f30476,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5458,f1464,f28786,f1329])).
fof(f33763,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30477,f928])).
fof(f30477,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6009,f1464,f28786,f1329])).
fof(f33762,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30478,f928])).
fof(f30478,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f3819,f1464,f28786,f1329])).
fof(f33761,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30480,f28786])).
fof(f30480,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1511,f1464,f28786,f1329])).
fof(f33760,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30481,f28786])).
fof(f30481,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6360,f1464,f28786,f1329])).
fof(f33759,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30482,f28786])).
fof(f30482,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1809,f1464,f28786,f1329])).
fof(f33758,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30483,f28786])).
fof(f30483,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1513,f1464,f28786,f1329])).
fof(f33757,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30484,f30488])).
fof(f30484,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1464,f28786,f1329])).
fof(f33756,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30486,f30488])).
fof(f30486,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1464,f28786,f1329])).
fof(f33755,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30491,f28786])).
fof(f30491,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1464,f28786,f1333])).
fof(f1333,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f897])).
fof(f33754,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30492,f28786])).
fof(f30492,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6370,f1464,f28786,f1333])).
fof(f33753,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30493,f28786])).
fof(f30493,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1810,f1464,f28786,f1333])).
fof(f33752,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30494,f28786])).
fof(f30494,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1464,f28786,f1333])).
fof(f33751,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30495,f28786])).
fof(f30495,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1520,f1464,f28786,f1333])).
fof(f33750,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30496,f28786])).
fof(f30496,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1519,f1464,f28786,f1333])).
fof(f33749,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30497,f28786])).
fof(f30497,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6371,f1464,f28786,f1333])).
fof(f33748,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30498,f28786])).
fof(f30498,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1811,f1464,f28786,f1333])).
fof(f33747,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30499,f28786])).
fof(f30499,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1521,f1464,f28786,f1333])).
fof(f33746,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30510,f30568])).
fof(f30510,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f928,f28786,f1333])).
fof(f33745,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30527,f30568])).
fof(f30527,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f28786,f1333])).
fof(f33744,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30554,f28786])).
fof(f30554,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1464,f28786,f1333])).
fof(f33743,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30555,f28786])).
fof(f30555,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1516,f1464,f28786,f1333])).
fof(f33742,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30556,f28786])).
fof(f30556,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6370,f1464,f28786,f1333])).
fof(f33741,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30557,f28786])).
fof(f30557,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1810,f1464,f28786,f1333])).
fof(f33740,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30558,f28786])).
fof(f30558,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1518,f1464,f28786,f1333])).
fof(f33739,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30560,f28786])).
fof(f30560,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1519,f1464,f28786,f1333])).
fof(f33738,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30561,f28786])).
fof(f30561,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6371,f1464,f28786,f1333])).
fof(f33737,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30562,f28786])).
fof(f30562,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1811,f1464,f28786,f1333])).
fof(f33736,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30563,f28786])).
fof(f30563,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1521,f1464,f28786,f1333])).
fof(f33735,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30564,f30568])).
fof(f30564,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1464,f28786,f1333])).
fof(f33734,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30566,f30568])).
fof(f30566,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1464,f28786,f1333])).
fof(f33733,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30581,f32168])).
fof(f30581,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f21730,f28786,f1342])).
fof(f21730,plain,(
  ~sP6('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f19940,f1324])).
fof(f19940,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1464,f19033,f1329])).
fof(f33732,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30583,f32168])).
fof(f30583,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f28786,f1342])).
fof(f33731,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30585,f1103])).
fof(f30585,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK31(X0,bool))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3559,f28786,f1342])).
fof(f3559,plain,(
  ( ! [X0] : (~sP6('F_0',X0,sK31(X0,bool))) )),
  inference(unit_resulting_resolution,[],[f1464,f1103,f1344])).
fof(f33730,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30586,f1059])).
fof(f30586,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f28786,f1342])).
fof(f33729,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30588,f1060])).
fof(f30588,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3544,f28786,f1342])).
fof(f3544,plain,(
  ( ! [X0,X1] : (~sP6('F_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f1464,f1060,f1344])).
fof(f33728,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30589,f19854])).
fof(f19854,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f19033,f1189])).
fof(f30589,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f21968,f28786,f1342])).
fof(f21968,plain,(
  ~sP6('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f19967,f1323])).
fof(f1323,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f891])).
fof(f19967,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1464,f19033,f1333])).
fof(f33727,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30590,f1639])).
fof(f30590,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f3547,f28786,f1342])).
fof(f3547,plain,(
  ~sP6('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f1464,f1639,f1344])).
fof(f33726,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30592,f1672])).
fof(f30592,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3549,f28786,f1342])).
fof(f3549,plain,(
  ~sP6('F_0','F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1672,f1344])).
fof(f33725,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30593,f1674])).
fof(f30593,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f3551,f28786,f1342])).
fof(f3551,plain,(
  ~sP6('F_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f1464,f1674,f1344])).
fof(f33724,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30594,f5343])).
fof(f30594,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5368,f28786,f1342])).
fof(f5368,plain,(
  ~sP6('F_0','F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f5343,f1344])).
fof(f33723,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30595,f5716])).
fof(f30595,plain,(
  s(bool,'F_0') != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5807,f28786,f1342])).
fof(f5807,plain,(
  ~sP6('F_0','F_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f5716,f1344])).
fof(f33722,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30596,f19032])).
fof(f30596,plain,(
  s(bool,'F_0') != s(bool,sK40('F_0','=_1'(s(bool,'F_0')),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28261,f28786,f1342])).
fof(f28261,plain,(
  ~sP6('F_0','F_0',sK40('F_0','=_1'(s(bool,'F_0')),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19032,f1344])).
fof(f33721,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30597,f1672])).
fof(f30597,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3533,f28786,f1342])).
fof(f3533,plain,(
  ~sP6('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1672,f1344])).
fof(f33720,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30598,f32168])).
fof(f30598,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f4626,f28786,f1342])).
fof(f4626,plain,(
  ~sP6('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3555,f1672])).
fof(f33719,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30599,f1060])).
fof(f30599,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3528,f28786,f1342])).
fof(f3528,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.bool.DATATYPE_1'(s(X0,X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1060,f1344])).
fof(f33718,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30600,f1059])).
fof(f30600,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f28786,f1342])).
fof(f33717,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30601,f32168])).
fof(f30601,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f4484,f28786,f1342])).
fof(f4484,plain,(
  ~sP6('F_0','const.combin.I_1'(s(bool,'F_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(superposition,[],[f3539,f1672])).
fof(f33716,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30602,f32168])).
fof(f30602,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5387,f28786,f1342])).
fof(f5387,plain,(
  ~sP6('F_0','const.combin.I_1'(s(bool,'F_0')),sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | $spl1),
  inference(superposition,[],[f3539,f5343])).
fof(f33715,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30603,f32168])).
fof(f30603,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5826,f28786,f1342])).
fof(f5826,plain,(
  ~sP6('F_0','const.combin.I_1'(s(bool,'F_0')),sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | $spl1),
  inference(superposition,[],[f3539,f5716])).
fof(f33714,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30604,f19854])).
fof(f30604,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f21961,f28786,f1342])).
fof(f21961,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f19955,f1323])).
fof(f19955,plain,(
  ~sP4('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1464,f19033,f1333])).
fof(f33713,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30605,f1674])).
fof(f30605,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f3535,f28786,f1342])).
fof(f3535,plain,(
  ~sP6('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1674,f1344])).
fof(f33712,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30606,f1639])).
fof(f30606,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f3531,f28786,f1342])).
fof(f3531,plain,(
  ~sP6('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f1464,f1639,f1344])).
fof(f33711,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30607,f1103])).
fof(f30607,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK31(X0,bool))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3543,f28786,f1342])).
fof(f3543,plain,(
  ( ! [X0] : (~sP6('F_0',sK31(X0,bool),X0)) )),
  inference(unit_resulting_resolution,[],[f1464,f1103,f1344])).
fof(f33710,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30608,f5343])).
fof(f30608,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5365,f28786,f1342])).
fof(f5365,plain,(
  ~sP6('F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f5343,f1344])).
fof(f33709,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30609,f5716])).
fof(f30609,plain,(
  s(bool,'F_0') != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5804,f28786,f1342])).
fof(f5804,plain,(
  ~sP6('F_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f5716,f1344])).
fof(f33708,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30610,f32168])).
fof(f30610,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5388,f28786,f1342])).
fof(f5388,plain,(
  ~sP6('F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3555,f5343])).
fof(f33707,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30611,f32168])).
fof(f30611,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5827,f28786,f1342])).
fof(f5827,plain,(
  ~sP6('F_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3555,f5716])).
fof(f33706,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30612,f19032])).
fof(f30612,plain,(
  s(bool,'F_0') != s(bool,sK40('F_0','=_1'(s(bool,'F_0')),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28254,f28786,f1342])).
fof(f28254,plain,(
  ~sP6('F_0',sK40('F_0','=_1'(s(bool,'F_0')),bool),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19032,f1344])).
fof(f33705,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30613,f1672])).
fof(f30613,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6425,f28786,f1342])).
fof(f33704,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30614,f1672])).
fof(f30614,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6447,f28786,f1342])).
fof(f6447,plain,(
  ( ! [X0] : (~sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1672,f6300,f1344])).
fof(f33703,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30615,f5343])).
fof(f30615,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6453,f28786,f1342])).
fof(f6453,plain,(
  ( ! [X0] : (~sP6('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f5343,f6300,f1344])).
fof(f33702,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30616,f5343])).
fof(f30616,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6431,f28786,f1342])).
fof(f33655,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30824,f31197])).
fof(f31197,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f28786,f28786,f1358])).
fof(f1358,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f907])).
fof(f907,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f906])).
fof(f906,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f738])).
fof(f738,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f620,f737,f736])).
fof(f736,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f737,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f620,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f206])).
fof(f206,axiom,(
  ! [X66,X67,X68] : ((p(s(bool,X68)) <=> (p(s(bool,X67)) & p(s(bool,X66)))) <=> ((p(s(bool,X68)) | ~p(s(bool,X67)) | ~p(s(bool,X66))) & (p(s(bool,X67)) | ~p(s(bool,X68))) & (p(s(bool,X66)) | ~p(s(bool,X68)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sat.sat_23_prove_ltb')).
fof(f30824,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1346])).
fof(f1346,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f903,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f902])).
fof(f902,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f901])).
fof(f901,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f737])).
fof(f33654,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30841,f31197])).
fof(f30841,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1347])).
fof(f1347,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f33653,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30843,f28786])).
fof(f30843,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1933,f28786,f1348])).
fof(f1348,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f903])).
fof(f1933,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1346])).
fof(f1531,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1354])).
fof(f1354,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f905])).
fof(f905,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f904])).
fof(f904,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f736])).
fof(f33652,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30844,f28786])).
fof(f30844,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f1933,f28786,f1348])).
fof(f1528,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1353])).
fof(f1353,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f905])).
fof(f33651,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30846,f28786])).
fof(f30846,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f1933,f28786,f1348])).
fof(f1533,plain,(
  ( ! [X0,X1] : (sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1355])).
fof(f1355,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f905])).
fof(f33650,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30847,f28786])).
fof(f30847,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f1933,f28786,f1348])).
fof(f6482,plain,(
  ( ! [X2,X0,X1] : (sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1355])).
fof(f33649,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30848,f28786])).
fof(f30848,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1813,f1933,f28786,f1348])).
fof(f1813,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1355])).
fof(f33648,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30849,f28786])).
fof(f30849,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1535,f1933,f28786,f1348])).
fof(f1535,plain,(
  ( ! [X0,X1] : (sP7('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1355])).
fof(f33647,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30850,f28786])).
fof(f30850,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6464,f28786,f1348])).
fof(f6464,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1346])).
fof(f6481,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1354])).
fof(f33646,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30851,f28786])).
fof(f30851,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f6464,f28786,f1348])).
fof(f33645,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30853,f28786])).
fof(f30853,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f6464,f28786,f1348])).
fof(f33644,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30854,f28786])).
fof(f30854,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6464,f28786,f1348])).
fof(f33643,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30855,f28786])).
fof(f30855,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1813,f6464,f28786,f1348])).
fof(f33642,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30856,f28786])).
fof(f30856,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1535,f6464,f28786,f1348])).
fof(f33641,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30858,f28786])).
fof(f30858,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6468,f28786,f1348])).
fof(f6468,plain,(
  ( ! [X0,X1] : (~sP8('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1347])).
fof(f33640,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30859,f28786])).
fof(f30859,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f6468,f28786,f1348])).
fof(f33639,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30860,f28786])).
fof(f30860,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f6468,f28786,f1348])).
fof(f33638,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30861,f28786])).
fof(f30861,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f6468,f28786,f1348])).
fof(f1530,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1354])).
fof(f33637,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30862,f28786])).
fof(f30862,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6468,f28786,f1348])).
fof(f33636,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30863,f28786])).
fof(f30863,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1812,f6468,f28786,f1348])).
fof(f1812,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.bag.BAG_IN_2'(s(X1,X2),s(fun(X1,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1354])).
fof(f33635,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30864,f28786])).
fof(f30864,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1532,f6468,f28786,f1348])).
fof(f1532,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1354])).
fof(f33634,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30865,f28786])).
fof(f30865,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1531,f20065,f28786,f1348])).
fof(f20065,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1346])).
fof(f33633,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30866,f28786])).
fof(f30866,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20113,f20065,f28786,f1348])).
fof(f20113,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1353])).
fof(f33632,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30868,f28786])).
fof(f30868,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f20065,f28786,f1348])).
fof(f33631,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30869,f28786])).
fof(f30869,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6482,f20065,f28786,f1348])).
fof(f33630,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30870,f28786])).
fof(f30870,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1813,f20065,f28786,f1348])).
fof(f33629,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30871,f28786])).
fof(f30871,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1535,f20065,f28786,f1348])).
fof(f33628,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30883,f31081])).
fof(f31081,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1353])).
fof(f30883,plain,(
  ~sP7('T_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f28786,f1348])).
fof(f33627,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30885,f31081])).
fof(f30885,plain,(
  ( ! [X0] : (~sP7('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f928,f28786,f1348])).
fof(f33626,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30890,f31081])).
fof(f30890,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f928,f28786,f1348])).
fof(f33625,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30907,f31081])).
fof(f30907,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1348])).
fof(f33624,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30909,f31081])).
fof(f30909,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1348])).
fof(f33623,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30911,f31081])).
fof(f30911,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1348])).
fof(f1934,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1347])).
fof(f33622,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30913,f31081])).
fof(f30913,plain,(
  ( ! [X0,X1] : (~sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1348])).
fof(f33621,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30919,f31081])).
fof(f30919,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1348])).
fof(f33620,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30921,f31081])).
fof(f30921,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1348])).
fof(f20066,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1347])).
fof(f33619,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30924,f28786])).
fof(f30924,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6464,f28786,f1348])).
fof(f33618,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30925,f28786])).
fof(f30925,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f6464,f28786,f1348])).
fof(f33617,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30926,f28786])).
fof(f30926,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f6464,f28786,f1348])).
fof(f1534,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1355])).
fof(f33616,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30927,f28786])).
fof(f30927,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f6464,f28786,f1348])).
fof(f33615,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30928,f28786])).
fof(f30928,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6464,f28786,f1348])).
fof(f33614,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30929,f28786])).
fof(f30929,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1813,f6464,f28786,f1348])).
fof(f33613,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30930,f28786])).
fof(f30930,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1535,f6464,f28786,f1348])).
fof(f33612,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30931,f28786])).
fof(f30931,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1934,f28786,f1348])).
fof(f33611,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30932,f28786])).
fof(f30932,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f1934,f28786,f1348])).
fof(f33610,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30934,f28786])).
fof(f30934,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f1934,f28786,f1348])).
fof(f33609,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30935,f28786])).
fof(f30935,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f1934,f28786,f1348])).
fof(f33608,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30936,f28786])).
fof(f30936,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1812,f1934,f28786,f1348])).
fof(f33607,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30937,f28786])).
fof(f30937,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1532,f1934,f28786,f1348])).
fof(f33606,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30938,f28786])).
fof(f30938,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6468,f28786,f1348])).
fof(f33605,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30939,f28786])).
fof(f30939,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1528,f6468,f28786,f1348])).
fof(f33604,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30941,f28786])).
fof(f30941,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f6468,f28786,f1348])).
fof(f33603,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30942,f28786])).
fof(f30942,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6468,f28786,f1348])).
fof(f33602,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30943,f28786])).
fof(f30943,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1812,f6468,f28786,f1348])).
fof(f33601,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30944,f28786])).
fof(f30944,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1532,f6468,f28786,f1348])).
fof(f33600,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30946,f28786])).
fof(f30946,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1534,f20066,f28786,f1348])).
fof(f33599,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30947,f28786])).
fof(f30947,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20113,f20066,f28786,f1348])).
fof(f33598,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30949,f28786])).
fof(f30949,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f20066,f28786,f1348])).
fof(f33597,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30950,f28786])).
fof(f30950,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6481,f20066,f28786,f1348])).
fof(f33596,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30951,f28786])).
fof(f30951,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1812,f20066,f28786,f1348])).
fof(f33595,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30952,f28786])).
fof(f30952,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1532,f20066,f28786,f1348])).
fof(f33594,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30963,f31081])).
fof(f30963,plain,(
  ~sP7('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f928,f28786,f1348])).
fof(f33593,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30965,f31081])).
fof(f30965,plain,(
  ( ! [X0] : (~sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f928,f28786,f1348])).
fof(f33592,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30970,f31081])).
fof(f30970,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f928,f28786,f1348])).
fof(f33591,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30987,f31081])).
fof(f30987,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1348])).
fof(f33590,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30989,f31081])).
fof(f30989,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1348])).
fof(f33589,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30991,f31081])).
fof(f30991,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1348])).
fof(f33588,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f30993,f31081])).
fof(f30993,plain,(
  ( ! [X0,X1] : (~sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1348])).
fof(f33587,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f30999,f31081])).
fof(f30999,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1348])).
fof(f33586,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31001,f31081])).
fof(f31001,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1348])).
fof(f33585,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31004,f28786])).
fof(f31004,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f1464,f28786,f1352])).
fof(f1352,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f905])).
fof(f33584,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31005,f28786])).
fof(f31005,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f1464,f28786,f1352])).
fof(f33583,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31006,f28786])).
fof(f31006,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1812,f1464,f28786,f1352])).
fof(f33582,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31007,f28786])).
fof(f31007,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1532,f1464,f28786,f1352])).
fof(f33581,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31008,f28786])).
fof(f31008,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1464,f28786,f1352])).
fof(f33580,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31009,f28786])).
fof(f31009,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f1464,f28786,f1352])).
fof(f33579,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31010,f28786])).
fof(f31010,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f1464,f28786,f1352])).
fof(f33578,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31011,f28786])).
fof(f31011,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1813,f1464,f28786,f1352])).
fof(f33577,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31012,f28786])).
fof(f31012,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1535,f1464,f28786,f1352])).
fof(f33576,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31023,f31081])).
fof(f31023,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f928,f28786,f1352])).
fof(f33575,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31040,f31081])).
fof(f31040,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f28786,f1352])).
fof(f33574,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31067,f28786])).
fof(f31067,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1464,f28786,f1352])).
fof(f33573,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31068,f28786])).
fof(f31068,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1530,f1464,f28786,f1352])).
fof(f33572,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31069,f28786])).
fof(f31069,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f1464,f28786,f1352])).
fof(f33571,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31070,f28786])).
fof(f31070,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1812,f1464,f28786,f1352])).
fof(f33570,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31071,f28786])).
fof(f31071,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1532,f1464,f28786,f1352])).
fof(f33569,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31073,f28786])).
fof(f31073,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1533,f1464,f28786,f1352])).
fof(f33568,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31074,f28786])).
fof(f31074,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f1464,f28786,f1352])).
fof(f33567,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31075,f28786])).
fof(f31075,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1813,f1464,f28786,f1352])).
fof(f33566,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31076,f28786])).
fof(f31076,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1535,f1464,f28786,f1352])).
fof(f33565,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31077,f31081])).
fof(f31077,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1464,f28786,f1352])).
fof(f33564,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31079,f31081])).
fof(f31079,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1464,f28786,f1352])).
fof(f33563,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31084,f28786])).
fof(f31084,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f928,f28786,f1358])).
fof(f33562,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31085,f28786])).
fof(f31085,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f928,f28786,f1358])).
fof(f33561,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31086,f28786])).
fof(f31086,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f928,f28786,f1358])).
fof(f33560,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31088,f28786])).
fof(f31088,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33559,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31089,f28786])).
fof(f31089,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1358])).
fof(f33558,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31090,f28786])).
fof(f31090,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33557,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31091,f928])).
fof(f31091,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f28786,f28786,f1358])).
fof(f5135,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f5118,f1345])).
fof(f1345,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f903])).
fof(f5118,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f928,f1464,f928,f1352])).
fof(f33556,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31092,f928])).
fof(f31092,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20130,f28786,f28786,f1358])).
fof(f20130,plain,(
  ~sP8('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f928,f19033,f1361])).
fof(f1361,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f907])).
fof(f33555,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31093,f28786])).
fof(f31093,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20136,f28786,f28786,f1358])).
fof(f20136,plain,(
  ~sP8('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f928,f19033,f1361])).
fof(f33554,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31094,f928])).
fof(f31094,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f6493,plain,(
  ( ! [X0] : (~sP8('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f928,f6300,f1361])).
fof(f33553,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31096,f28786])).
fof(f31096,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1358])).
fof(f33552,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31097,f928])).
fof(f31097,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f928,f28786,f1358])).
fof(f33551,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31098,f928])).
fof(f31098,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f928,f28786,f1358])).
fof(f33550,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31101,f28786])).
fof(f31101,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f928,f28786,f1358])).
fof(f33549,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31102,f28786])).
fof(f31102,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f928,f28786,f1358])).
fof(f33548,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31107,f928])).
fof(f31107,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1358])).
fof(f33547,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31108,f928])).
fof(f31108,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33546,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31109,f928])).
fof(f31109,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1358])).
fof(f33545,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31110,f928])).
fof(f31110,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33544,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31114,f28786])).
fof(f31114,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f33543,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31115,f28786])).
fof(f31115,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1358])).
fof(f33542,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31116,f28786])).
fof(f31116,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1358])).
fof(f33541,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31122,f28786])).
fof(f31122,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f928,f28786,f1358])).
fof(f33540,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31123,f928])).
fof(f31123,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f928,f28786,f1358])).
fof(f33539,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31124,f28786])).
fof(f31124,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20130,f928,f28786,f1358])).
fof(f33538,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31125,f928])).
fof(f31125,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f928,f28786,f1358])).
fof(f33537,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31127,f28786])).
fof(f31127,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1358])).
fof(f33536,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31128,f28786])).
fof(f31128,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33535,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31130,f28786])).
fof(f31130,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33534,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31131,f928])).
fof(f31131,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f28786,f28786,f1358])).
fof(f33533,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31132,f28786])).
fof(f31132,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20130,f28786,f28786,f1358])).
fof(f33532,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31133,f928])).
fof(f31133,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20136,f28786,f28786,f1358])).
fof(f33531,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31134,f928])).
fof(f31134,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f33530,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31135,f28786])).
fof(f31135,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1358])).
fof(f33529,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31138,f28786])).
fof(f31138,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f928,f28786,f1358])).
fof(f33528,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31139,f928])).
fof(f31139,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f928,f28786,f1358])).
fof(f33527,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31140,f28786])).
fof(f31140,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20136,f928,f28786,f1358])).
fof(f33526,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31141,f928])).
fof(f31141,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f928,f28786,f1358])).
fof(f33525,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31148,f28786])).
fof(f31148,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33524,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31149,f28786])).
fof(f31149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1358])).
fof(f33523,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31150,f28786])).
fof(f31150,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33522,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31151,f928])).
fof(f31151,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f28786,f28786,f1358])).
fof(f33521,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31152,f928])).
fof(f31152,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20130,f28786,f28786,f1358])).
fof(f33520,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31153,f28786])).
fof(f31153,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20136,f28786,f28786,f1358])).
fof(f33519,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31154,f928])).
fof(f31154,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f33518,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31156,f28786])).
fof(f31156,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1358])).
fof(f33517,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31161,f28786])).
fof(f31161,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f28786,f1358])).
fof(f33516,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31162,f28786])).
fof(f31162,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f928,f28786,f1358])).
fof(f33515,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31164,f28786])).
fof(f31164,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f928,f28786,f1358])).
fof(f33514,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31165,f28786])).
fof(f31165,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1358])).
fof(f33513,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31166,f28786])).
fof(f31166,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33512,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31168,f28786])).
fof(f31168,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33511,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31169,f928])).
fof(f31169,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f28786,f28786,f1358])).
fof(f33510,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31170,f28786])).
fof(f31170,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20130,f28786,f28786,f1358])).
fof(f33509,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31171,f928])).
fof(f31171,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20136,f28786,f28786,f1358])).
fof(f33508,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31172,f928])).
fof(f31172,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f33507,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31173,f28786])).
fof(f31173,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1358])).
fof(f33506,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31175,f928])).
fof(f31175,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f28786,f1358])).
fof(f33505,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31176,f928])).
fof(f31176,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f928,f28786,f1358])).
fof(f33504,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31179,f28786])).
fof(f31179,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f928,f28786,f1358])).
fof(f33503,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31180,f28786])).
fof(f31180,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f928,f28786,f1358])).
fof(f33502,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31185,f928])).
fof(f31185,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f1358])).
fof(f33501,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31186,f928])).
fof(f31186,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6464,f28786,f28786,f1358])).
fof(f33500,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31187,f928])).
fof(f31187,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f28786,f1358])).
fof(f33499,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31188,f928])).
fof(f31188,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6468,f28786,f28786,f1358])).
fof(f33498,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31192,f28786])).
fof(f31192,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6493,f28786,f28786,f1358])).
fof(f33497,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31193,f28786])).
fof(f31193,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20065,f28786,f28786,f1358])).
fof(f33496,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31194,f28786])).
fof(f31194,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20066,f28786,f28786,f1358])).
fof(f33495,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31200,f28786])).
fof(f31200,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2020,f1464,f28786,f1361])).
fof(f2020,plain,(
  ( ! [X0] : (sP8('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1530,f1464,f1447])).
fof(f1447,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1351])).
fof(f1351,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f903])).
fof(f33494,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31201,f28786])).
fof(f31201,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6483,f1464,f28786,f1361])).
fof(f6483,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1356])).
fof(f1356,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f907])).
fof(f33493,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31202,f28786])).
fof(f31202,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f2021,f1464,f28786,f1361])).
fof(f2021,plain,(
  ( ! [X0] : (sP8('F_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1532,f1464,f1447])).
fof(f33492,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31203,f28786])).
fof(f31203,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1943,f1464,f28786,f1361])).
fof(f1943,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1357])).
fof(f1357,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f907])).
fof(f33491,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31204,f28786])).
fof(f31204,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2023,f1464,f28786,f1361])).
fof(f2023,plain,(
  ( ! [X0] : (sP8('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1533,f1464,f1447])).
fof(f33490,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31205,f28786])).
fof(f31205,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6488,f1464,f28786,f1361])).
fof(f6488,plain,(
  ( ! [X0,X1] : (sP8('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1357])).
fof(f33489,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31206,f28786])).
fof(f31206,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f2024,f1464,f28786,f1361])).
fof(f2024,plain,(
  ( ! [X0] : (sP8('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1535,f1464,f1447])).
fof(f33488,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31235,f28786])).
fof(f31235,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1937,f1464,f28786,f1361])).
fof(f1937,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1356])).
fof(f33487,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31236,f28786])).
fof(f31236,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2020,f1464,f28786,f1361])).
fof(f33486,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31237,f28786])).
fof(f31237,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6483,f1464,f28786,f1361])).
fof(f33485,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31238,f28786])).
fof(f31238,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f2021,f1464,f28786,f1361])).
fof(f33484,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31240,f28786])).
fof(f31240,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2023,f1464,f28786,f1361])).
fof(f33483,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31241,f28786])).
fof(f31241,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6488,f1464,f28786,f1361])).
fof(f33482,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31242,f28786])).
fof(f31242,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f2024,f1464,f28786,f1361])).
fof(f33481,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31292,f32142])).
fof(f32142,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1541,f28786,f1446])).
fof(f1446,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1365])).
fof(f1365,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f910])).
fof(f910,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f909])).
fof(f909,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f908])).
fof(f908,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f1541,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1372])).
fof(f1372,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f913])).
fof(f913,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f912])).
fof(f912,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f911])).
fof(f911,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f739])).
fof(f739,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f31292,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1362])).
fof(f1362,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f910])).
fof(f33480,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31315,f32142])).
fof(f31315,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1363])).
fof(f1363,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f910])).
fof(f33479,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31353,f31435])).
fof(f31435,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1371])).
fof(f1371,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f913])).
fof(f31353,plain,(
  ~sP9('F_0','T_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f1464,f28786,f1366])).
fof(f1366,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f910])).
fof(f1948,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1363])).
fof(f33478,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31355,f31435])).
fof(f31355,plain,(
  ~sP9('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20147,f1464,f28786,f1366])).
fof(f20147,plain,(
  ( ! [X0] : (~sP10('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1363])).
fof(f33477,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31357,f31435])).
fof(f31357,plain,(
  ( ! [X0] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f1464,f28786,f1366])).
fof(f6495,plain,(
  ( ! [X0,X1] : (~sP10('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1363])).
fof(f33476,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31393,f31435])).
fof(f31393,plain,(
  ~sP9('F_0','F_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1464,f28786,f1367])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f910])).
fof(f1947,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1362])).
fof(f33475,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31394,f31435])).
fof(f31394,plain,(
  ~sP9('F_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f1464,f28786,f1367])).
fof(f20142,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1362])).
fof(f33474,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31395,f31435])).
fof(f31395,plain,(
  ( ! [X0] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f1464,f28786,f1367])).
fof(f6494,plain,(
  ( ! [X0,X1] : (~sP10('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1362])).
fof(f33473,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31431,f31435])).
fof(f31431,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1464,f28786,f1369])).
fof(f1369,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f913])).
fof(f33472,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31439,f28786])).
fof(f31439,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6571,f28786,f1374])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f914])).
fof(f914,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f741])).
fof(f741,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f622,f740,f739])).
fof(f622,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f211])).
fof(f211,axiom,(
  ! [X66,X67,X68] : ((p(s(bool,X68)) <=> (p(s(bool,X67)) | p(s(bool,X66)))) <=> ((p(s(bool,X68)) | ~p(s(bool,X67))) & (p(s(bool,X68)) | ~p(s(bool,X66))) & (p(s(bool,X67)) | p(s(bool,X66)) | ~p(s(bool,X68)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f6571,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f1464,f928,f6300,f1376])).
fof(f1376,plain,(
  ( ! [X2,X0,X1] : (~sP10(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f915])).
fof(f33471,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31440,f928])).
fof(f31440,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f28786,f1374])).
fof(f33470,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31441,f28786])).
fof(f31441,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f28786,f1374])).
fof(f33469,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31442,f28786])).
fof(f31442,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f28786,f1374])).
fof(f33468,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31443,f28786])).
fof(f31443,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20147,f28786,f1374])).
fof(f33467,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31444,f928])).
fof(f31444,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f28786,f1374])).
fof(f33466,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31445,f28786])).
fof(f31445,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f28786,f1374])).
fof(f33465,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31451,f928])).
fof(f31451,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f1374])).
fof(f5326,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f5309,f1364])).
fof(f1364,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f910])).
fof(f5309,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f928,f1464,f1464,f1369])).
fof(f33464,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31452,f928])).
fof(f31452,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6565,f28786,f1374])).
fof(f6565,plain,(
  ( ! [X0] : (~sP10('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f928,f6300,f1376])).
fof(f33463,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31453,f928])).
fof(f31453,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6571,f28786,f1374])).
fof(f33462,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31458,f28786])).
fof(f31458,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f28786,f1374])).
fof(f33461,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31459,f28786])).
fof(f31459,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f28786,f1374])).
fof(f33460,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31460,f28786])).
fof(f31460,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20183,f28786,f1374])).
fof(f20183,plain,(
  ~sP10('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1464,f19033,f1376])).
fof(f33459,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31465,f928])).
fof(f31465,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f1375])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f915])).
fof(f33458,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31466,f928])).
fof(f31466,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6565,f28786,f1375])).
fof(f33457,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31467,f928])).
fof(f31467,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6571,f28786,f1375])).
fof(f33456,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31472,f28786])).
fof(f31472,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f28786,f1375])).
fof(f33455,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31473,f28786])).
fof(f31473,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f28786,f1375])).
fof(f33454,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31474,f28786])).
fof(f31474,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20183,f28786,f1375])).
fof(f33453,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31480,f28786])).
fof(f31480,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6565,f28786,f1375])).
fof(f33452,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31482,f28786])).
fof(f31482,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f28786,f1375])).
fof(f33451,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31483,f28786])).
fof(f31483,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f28786,f1375])).
fof(f33450,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31484,f928])).
fof(f31484,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f28786,f1375])).
fof(f33449,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31485,f28786])).
fof(f31485,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20147,f28786,f1375])).
fof(f33448,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31486,f28786])).
fof(f31486,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f28786,f1375])).
fof(f33447,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31487,f928])).
fof(f31487,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f28786,f1375])).
fof(f33446,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31535,f32107])).
fof(f32107,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f28786,f1445])).
fof(f1445,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1382])).
fof(f1382,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f918])).
fof(f918,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f917])).
fof(f917,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f916])).
fof(f916,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f743])).
fof(f743,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1549,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1464,f1389])).
fof(f1389,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f921])).
fof(f921,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f920])).
fof(f920,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f919])).
fof(f919,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f742])).
fof(f742,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f31535,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1380])).
fof(f1380,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f918])).
fof(f33445,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31538,f28786])).
fof(f31538,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f6579,f28786,f1383])).
fof(f1383,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f918])).
fof(f6579,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1379])).
fof(f1379,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f918])).
fof(f33444,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31539,f28786])).
fof(f31539,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1546,f6579,f28786,f1383])).
fof(f1546,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1388])).
fof(f1388,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f921])).
fof(f33443,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31540,f28786])).
fof(f31540,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1593,f6579,f28786,f1383])).
fof(f1593,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f932,f1388])).
fof(f33442,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31541,f28786])).
fof(f31541,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1609,f6579,f28786,f1383])).
fof(f1609,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f943,f1388])).
fof(f33441,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31542,f28786])).
fof(f31542,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20240,f6579,f28786,f1383])).
fof(f20240,plain,(
  ( ! [X0,X1] : (sP11(X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1388])).
fof(f33440,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31543,f28786])).
fof(f31543,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1547,f6579,f28786,f1383])).
fof(f1547,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f930,f1388])).
fof(f33439,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31544,f28786])).
fof(f31544,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f1966,f28786,f1383])).
fof(f1966,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1380])).
fof(f33438,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31545,f28786])).
fof(f31545,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1546,f1966,f28786,f1383])).
fof(f33437,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31547,f28786])).
fof(f31547,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1543,f1966,f28786,f1383])).
fof(f1543,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1387])).
fof(f1387,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f921])).
fof(f33436,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31548,f28786])).
fof(f31548,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1815,f1966,f28786,f1383])).
fof(f1815,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,X1,'const.bag.BAG_IN_2'(s(X2,X3),s(fun(X2,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(unit_resulting_resolution,[],[f1102,f1387])).
fof(f33435,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31549,f28786])).
fof(f31549,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1545,f1966,f28786,f1383])).
fof(f1545,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f940,f1387])).
fof(f33434,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31550,f28786])).
fof(f31550,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1549,f20184,f28786,f1383])).
fof(f20184,plain,(
  ( ! [X0] : (~sP12('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f19033,f1380])).
fof(f33433,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31551,f28786])).
fof(f31551,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20240,f20184,f28786,f1383])).
fof(f33432,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31553,f28786])).
fof(f31553,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1543,f20184,f28786,f1383])).
fof(f33431,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31554,f28786])).
fof(f31554,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1815,f20184,f28786,f1383])).
fof(f33430,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31555,f28786])).
fof(f31555,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1545,f20184,f28786,f1383])).
fof(f33429,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31557,f28786])).
fof(f31557,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1546,f6585,f28786,f1383])).
fof(f6585,plain,(
  ( ! [X0,X1] : (~sP12('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1380])).
fof(f33428,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31559,f28786])).
fof(f31559,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1543,f6585,f28786,f1383])).
fof(f33427,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31560,f28786])).
fof(f31560,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1815,f6585,f28786,f1383])).
fof(f33426,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31561,f28786])).
fof(f31561,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1545,f6585,f28786,f1383])).
fof(f33425,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31572,f31833])).
fof(f31833,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1388])).
fof(f31572,plain,(
  ~sP11('F_0','T_0','T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f928,f28786,f1383])).
fof(f33424,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31574,f31833])).
fof(f31574,plain,(
  ~sP11('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'T_0') | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f928,f28786,f1383])).
fof(f33423,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31576,f31833])).
fof(f31576,plain,(
  ( ! [X0] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f928,f28786,f1383])).
fof(f33422,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31594,f31833])).
fof(f31594,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f28786,f28786,f1383])).
fof(f1961,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1379])).
fof(f33421,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31596,f31833])).
fof(f31596,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f28786,f28786,f1383])).
fof(f33420,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31598,f31833])).
fof(f31598,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f28786,f28786,f1383])).
fof(f33419,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31600,f31833])).
fof(f31600,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f28786,f28786,f1383])).
fof(f33418,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31602,f31833])).
fof(f31602,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f28786,f28786,f1383])).
fof(f6577,plain,(
  ( ! [X0,X1] : (~sP12('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f6300,f1379])).
fof(f33417,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31604,f31833])).
fof(f31604,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f28786,f28786,f1383])).
fof(f33416,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31610,f28786])).
fof(f31610,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1544,f28786,f1383])).
fof(f1544,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1464,f1387])).
fof(f33415,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31611,f28786])).
fof(f31611,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f1544,f28786,f1383])).
fof(f33414,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31614,f28786])).
fof(f31614,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6585,f1543,f28786,f1383])).
fof(f33413,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31617,f28786])).
fof(f31617,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f1815,f28786,f1383])).
fof(f33412,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31620,f28786])).
fof(f31620,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f1545,f28786,f1383])).
fof(f33411,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31624,f28786])).
fof(f31624,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f1546,f28786,f1383])).
fof(f33410,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31625,f28786])).
fof(f31625,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1546,f28786,f1383])).
fof(f33409,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31628,f28786])).
fof(f31628,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1593,f28786,f1383])).
fof(f33408,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31631,f28786])).
fof(f31631,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1609,f28786,f1383])).
fof(f33407,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31635,f28786])).
fof(f31635,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f6577,f20240,f28786,f1383])).
fof(f33406,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31638,f28786])).
fof(f31638,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1547,f28786,f1383])).
fof(f33405,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31653,f31833])).
fof(f31653,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f28786,f28786,f1383])).
fof(f33404,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31654,f31833])).
fof(f31654,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f28786,f28786,f1383])).
fof(f33403,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31655,f31833])).
fof(f31655,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f28786,f28786,f1383])).
fof(f33402,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31656,f31833])).
fof(f31656,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f28786,f28786,f1383])).
fof(f33401,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31657,f31833])).
fof(f31657,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f28786,f28786,f1383])).
fof(f33400,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31658,f31833])).
fof(f31658,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f28786,f28786,f1383])).
fof(f33399,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31710,f31833])).
fof(f31710,plain,(
  ~sP11('F_0','F_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1464,f28786,f1384])).
fof(f1384,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f918])).
fof(f33398,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31711,f31833])).
fof(f31711,plain,(
  ( ! [X0] : (~sP11('F_0','F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f1464,f28786,f1384])).
fof(f33397,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31712,f31833])).
fof(f31712,plain,(
  ( ! [X0] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1464,f28786,f1384])).
fof(f33396,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31751,f31833])).
fof(f31751,plain,(
  ~sP11('F_0','F_0','F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1464,f28786,f1385])).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f918])).
fof(f33395,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31752,f31833])).
fof(f31752,plain,(
  ( ! [X0] : (~sP11('F_0','F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f1464,f28786,f1385])).
fof(f33394,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31753,f31833])).
fof(f31753,plain,(
  ( ! [X0] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0','F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f1464,f28786,f1385])).
fof(f33393,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31788,f28786])).
fof(f31788,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1543,f1464,f28786,f1386])).
fof(f1386,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f921])).
fof(f33392,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31789,f28786])).
fof(f31789,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1815,f1464,f28786,f1386])).
fof(f33391,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31790,f28786])).
fof(f31790,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1545,f1464,f28786,f1386])).
fof(f33390,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31791,f28786])).
fof(f31791,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f1464,f28786,f1386])).
fof(f33389,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31792,f28786])).
fof(f31792,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1548,f1464,f28786,f1386])).
fof(f1548,plain,(
  ( ! [X0,X1] : (sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1461,f1389])).
fof(f33388,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31793,f28786])).
fof(f31793,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1816,f1464,f28786,f1386])).
fof(f1816,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1102,f1389])).
fof(f33387,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31794,f28786])).
fof(f31794,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1550,f1464,f28786,f1386])).
fof(f1550,plain,(
  ( ! [X0,X1] : (sP11('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f940,f1389])).
fof(f33386,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31795,f31833])).
fof(f31795,plain,(
  ( ! [X0] : (~sP11(X0,'F_0','T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1464,f28786,f1386])).
fof(f33385,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31797,f31833])).
fof(f31797,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1464,f28786,f1386])).
fof(f33384,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31821,f28786])).
fof(f31821,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1544,f1464,f28786,f1386])).
fof(f33383,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31822,f28786])).
fof(f31822,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1543,f1464,f28786,f1386])).
fof(f33382,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31823,f28786])).
fof(f31823,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1815,f1464,f28786,f1386])).
fof(f33381,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31824,f28786])).
fof(f31824,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1545,f1464,f28786,f1386])).
fof(f33380,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31826,f28786])).
fof(f31826,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1548,f1464,f28786,f1386])).
fof(f33379,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31827,f28786])).
fof(f31827,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1816,f1464,f28786,f1386])).
fof(f33378,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31828,f28786])).
fof(f31828,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1550,f1464,f28786,f1386])).
fof(f33377,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31829,f31833])).
fof(f31829,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1464,f28786,f1386])).
fof(f33376,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31831,f31833])).
fof(f31831,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1464,f28786,f1386])).
fof(f33375,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31858,f928])).
fof(f31858,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5498,f28786,f1392])).
fof(f1392,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f923])).
fof(f923,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f922])).
fof(f922,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f716,f743,f742])).
fof(f716,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f623])).
fof(f623,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f226])).
fof(f226,axiom,(
  ! [X66,X67,X68] : ((p(s(bool,X68)) <=> (p(s(bool,X67)) => p(s(bool,X66)))) <=> ((p(s(bool,X68)) | p(s(bool,X67))) & (p(s(bool,X68)) | ~p(s(bool,X66))) & (~p(s(bool,X67)) | p(s(bool,X66)) | ~p(s(bool,X68)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f5498,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f5481,f1381])).
fof(f1381,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f918])).
fof(f5481,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f928,f928,f1464,f1386])).
fof(f33374,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31859,f928])).
fof(f31859,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20253,f28786,f1392])).
fof(f20253,plain,(
  ~sP12('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f928,f19033,f1393])).
fof(f1393,plain,(
  ( ! [X2,X0,X1] : (~sP12(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f923])).
fof(f33373,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31864,f28786])).
fof(f31864,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f28786,f1392])).
fof(f33372,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31865,f28786])).
fof(f31865,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f28786,f1392])).
fof(f33371,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31866,f28786])).
fof(f31866,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20251,f28786,f1392])).
fof(f20251,plain,(
  ~sP12('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f928,f19033,f1393])).
fof(f33370,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31873,f28786])).
fof(f31873,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f28786,f1392])).
fof(f33369,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31874,f28786])).
fof(f31874,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f28786,f1392])).
fof(f33368,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31875,f928])).
fof(f31875,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f28786,f1392])).
fof(f33367,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31876,f28786])).
fof(f31876,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f28786,f1392])).
fof(f33366,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31877,f28786])).
fof(f31877,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f28786,f1392])).
fof(f33365,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31878,f928])).
fof(f31878,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f28786,f1392])).
fof(f33364,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31885,f28786])).
fof(f31885,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1988,f1464,f28786,f1393])).
fof(f1988,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1543,f928,f1445])).
fof(f33363,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31886,f28786])).
fof(f31886,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f1990,f1464,f28786,f1393])).
fof(f1990,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1545,f928,f1445])).
fof(f33362,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31887,f928])).
fof(f31887,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5465,f1464,f28786,f1393])).
fof(f5465,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1549,f928,f1464,f1385])).
fof(f33361,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31888,f28786])).
fof(f31888,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20241,f1464,f28786,f1393])).
fof(f20241,plain,(
  sP12('F_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1464,f19033,f1390])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f923])).
fof(f33360,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31889,f928])).
fof(f31889,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5466,f1464,f28786,f1393])).
fof(f5466,plain,(
  sP12('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1548,f928,f1464,f1385])).
fof(f33359,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31890,f28786])).
fof(f31890,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2712,f1464,f28786,f1393])).
fof(f2712,plain,(
  ( ! [X0,X1] : (sP12('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1543,f959,f1445])).
fof(f959,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X6] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X6,X6),'_i_0_0'),s(fun(X6,X6),'_i_0_0'))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.ind_type.ISO_REFL')).
fof(f33358,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31892,f928])).
fof(f31892,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f5468,f1464,f28786,f1393])).
fof(f5468,plain,(
  sP12('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1550,f928,f1464,f1385])).
fof(f33357,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31923,f928])).
fof(f31923,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1969,f1464,f28786,f1393])).
fof(f1969,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1464,f1391])).
fof(f1391,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f923])).
fof(f33356,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31924,f928])).
fof(f31924,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1988,f1464,f28786,f1393])).
fof(f33355,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31925,f928])).
fof(f31925,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(unit_resulting_resolution,[],[f1990,f1464,f28786,f1393])).
fof(f33354,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31928,f28786])).
fof(f31928,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5466,f1464,f28786,f1393])).
fof(f33353,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31929,f28786])).
fof(f31929,plain,(
  ( ! [X0] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f2712,f1464,f28786,f1393])).
fof(f33352,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f31930,f28786])).
fof(f31930,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20276,f1464,f28786,f1393])).
fof(f20276,plain,(
  ( ! [X0] : (sP12('|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1544,f19033,f1445])).
fof(f33351,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31931,f28786])).
fof(f31931,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl36),
  inference(unit_resulting_resolution,[],[f5468,f1464,f28786,f1393])).
fof(f33350,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f31960,f32168])).
fof(f31960,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786,f1444])).
fof(f1444,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1443])).
fof(f1443,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1412])).
fof(f1412,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f925])).
fof(f925,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f924])).
fof(f924,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f718])).
fof(f718,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f717])).
fof(f717,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X9,X52,X17,X53] : ((s(bool,X9) = s(bool,X52) & (p(s(bool,X52)) => s(bool,X17) = s(bool,X53))) => ((p(s(bool,X9)) => p(s(bool,X17))) <=> (p(s(bool,X52)) => p(s(bool,X53)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.bool.IMP_CONG')).
fof(f33349,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32082,f1544])).
fof(f32082,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f28786,f1445])).
fof(f33348,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32083,f6607])).
fof(f6607,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X2),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f6300,f1387])).
fof(f32083,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6579,f28786,f1445])).
fof(f33347,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32084,f1546])).
fof(f32084,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f28786,f1445])).
fof(f33346,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f32085,f20240])).
fof(f32085,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f28786,f1445])).
fof(f33345,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32086,f1544])).
fof(f32086,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'F_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6577,f28786,f1445])).
fof(f33344,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32087,f1546])).
fof(f32087,plain,(
  ( ! [X0,X1] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6585,f28786,f1445])).
fof(f33343,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32115,f1536])).
fof(f1536,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1370])).
fof(f1370,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f913])).
fof(f32115,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f28786,f1446])).
fof(f33342,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f32116,f20169])).
fof(f20169,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1370])).
fof(f32116,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))))) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f28786,f1446])).
fof(f33341,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32117,f1538])).
fof(f1538,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1371])).
fof(f32117,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f28786,f1446])).
fof(f33340,plain,(
  $false | ($spl1 | $spl36)),
  inference(subsumption_resolution,[],[f32118,f20170])).
fof(f20170,plain,(
  ( ! [X0,X1] : (sP9(X0,'|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f19033,f1371])).
fof(f32118,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun(bool,bool),'=_1'(s(bool,'F_0')))),X0)) ) | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20147,f28786,f1446])).
fof(f33339,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32119,f1536])).
fof(f32119,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f28786,f1446])).
fof(f33338,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32120,f1538])).
fof(f32120,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6495,f28786,f1446])).
fof(f33337,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32157,f32168])).
fof(f32157,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1086])).
fof(f33336,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32159,f32168])).
fof(f32159,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1283])).
fof(f1283,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4) | ~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))) )),
  inference(cnf_transformation,[],[f877])).
fof(f877,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,sK60(X4,X3,X2,X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X1)),sK61(X4,X3,X2,X1,X0)))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK60,sK61])],[f876])).
fof(f876,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X7,X8] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X7),s('type.list.list'('type.PairingHeap.heap'(X1)),X8))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(rectify,[],[f875])).
fof(f875,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(flattening,[],[f874])).
fof(f874,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & ((! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4)) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(nnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4)))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & ~s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X6] : ! [X8,X12,X13] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X6,X11),X8),s(fun(X11,fun(X11,bool)),X12),s('type.PairingHeap.heap'(X6),X13)))) <=> (? [X9,X14] : s('type.PairingHeap.heap'(X6),X13) = s('type.PairingHeap.heap'(X6),'const.PairingHeap.Tree_2'(s(X6,X9),s('type.list.list'('type.PairingHeap.heap'(X6)),X14))) & ~s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),X13)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.delete_min_side_def')).
fof(f33335,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32176,f28786])).
fof(f32176,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1194])).
fof(f1194,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f840])).
fof(f840,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f502])).
fof(f502,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X55,X69] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X69)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X69),s('type.num.num',X55)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_lte|split|5')).
fof(f33334,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32178,f28786])).
fof(f32178,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1226])).
fof(f1226,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f853])).
fof(f853,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f513])).
fof(f513,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f512])).
fof(f512,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f154])).
fof(f154,axiom,(
  ! [X69,X55] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X55)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X55),s('type.num.num',X69)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f33333,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32180,f28786])).
fof(f32180,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1194])).
fof(f33332,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32182,f28786])).
fof(f32182,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1226])).
fof(f33329,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33328,f28786])).
fof(f33328,plain,(
  ( ! [X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32860,f28786])).
fof(f32860,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))),s('type.num.num',X16))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f10801])).
fof(f10801,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))),s('type.num.num',X16))))) )),
  inference(superposition,[],[f1226,f1156])).
fof(f1156,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X69,X55] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X69))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.arithmetic.ADD_SYM')).
fof(f32184,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1007])).
fof(f1007,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f755])).
fof(f755,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X55] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X55)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X55))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|30')).
fof(f33327,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33326,f28786])).
fof(f33326,plain,(
  ( ! [X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32859,f28786])).
fof(f32859,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f10759])).
fof(f10759,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))),s('type.num.num',X16)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))))) )),
  inference(superposition,[],[f1226,f1156])).
fof(f33318,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33317,f28786])).
fof(f33317,plain,(
  ( ! [X64] : (~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f33316,f1622])).
fof(f33316,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))))) ) | $spl36),
  inference(backward_demodulation,[],[f33315,f16568])).
fof(f16568,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X63)),X64,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f16495,f958])).
fof(f958,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X55] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|26')).
fof(f16495,plain,(
  ( ! [X64,X63] : (~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0')))) | s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X63))) = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X63)),X64,bool))) )),
  inference(superposition,[],[f1134,f958])).
fof(f33315,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71)),X72,bool))) ) | $spl36),
  inference(forward_demodulation,[],[f33314,f33146])).
fof(f33146,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl36),
  inference(forward_demodulation,[],[f32473,f948])).
fof(f948,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X55] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|23')).
fof(f32473,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f1165])).
fof(f1165,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X55,X69] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X55))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X69)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|25')).
fof(f33314,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71)),X72,bool))) ) | $spl36),
  inference(subsumption_resolution,[],[f32789,f28786])).
fof(f32789,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X71)),X72,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f16815])).
fof(f16815,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))),s('type.num.num',X71)),X72,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))))))) )),
  inference(forward_demodulation,[],[f16756,f1157])).
fof(f1157,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f197])).
fof(f197,axiom,(
  ! [X69,X55] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X55))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.arithmetic.LESS_EQ')).
fof(f16756,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71)))))) | s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))),s('type.num.num',X71))) = s(bool,sK40('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))),s('type.num.num',X71)),X72,bool))) )),
  inference(superposition,[],[f1134,f1157])).
fof(f33313,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33312,f28786])).
fof(f33312,plain,(
  ( ! [X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32787,f28786])).
fof(f32787,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f16800])).
fof(f16800,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) )),
  inference(backward_demodulation,[],[f16799,f10882])).
fof(f10882,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8))),s('type.num.num',X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) )),
  inference(forward_demodulation,[],[f10798,f1157])).
fof(f10798,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8))))),s('type.num.num',X9))))) )),
  inference(superposition,[],[f1226,f973])).
fof(f973,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X55] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X55)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_add|split|6')).
fof(f16799,plain,(
  ( ! [X8,X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7))),s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f16681,f1157])).
fof(f16681,plain,(
  ( ! [X8,X7] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f1157,f973])).
fof(f33304,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33303,f1622])).
fof(f33303,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl36),
  inference(forward_demodulation,[],[f33302,f958])).
fof(f33302,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) ) | $spl36),
  inference(forward_demodulation,[],[f32764,f33146])).
fof(f32764,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f1157])).
fof(f33171,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33170,f28786])).
fof(f33170,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32567,f28786])).
fof(f32567,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f10881])).
fof(f10881,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num',X7)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) )),
  inference(forward_demodulation,[],[f10797,f1157])).
fof(f10797,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.num.num',X7))))) )),
  inference(superposition,[],[f1226,f974])).
fof(f974,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X55] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|14')).
fof(f33099,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33098,f1622])).
fof(f33098,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl36),
  inference(forward_demodulation,[],[f32411,f958])).
fof(f32411,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl36),
  inference(backward_demodulation,[],[f32184,f966])).
fof(f966,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X55] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_lte|split|2')).
fof(f33097,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33096,f28786])).
fof(f33096,plain,(
  ( ! [X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X16))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32409,f28786])).
fof(f32409,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X16))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))),s('type.num.num',X16))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f22028])).
fof(f22028,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))),s('type.num.num',X16))))) )),
  inference(superposition,[],[f1194,f1156])).
fof(f33093,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33092,f28786])).
fof(f33092,plain,(
  ( ! [X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32401,f28786])).
fof(f32401,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f22122])).
fof(f22122,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) )),
  inference(forward_demodulation,[],[f22121,f16799])).
fof(f22121,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8))),s('type.num.num',X9)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) )),
  inference(forward_demodulation,[],[f22025,f1157])).
fof(f22025,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8))))),s('type.num.num',X9))))) )),
  inference(superposition,[],[f1194,f973])).
fof(f33088,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33087,f28786])).
fof(f33087,plain,(
  ( ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32395,f28786])).
fof(f32395,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num',X5))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f10880])).
fof(f10880,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) )),
  inference(forward_demodulation,[],[f10796,f1157])).
fof(f10796,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))),s('type.num.num',X5))))) )),
  inference(superposition,[],[f1226,f971])).
fof(f971,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f207])).
fof(f207,axiom,(
  ! [X55] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_suc|split|2')).
fof(f33086,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33085,f28786])).
fof(f33085,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32392,f28786])).
fof(f32392,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num',X7))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f22120])).
fof(f22120,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num',X7)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))))) )),
  inference(forward_demodulation,[],[f22024,f1157])).
fof(f22024,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.num.num',X7))))) )),
  inference(superposition,[],[f1194,f974])).
fof(f33084,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33083,f28786])).
fof(f33083,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32390,f28786])).
fof(f32390,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num',X5))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f22119])).
fof(f22119,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))))) )),
  inference(forward_demodulation,[],[f22023,f1157])).
fof(f22023,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))),s('type.num.num',X5))))) )),
  inference(superposition,[],[f1194,f971])).
fof(f33075,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33074,f28786])).
fof(f33074,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32378,f28786])).
fof(f32378,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f22116])).
fof(f22116,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(forward_demodulation,[],[f22021,f1157])).
fof(f22021,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f1194,f933])).
fof(f933,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f209])).
fof(f209,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_suc|split|0')).
fof(f33058,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33057,f1622])).
fof(f33057,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl36),
  inference(forward_demodulation,[],[f32361,f958])).
fof(f32361,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(backward_demodulation,[],[f32184,f16167])).
fof(f16167,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f16097,f961])).
fof(f961,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X55] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_lte|split|1')).
fof(f16097,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f984,f8073])).
fof(f8073,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(backward_demodulation,[],[f8072,f934])).
fof(f934,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f109])).
fof(f109,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f8072,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f7996,f933])).
fof(f7996,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f973,f944])).
fof(f944,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X55] : s('type.num.num',X55) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|1')).
fof(f984,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X55] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X55),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|27')).
fof(f33054,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33053,f28786])).
fof(f33053,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32358,f28786])).
fof(f32358,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f10878])).
fof(f10878,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(forward_demodulation,[],[f10794,f1157])).
fof(f10794,plain,(
  ( ! [X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))))) )),
  inference(superposition,[],[f1226,f933])).
fof(f33039,plain,(
  $false | ($spl15 | $spl36)),
  inference(subsumption_resolution,[],[f32344,f28786])).
fof(f32344,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl15 | $spl36)),
  inference(backward_demodulation,[],[f32184,f8171])).
fof(f8171,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl15),
  inference(subsumption_resolution,[],[f8118,f8083])).
fof(f8083,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl15),
  inference(backward_demodulation,[],[f8072,f5911])).
fof(f8118,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1007,f8073])).
fof(f33015,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f33014,f28786])).
fof(f33014,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32325,f28786])).
fof(f32325,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f21980])).
fof(f21980,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X16)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))))) )),
  inference(superposition,[],[f1194,f1156])).
fof(f32959,plain,(
  $false | ($spl25 | $spl36)),
  inference(subsumption_resolution,[],[f32958,f32184])).
fof(f32958,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.|48|_0') | ($spl25 | $spl36)),
  inference(backward_demodulation,[],[f32957,f5919])).
fof(f32957,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(subsumption_resolution,[],[f32272,f32184])).
fof(f32272,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f5878])).
fof(f5878,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1003,f965])).
fof(f965,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f208])).
fof(f208,axiom,(
  ! [X55] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_suc|split|1')).
fof(f1003,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X55] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X55))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|17')).
fof(f32956,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32955,f28786])).
fof(f32955,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32257,f28786])).
fof(f32257,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f1194])).
fof(f32944,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32943,f28786])).
fof(f32943,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(subsumption_resolution,[],[f32934,f928])).
fof(f32934,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl36),
  inference(backward_demodulation,[],[f32932,f1226])).
fof(f32932,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl36),
  inference(forward_demodulation,[],[f32246,f958])).
fof(f32246,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) ) | $spl36),
  inference(backward_demodulation,[],[f32184,f1158])).
fof(f1158,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X55,X69] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X55),s('type.num.num',X69))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X69)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_lte|split|3')).
fof(f32930,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f32929,f1622])).
fof(f32929,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl36),
  inference(forward_demodulation,[],[f32243,f958])).
fof(f32243,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl36),
  inference(backward_demodulation,[],[f32184,f961])).
fof(f32183,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f28786,f1226])).
fof(f32181,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f28786,f1194])).
fof(f32179,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f28786,f1226])).
fof(f32177,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f28786,f1194])).
fof(f32169,plain,(
  $false | ($spl3 | $spl36)),
  inference(unit_resulting_resolution,[],[f1462,f28786,f1304])).
fof(f32167,plain,(
  $false | ($spl3 | $spl36)),
  inference(unit_resulting_resolution,[],[f1462,f28786,f1302])).
fof(f1302,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f880])).
fof(f32160,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f941,f28786,f1283])).
fof(f941,plain,(
  ( ! [X0] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),'const.PairingHeap.empty_0')) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),'const.PairingHeap.empty_0')),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),'const.PairingHeap.empty_0')),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.empty_def')).
fof(f32158,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f941,f28786,f1086])).
fof(f32156,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1136,f28786,f1087])).
fof(f1087,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK28(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK29(X1,X0))))) )),
  inference(cnf_transformation,[],[f802])).
fof(f1136,plain,(
  ( ! [X2,X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X1)))) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2] : s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X1)))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X29,X28] : ~s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),'const.PairingHeap.Tree_2'(s(X6,X28),s('type.list.list'('type.PairingHeap.heap'(X6)),X29)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.heap_distinct')).
fof(f32150,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f1448])).
fof(f1448,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f1308])).
fof(f1308,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f884])).
fof(f32146,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f6550,f28786,f1446])).
fof(f6550,plain,(
  ( ! [X2,X0,X1] : (sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f6300,f1372])).
fof(f32143,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1541,f28786,f1446])).
fof(f32141,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1539,f28786,f1446])).
fof(f1539,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f930,f1371])).
fof(f32139,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20147,f20170,f28786,f1446])).
fof(f32137,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1608,f28786,f1446])).
fof(f1608,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f943,f1371])).
fof(f32135,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1592,f28786,f1446])).
fof(f1592,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f932,f1371])).
fof(f32133,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f1538,f28786,f1446])).
fof(f32131,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f1537,f28786,f1446])).
fof(f1537,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f930,f1370])).
fof(f32129,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f20169,f28786,f1446])).
fof(f32127,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f1607,f28786,f1446])).
fof(f1607,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f943,f1370])).
fof(f32125,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f1591,f28786,f1446])).
fof(f1591,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f932,f1370])).
fof(f32123,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1536,f28786,f1446])).
fof(f32108,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1549,f28786,f1445])).
fof(f32106,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1547,f28786,f1445])).
fof(f32104,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f20240,f28786,f1445])).
fof(f32102,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1609,f28786,f1445])).
fof(f32100,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1593,f28786,f1445])).
fof(f32098,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1546,f28786,f1445])).
fof(f32096,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1545,f28786,f1445])).
fof(f32094,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1815,f28786,f1445])).
fof(f32092,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1966,f1543,f28786,f1445])).
fof(f32090,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1544,f28786,f1445])).
fof(f32077,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1103,f28786,f1444])).
fof(f32075,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1245,f28786,f1444])).
fof(f1245,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f537])).
fof(f537,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f536])).
fof(f536,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f261])).
fof(f261,axiom,(
  ! [X11,X6] : ! [X9,X17] : s(X6,X9) = s(X6,'const.combin.K_2'(s(X6,X9),s(X11,X17)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.combin.K_THM')).
fof(f32073,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1248,f28786,f1444])).
fof(f1248,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X6,X11] : ! [X21,X27] : s(X11,X21) = s(X11,'const.list.list_CASE_3'(s('type.list.list'(X6),'const.list.NIL_0'),s(X11,X21),s(fun(X6,fun('type.list.list'(X6),X11)),X27)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.list.list_case_def|split|0')).
fof(f32071,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1247,f28786,f1444])).
fof(f1247,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f540])).
fof(f540,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f259])).
fof(f259,axiom,(
  ! [X11,X6] : ! [X9,X17] : s(X6,X9) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X11),'const.pair.,_2'(s(X6,X9),s(X11,X17)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.pair.FST')).
fof(f32069,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1246,f28786,f1444])).
fof(f1246,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f538])).
fof(f538,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f163])).
fof(f163,axiom,(
  ! [X6,X11] : ! [X9,X17] : s(X11,X17) = s(X11,'const.pair.SND_1'(s('type.pair.prod'(X6,X11),'const.pair.,_2'(s(X6,X9),s(X11,X17)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.pair.SND')).
fof(f32067,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1236,f28786,f1444])).
fof(f1236,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f520])).
fof(f520,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X11,X6] : ! [X9] : s(X6,X9) = s(X6,'const.sum.OUTL_1'(s('type.sum.sum'(X6,X11),'const.sum.INL_1'(s(X6,X9)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sum.OUTL')).
fof(f32065,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1235,f28786,f1444])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6,X11] : ! [X9] : s(X11,X9) = s(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X6,X11),'const.sum.INR_1'(s(X11,X9)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.sum.OUTR')).
fof(f32063,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1104,f28786,f1444])).
fof(f1104,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : ! [X19,X27] : s(X6,X19) = s(X6,'const.ind_type.FCONS_3'(s(X6,X19),s(fun('type.num.num',X6),X27),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.ind_type.FCONS|split|0')).
fof(f32061,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1059,f28786,f1444])).
fof(f32059,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1249,f28786,f1444])).
fof(f1249,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.PairingHeap.heap_CASE_3'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.PairingHeap.heap'(X0)),X1)),X3)))) )),
  inference(cnf_transformation,[],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.PairingHeap.heap_CASE_3'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.PairingHeap.heap'(X0)),X1)),X3)))),
  inference(flattening,[],[f544])).
fof(f544,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.PairingHeap.heap_CASE_3'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.PairingHeap.heap'(X0)),X1)),X3)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X11] : ! [X21,X27] : s(X11,X21) = s(X11,'const.PairingHeap.heap_CASE_3'(s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0'),s(X11,X21),s(fun(X6,fun('type.list.list'('type.PairingHeap.heap'(X6)),X11)),X27)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.heap_case_def|split|0')).
fof(f32057,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1269,f28786,f1444])).
fof(f1269,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK59(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f873])).
fof(f873,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK59(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK59(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK59(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK59])],[f569])).
fof(f569,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f568])).
fof(f568,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X11,X6] : ! [X57,X58] : ? [X59] : (s(X11,X57) = s(X11,happ(s(fun('type.list.list'(X6),X11),X59),s('type.list.list'(X6),'const.list.NIL_0'))) & ! [X28,X29] : s(X11,happ(s(fun('type.list.list'(X6),X11),X59),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X28),s('type.list.list'(X6),X29))))) = s(X11,happ(s(fun(X11,X11),happ(s(fun('type.list.list'(X6),fun(X11,X11)),happ(s(fun(X6,fun('type.list.list'(X6),fun(X11,X11))),X58),s(X6,X28))),s('type.list.list'(X6),X29))),s(X11,happ(s(fun('type.list.list'(X6),X11),X59),s('type.list.list'(X6),X29))))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.list.list_Axiom')).
fof(f32055,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1479,f28786,f1444])).
fof(f1479,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f1252,f1245])).
fof(f1252,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f262])).
fof(f262,axiom,(
  ! [X6,X11] : ! [X34,X35] : s(X6,'const.combin.K_2'(s(X6,X34),s(X11,X35))) = s(X6,happ(s(fun(X11,X6),'const.combin.K_1'(s(X6,X34))),s(X11,X35)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','ARITY_1_2_const.combin.K')).
fof(f32053,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1470,f28786,f1444])).
fof(f32051,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1061,f28786,f1444])).
fof(f32049,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1105,f28786,f1444])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.PairingHeap.find_min_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X1),s('type.list.list'('type.PairingHeap.heap'(X0)),X2)))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.PairingHeap.find_min_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X1),s('type.list.list'('type.PairingHeap.heap'(X0)),X2)))))),
  inference(flattening,[],[f409])).
fof(f409,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.PairingHeap.find_min_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X1),s('type.list.list'('type.PairingHeap.heap'(X0)),X2)))))),
  inference(rectify,[],[f198])).
fof(f198,axiom,(
  ! [X6] : ! [X9,X10] : s(X6,X9) = s(X6,'const.PairingHeap.find_min_1'(s('type.PairingHeap.heap'(X6),'const.PairingHeap.Tree_2'(s(X6,X9),s('type.list.list'('type.PairingHeap.heap'(X6)),X10)))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.find_min_def')).
fof(f32032,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1464,f19032,f28786,f1444])).
fof(f32030,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1464,f5716,f28786,f1444])).
fof(f32028,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1464,f5343,f28786,f1444])).
fof(f32026,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1067,f28786,f1444])).
fof(f1067,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f365])).
fof(f365,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f190])).
fof(f190,axiom,(
  ! [X6] : ! [X9] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.list.MEM|split|0')).
fof(f32024,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f947,f28786,f1444])).
fof(f947,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f135])).
fof(f135,axiom,(
  ! [X55] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|20')).
fof(f32022,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f948,f28786,f1444])).
fof(f32020,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f1674,f28786,f1444])).
fof(f32018,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f966,f28786,f1444])).
fof(f32016,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f961,f28786,f1444])).
fof(f32014,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1464,f1672,f28786,f1444])).
fof(f31959,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1969,f1464,f28786,f28786,f1393])).
fof(f31941,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1969,f1464,f928,f28786,f1393])).
fof(f31922,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1969,f1464,f28786,f28786,f1393])).
fof(f31904,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5465,f1464,f928,f28786,f1393])).
fof(f31883,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5498,f28786,f28786,f1392])).
fof(f31881,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f928,f28786,f1392])).
fof(f31870,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5498,f28786,f28786,f1392])).
fof(f31868,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5498,f928,f28786,f1392])).
fof(f31857,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1464,f28786,f1391])).
fof(f31846,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1464,f1464,f28786,f1390])).
fof(f31834,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5481,f28786,f1388])).
fof(f31832,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1544,f28786,f1464,f28786,f1386])).
fof(f31830,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1544,f928,f1464,f28786,f1386])).
fof(f31798,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1544,f28786,f1464,f28786,f1386])).
fof(f31796,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f928,f1464,f28786,f1386])).
fof(f31761,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1549,f1464,f28786,f1385])).
fof(f31756,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1544,f1464,f28786,f1385])).
fof(f31720,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1549,f1464,f28786,f1384])).
fof(f31715,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1544,f1464,f28786,f1384])).
fof(f31679,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1549,f28786,f28786,f1383])).
fof(f31677,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1547,f28786,f28786,f1383])).
fof(f31675,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20184,f20240,f28786,f28786,f1383])).
fof(f31673,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1609,f28786,f28786,f1383])).
fof(f31671,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1593,f28786,f28786,f1383])).
fof(f31669,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1546,f28786,f28786,f1383])).
fof(f31667,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1545,f28786,f28786,f1383])).
fof(f31665,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1966,f1815,f28786,f28786,f1383])).
fof(f31663,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1966,f1543,f28786,f28786,f1383])).
fof(f31661,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1961,f1544,f28786,f28786,f1383])).
fof(f31605,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1546,f6585,f28786,f28786,f1383])).
fof(f31603,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1544,f6577,f28786,f28786,f1383])).
fof(f31601,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1549,f20184,f28786,f28786,f1383])).
fof(f31599,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f1966,f28786,f28786,f1383])).
fof(f31597,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f6579,f28786,f28786,f1383])).
fof(f31595,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f1961,f28786,f28786,f1383])).
fof(f31577,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1546,f6585,f928,f28786,f1383])).
fof(f31575,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1549,f20184,f928,f28786,f1383])).
fof(f31573,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1549,f1966,f928,f28786,f1383])).
fof(f31536,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5465,f1464,f28786,f1380])).
fof(f31518,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5293,f1464,f1464,f28786,f1376])).
fof(f5293,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1541,f1464,f1464,f1368])).
fof(f1368,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f910])).
fof(f31492,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f28786,f1375])).
fof(f31490,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1948,f928,f28786,f1375])).
fof(f31478,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f28786,f1375])).
fof(f31476,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f928,f28786,f1375])).
fof(f31464,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f28786,f1374])).
fof(f31462,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f928,f28786,f1374])).
fof(f31450,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5326,f28786,f28786,f1374])).
fof(f31448,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f928,f28786,f1374])).
fof(f31436,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5309,f28786,f1371])).
fof(f31434,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5309,f28786,f1370])).
fof(f31432,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1541,f1464,f1464,f28786,f1369])).
fof(f31408,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6494,f6550,f1464,f28786,f1367])).
fof(f31405,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1541,f1464,f28786,f1367])).
fof(f31402,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20142,f20169,f1464,f28786,f1367])).
fof(f31398,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1947,f1536,f1464,f28786,f1367])).
fof(f31358,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6550,f6495,f1464,f28786,f1366])).
fof(f31356,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1541,f20147,f1464,f28786,f1366])).
fof(f31354,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1541,f1948,f1464,f28786,f1366])).
fof(f31316,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5293,f1464,f28786,f1363])).
fof(f31293,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5293,f1464,f28786,f1362])).
fof(f31270,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1937,f1464,f28786,f28786,f1361])).
fof(f31253,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1937,f1464,f928,f28786,f1361])).
fof(f31234,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1937,f1464,f28786,f28786,f1361])).
fof(f31217,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1943,f1464,f928,f28786,f1361])).
fof(f31198,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f28786,f1358])).
fof(f31196,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f28786,f28786,f1358])).
fof(f31184,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f928,f28786,f1358])).
fof(f31182,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f928,f28786,f1358])).
fof(f31160,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f28786,f1358])).
fof(f31158,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f928,f28786,f28786,f1358])).
fof(f31146,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f928,f28786,f1358])).
fof(f31144,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5135,f928,f928,f28786,f1358])).
fof(f31120,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f28786,f28786,f28786,f1358])).
fof(f31118,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1933,f928,f28786,f28786,f1358])).
fof(f31106,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f28786,f928,f28786,f1358])).
fof(f31104,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1934,f928,f928,f28786,f1358])).
fof(f31082,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5118,f28786,f1353])).
fof(f31080,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f28786,f1464,f28786,f1352])).
fof(f31078,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f928,f1464,f28786,f1352])).
fof(f31041,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1464,f28786,f28786,f1352])).
fof(f31024,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1464,f928,f28786,f1352])).
fof(f31002,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1534,f20066,f28786,f28786,f1348])).
fof(f31000,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1531,f20065,f28786,f28786,f1348])).
fof(f30994,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6468,f28786,f28786,f1348])).
fof(f30992,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1934,f28786,f28786,f1348])).
fof(f30990,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6464,f28786,f28786,f1348])).
fof(f30988,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1933,f28786,f28786,f1348])).
fof(f30971,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1534,f20066,f928,f28786,f1348])).
fof(f30966,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6468,f928,f28786,f1348])).
fof(f30964,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1934,f928,f28786,f1348])).
fof(f30922,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1534,f20066,f28786,f28786,f1348])).
fof(f30920,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1531,f20065,f28786,f28786,f1348])).
fof(f30914,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6482,f6468,f28786,f28786,f1348])).
fof(f30912,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1534,f1934,f28786,f28786,f1348])).
fof(f30910,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6464,f28786,f28786,f1348])).
fof(f30908,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1933,f28786,f28786,f1348])).
fof(f30891,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1531,f20065,f928,f28786,f1348])).
fof(f30886,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6481,f6464,f928,f28786,f1348])).
fof(f30884,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1531,f1933,f928,f28786,f1348])).
fof(f30842,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1943,f1464,f28786,f1347])).
fof(f30825,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1937,f1464,f28786,f1346])).
fof(f30808,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3466,f1623,f28786,f1343])).
fof(f3466,plain,(
  sP6('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f1464,f1623,f1341])).
fof(f30806,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3465,f1622,f28786,f1343])).
fof(f3465,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1464,f1622,f1341])).
fof(f30804,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3464,f1621,f28786,f1343])).
fof(f3464,plain,(
  sP6('F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1621,f1341])).
fof(f30802,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3469,f1623,f28786,f1343])).
fof(f3469,plain,(
  sP6('F_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f1464,f1623,f1341])).
fof(f30800,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3468,f1622,f28786,f1343])).
fof(f30798,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3467,f1621,f28786,f1343])).
fof(f3467,plain,(
  sP6('F_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1464,f1621,f1341])).
fof(f30728,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3559,f1103,f28786,f1342])).
fof(f30719,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3555,f1059,f28786,f1342])).
fof(f30696,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28261,f19032,f28786,f1342])).
fof(f30694,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5807,f5716,f28786,f1342])).
fof(f30692,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5368,f5343,f28786,f1342])).
fof(f30687,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3551,f1674,f28786,f1342])).
fof(f30683,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3549,f1672,f28786,f1342])).
fof(f30680,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3547,f1639,f28786,f1342])).
fof(f30676,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3544,f1060,f28786,f1342])).
fof(f30672,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3543,f1103,f28786,f1342])).
fof(f30663,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3539,f1059,f28786,f1342])).
fof(f30640,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28254,f19032,f28786,f1342])).
fof(f30638,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5804,f5716,f28786,f1342])).
fof(f30636,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f5365,f5343,f28786,f1342])).
fof(f30631,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3535,f1674,f28786,f1342])).
fof(f30627,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3533,f1672,f28786,f1342])).
fof(f30624,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3531,f1639,f28786,f1342])).
fof(f30620,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3528,f1060,f28786,f1342])).
fof(f32238,plain,(
  $false | $spl36),
  inference(trivial_inequality_removal,[],[f30591])).
fof(f30591,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f5097,f28786,f1342])).
fof(f5097,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f5080,f1322])).
fof(f1322,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f891])).
fof(f5080,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1464,f1337])).
fof(f1337,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f899])).
fof(f32237,plain,(
  $false | $spl36),
  inference(trivial_inequality_removal,[],[f30587])).
fof(f30587,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl36),
  inference(unit_resulting_resolution,[],[f5070,f28786,f1342])).
fof(f5070,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f5053,f1323])).
fof(f5053,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f928,f1464,f928,f1333])).
fof(f30575,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5080,f28786,f1340])).
fof(f30573,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5080,f28786,f1339])).
fof(f30571,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5080,f28786,f1338])).
fof(f30569,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5053,f28786,f1334])).
fof(f30567,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f28786,f1464,f28786,f1333])).
fof(f30565,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f928,f1464,f28786,f1333])).
fof(f30528,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1517,f1464,f28786,f28786,f1333])).
fof(f30511,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1520,f1464,f928,f28786,f1333])).
fof(f30489,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f5023,f28786,f1330])).
fof(f5023,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f928,f1464,f928,f1329])).
fof(f30487,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1509,f28786,f1464,f28786,f1329])).
fof(f30485,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f4299,f928,f1464,f28786,f1329])).
fof(f30418,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1509,f1464,f28786,f28786,f1329])).
fof(f30388,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1512,f1464,f928,f28786,f1329])).
fof(f29936,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f1509,f19981,f28786,f1326])).
fof(f29934,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f1509,f19981,f28786,f1326])).
fof(f29927,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f19954,f19981,f28786,f1326])).
fof(f29921,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f19954,f19981,f28786,f1326])).
fof(f29734,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1524,f28786,f1326])).
fof(f29732,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1524,f28786,f1326])).
fof(f29719,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1524,f28786,f1326])).
fof(f29717,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1524,f28786,f1326])).
fof(f29710,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1524,f28786,f1326])).
fof(f29708,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1524,f28786,f1326])).
fof(f29704,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1509,f1524,f28786,f1326])).
fof(f29700,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1524,f28786,f1326])).
fof(f29698,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1524,f28786,f1326])).
fof(f29694,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1524,f28786,f1326])).
fof(f29691,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1524,f28786,f1326])).
fof(f29689,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1524,f28786,f1326])).
fof(f29687,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1524,f28786,f1326])).
fof(f29684,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1506,f1524,f28786,f1326])).
fof(f29680,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1524,f28786,f1326])).
fof(f29559,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1509,f19980,f28786,f1326])).
fof(f29557,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f1509,f19980,f28786,f1326])).
fof(f29526,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f1517,f1506,f19980,f28786,f1326])).
fof(f29524,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f20060,f19979,f1506,f19980,f28786,f1326])).
fof(f29382,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1522,f28786,f1326])).
fof(f29380,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1522,f28786,f1326])).
fof(f29365,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1522,f28786,f1326])).
fof(f29363,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1522,f28786,f1326])).
fof(f29355,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1522,f28786,f1326])).
fof(f29353,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1522,f28786,f1326])).
fof(f29345,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1522,f28786,f1326])).
fof(f29343,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1522,f28786,f1326])).
fof(f29329,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1517,f19954,f1522,f28786,f1326])).
fof(f29327,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f19890,f1514,f19954,f1522,f28786,f1326])).
fof(f29312,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1522,f28786,f1326])).
fof(f29309,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1522,f28786,f1326])).
fof(f29307,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1522,f28786,f1326])).
fof(f29305,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1522,f28786,f1326])).
fof(f29303,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1522,f28786,f1326])).
fof(f29149,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f11162,f28786,f1464,f28786,f1325])).
fof(f29147,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f11162,f928,f1464,f28786,f1325])).
fof(f29035,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f11162,f1464,f28786,f28786,f1325])).
fof(f28967,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f3468,f1464,f928,f28786,f1325])).
fof(f28938,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1311])).
fof(f28936,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f928,f28786,f1311])).
fof(f28928,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1311])).
fof(f28926,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1905,f928,f28786,f1311])).
fof(f28918,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1310])).
fof(f28916,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f928,f28786,f1310])).
fof(f28908,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1310])).
fof(f28899,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1309])).
fof(f28897,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1905,f928,f28786,f1309])).
fof(f28889,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1891,f28786,f28786,f1309])).
fof(f28880,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1700,f1464,f28786,f1220])).
fof(f1700,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1215])).
fof(f1215,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f851])).
fof(f28870,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1700,f1464,f28786,f1217])).
fof(f28860,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1704,f28786,f28786,f1216])).
fof(f28858,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1704,f928,f28786,f1216])).
fof(f28844,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1704,f28786,f28786,f1216])).
fof(f28842,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1716,f928,f28786,f1216])).
fof(f28828,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1762,f28786,f1202])).
fof(f1202,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f844])).
fof(f844,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f843])).
fof(f843,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f725])).
fof(f725,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1762,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1464,f1464,f1478])).
fof(f1478,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1200,f1464])).
fof(f1200,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f844])).
fof(f28826,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1762,f28786,f1201])).
fof(f1201,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f844])).
fof(f28824,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1623,f28786,f1189])).
fof(f28822,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1622,f28786,f1189])).
fof(f28820,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28786,f1621,f28786,f1189])).
fof(f28818,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1623,f28786,f1189])).
fof(f28816,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1622,f28786,f1189])).
fof(f28814,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f928,f1621,f28786,f1189])).
fof(f28810,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1623,f28786,f1189])).
fof(f28808,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f28786,f1622,f28786,f1189])).
fof(f28806,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f28786,f1621,f28786,f1189])).
fof(f28804,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1623,f28786,f1189])).
fof(f28802,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f928,f1622,f28786,f1189])).
fof(f28800,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f928,f1621,f28786,f1189])).
fof(f28795,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1672,f28786,f999])).
fof(f28793,plain,(
  $false | ($spl1 | $spl36)),
  inference(unit_resulting_resolution,[],[f1621,f28786,f996])).
fof(f28791,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f1464,f28786])).
fof(f28788,plain,(
  ( ! [X9] : (~p(s(bool,sK36('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) ) | $spl38),
  inference(cnf_transformation,[],[f28788_D])).
fof(f28788_D,plain,(
  ( ! [X9] : (~p(s(bool,sK36('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f28787,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | ~p(s(bool,sK36('const.combin.I_0',X9,bool))) | ~p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f28760,f1470])).
fof(f28760,plain,(
  ( ! [X10,X9] : (~p(s(bool,sK36('const.combin.I_0',X9,bool))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X10)))) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f1126,f1470])).
fof(f1126,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f826])).
fof(f28785,plain,(
  ( ! [X7] : (~p(s(bool,sK36('_i_0_0',X7,bool))) | ~p(s(bool,X7))) ) | $spl34),
  inference(cnf_transformation,[],[f28785_D])).
fof(f28785_D,plain,(
  ( ! [X7] : (~p(s(bool,sK36('_i_0_0',X7,bool))) | ~p(s(bool,X7))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f28784,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | ~p(s(bool,sK36('_i_0_0',X7,bool))) | ~p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f28759,f1061])).
fof(f28759,plain,(
  ( ! [X8,X7] : (~p(s(bool,sK36('_i_0_0',X7,bool))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X8)))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f1126,f1061])).
fof(f15733,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12042,f12026])).
fof(f12042,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1621,f12026,f1186])).
fof(f15732,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12044,f1464])).
fof(f12044,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1622,f12026,f1186])).
fof(f15731,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12046,f12026])).
fof(f12046,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1623,f12026,f1186])).
fof(f15730,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12056,f12026])).
fof(f12056,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1621,f12026,f1186])).
fof(f15729,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12058,f1464])).
fof(f12058,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1622,f12026,f1186])).
fof(f15728,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12060,f12026])).
fof(f12060,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1623,f12026,f1186])).
fof(f15727,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12062,f1464])).
fof(f12062,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1704,f12026,f1215])).
fof(f15726,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12063,f12026])).
fof(f12063,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1806,f12026,f1215])).
fof(f15725,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12064,f12026])).
fof(f12064,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6325,f12026,f1215])).
fof(f15724,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12066,f12026])).
fof(f12066,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1894,f12026,f1215])).
fof(f15723,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12067,f12026])).
fof(f12067,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1925,f12026,f1215])).
fof(f15722,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12070,f1464])).
fof(f12070,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1870,f12026,f1215])).
fof(f15721,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12071,f1464])).
fof(f12071,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1920,f12026,f1215])).
fof(f15720,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12079,f1464])).
fof(f12079,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1716,f12026,f1215])).
fof(f15719,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12080,f1464])).
fof(f12080,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1894,f12026,f1215])).
fof(f15718,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12081,f1464])).
fof(f12081,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1925,f12026,f1215])).
fof(f15717,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12082,f12026])).
fof(f12082,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6326,f12026,f1215])).
fof(f15716,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12083,f12026])).
fof(f12083,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1807,f12026,f1215])).
fof(f15715,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12084,f12026])).
fof(f12084,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1870,f12026,f1215])).
fof(f15714,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12085,f12026])).
fof(f12085,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1920,f12026,f1215])).
fof(f15713,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12106,f12026])).
fof(f12106,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f1306])).
fof(f1551,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1448])).
fof(f15712,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12107,f12026])).
fof(f12107,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1552,f12026,f1306])).
fof(f1552,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f930,f1448])).
fof(f15711,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12108,f12026])).
fof(f12108,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1594,f12026,f1306])).
fof(f1594,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f932,f1448])).
fof(f15710,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12109,f12026])).
fof(f12109,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1610,f12026,f1306])).
fof(f1610,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f943,f1448])).
fof(f15709,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12116,f12026])).
fof(f12116,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1552,f12026,f1306])).
fof(f15708,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12117,f12026])).
fof(f12117,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1594,f12026,f1306])).
fof(f15707,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12118,f12026])).
fof(f12118,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1610,f12026,f1306])).
fof(f15706,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12119,f12026])).
fof(f12119,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1930,f12026,f1306])).
fof(f1930,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f928,f1311])).
fof(f15705,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12125,f12026])).
fof(f12125,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1552,f12026,f1307])).
fof(f15704,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12126,f12026])).
fof(f12126,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1594,f12026,f1307])).
fof(f15703,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12127,f12026])).
fof(f12127,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1610,f12026,f1307])).
fof(f15702,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12128,f12026])).
fof(f12128,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1930,f12026,f1307])).
fof(f15701,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12133,f12026])).
fof(f12133,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f1307])).
fof(f15700,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12134,f12026])).
fof(f12134,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1552,f12026,f1307])).
fof(f15699,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12135,f12026])).
fof(f12135,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1594,f12026,f1307])).
fof(f15698,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12136,f12026])).
fof(f12136,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1610,f12026,f1307])).
fof(f15697,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12178,f13468])).
fof(f13468,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1514,f1506,f1524,f12026,f1328])).
fof(f12178,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f928,f12026,f1325])).
fof(f15696,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12266,f13980])).
fof(f13980,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1332])).
fof(f12266,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1522,f12026,f1327])).
fof(f15695,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12267,f13980])).
fof(f12267,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1522,f12026,f1327])).
fof(f15694,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12268,f13980])).
fof(f12268,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1522,f12026,f1327])).
fof(f15693,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12269,f13980])).
fof(f12269,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1522,f12026,f1327])).
fof(f15692,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12273,f13980])).
fof(f12273,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1522,f12026,f1327])).
fof(f15691,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12274,f13980])).
fof(f12274,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1522,f12026,f1327])).
fof(f15690,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12275,f13980])).
fof(f12275,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1522,f12026,f1327])).
fof(f15689,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12276,f13980])).
fof(f12276,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1522,f12026,f1327])).
fof(f15688,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12313,f14004])).
fof(f14004,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1336])).
fof(f12313,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1522,f12026,f1327])).
fof(f15687,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12314,f14004])).
fof(f12314,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1522,f12026,f1327])).
fof(f15686,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12315,f14004])).
fof(f12315,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1522,f12026,f1327])).
fof(f15685,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12316,f14004])).
fof(f12316,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1522,f12026,f1327])).
fof(f15684,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12346,f14004])).
fof(f12346,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1522,f12026,f1327])).
fof(f15683,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12356,f14004])).
fof(f12356,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1522,f12026,f1327])).
fof(f15682,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12366,f14004])).
fof(f12366,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1522,f12026,f1327])).
fof(f15681,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12383,f14004])).
fof(f12383,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1522,f12026,f1327])).
fof(f15680,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12570,f13980])).
fof(f12570,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1524,f12026,f1327])).
fof(f15679,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12571,f13980])).
fof(f12571,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1524,f12026,f1327])).
fof(f15678,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12572,f13980])).
fof(f12572,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1524,f12026,f1327])).
fof(f15677,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12573,f13980])).
fof(f12573,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1524,f12026,f1327])).
fof(f15676,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12574,f13980])).
fof(f12574,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1524,f12026,f1327])).
fof(f15675,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12575,f13980])).
fof(f12575,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1524,f12026,f1327])).
fof(f15674,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12576,f13980])).
fof(f12576,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1524,f12026,f1327])).
fof(f15673,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12577,f13980])).
fof(f12577,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1524,f12026,f1327])).
fof(f15672,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12578,f14004])).
fof(f12578,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1524,f12026,f1327])).
fof(f15671,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12579,f14004])).
fof(f12579,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1524,f12026,f1327])).
fof(f15670,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12580,f14004])).
fof(f12580,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1524,f12026,f1327])).
fof(f15669,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12581,f14004])).
fof(f12581,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1524,f12026,f1327])).
fof(f15668,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12601,f14004])).
fof(f12601,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1524,f12026,f1327])).
fof(f15667,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f12609,f14004])).
fof(f12609,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1524,f12026,f1327])).
fof(f15666,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12617,f14004])).
fof(f12617,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1524,f12026,f1327])).
fof(f15665,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12630,f14004])).
fof(f12630,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1524,f12026,f1327])).
fof(f15664,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12671,f14004])).
fof(f12671,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1585,f1588,f12026,f1327])).
fof(f15663,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12709,f14004])).
fof(f12709,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1588,f12026,f1327])).
fof(f15662,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12736,f14004])).
fof(f12736,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1601,f1604,f12026,f1327])).
fof(f15661,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12774,f14004])).
fof(f12774,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1604,f12026,f1327])).
fof(f15660,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12802,f14004])).
fof(f12802,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1507,f1525,f12026,f1327])).
fof(f15659,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12841,f14004])).
fof(f12841,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1525,f12026,f1327])).
fof(f15658,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12942,f14004])).
fof(f12942,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1589,f12026,f1327])).
fof(f15657,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12987,f14004])).
fof(f12987,plain,(
  ( ! [X0] : (~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1605,f12026,f1327])).
fof(f15656,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13034,f14004])).
fof(f13034,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1527,f12026,f1327])).
fof(f15655,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13147,f13980])).
fof(f13147,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1522,f12026,f1328])).
fof(f15654,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13148,f13980])).
fof(f13148,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1522,f12026,f1328])).
fof(f15653,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13149,f13980])).
fof(f13149,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1522,f12026,f1328])).
fof(f15652,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13150,f13980])).
fof(f13150,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1522,f12026,f1328])).
fof(f15651,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13154,f13980])).
fof(f13154,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1522,f12026,f1328])).
fof(f15650,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13155,f13980])).
fof(f13155,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1522,f12026,f1328])).
fof(f15649,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13156,f13980])).
fof(f13156,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1522,f12026,f1328])).
fof(f15648,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13157,f13980])).
fof(f13157,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1522,f12026,f1328])).
fof(f15647,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13194,f14004])).
fof(f13194,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1522,f12026,f1328])).
fof(f15646,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13195,f14004])).
fof(f13195,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1522,f12026,f1328])).
fof(f15645,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13196,f14004])).
fof(f13196,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1522,f12026,f1328])).
fof(f15644,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13197,f14004])).
fof(f13197,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1522,f12026,f1328])).
fof(f15643,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13227,f14004])).
fof(f13227,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1522,f12026,f1328])).
fof(f15642,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13237,f14004])).
fof(f13237,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1522,f12026,f1328])).
fof(f15641,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13247,f14004])).
fof(f13247,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1522,f12026,f1328])).
fof(f15640,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13264,f14004])).
fof(f13264,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1522,f12026,f1328])).
fof(f15639,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13301,f14247])).
fof(f14247,plain,(
  ( ! [X0] : (~sP6(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0')) ) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1672,f12026,f1344])).
fof(f13301,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1519,f1509,f1523,f12026,f1328])).
fof(f15638,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13304,f14251])).
fof(f14251,plain,(
  ( ! [X0] : (~sP6(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1674,f12026,f1344])).
fof(f13304,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1521,f1509,f1523,f12026,f1328])).
fof(f15637,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13360,f14247])).
fof(f13360,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1519,f1509,f1587,f12026,f1328])).
fof(f1587,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f932,f1338])).
fof(f15636,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13363,f14251])).
fof(f13363,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1521,f1509,f1587,f12026,f1328])).
fof(f15635,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13419,f14247])).
fof(f13419,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'F_0')) ) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1519,f1509,f1603,f12026,f1328])).
fof(f1603,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f943,f1338])).
fof(f15634,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13422,f14251])).
fof(f13422,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1521,f1509,f1603,f12026,f1328])).
fof(f15633,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13451,f13980])).
fof(f13451,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1524,f12026,f1328])).
fof(f15632,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13452,f13980])).
fof(f13452,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f1524,f12026,f1328])).
fof(f15631,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13453,f13980])).
fof(f13453,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1524,f12026,f1328])).
fof(f15630,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13454,f13980])).
fof(f13454,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1524,f12026,f1328])).
fof(f15629,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13455,f13980])).
fof(f13455,plain,(
  ~sP5('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1524,f12026,f1328])).
fof(f15628,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13456,f13980])).
fof(f13456,plain,(
  ~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1524,f12026,f1328])).
fof(f15627,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13457,f13980])).
fof(f13457,plain,(
  ( ! [X0] : (~sP5('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1524,f12026,f1328])).
fof(f15626,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13458,f13980])).
fof(f13458,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1524,f12026,f1328])).
fof(f15625,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13459,f14004])).
fof(f13459,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1506,f1524,f12026,f1328])).
fof(f15624,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13460,f14004])).
fof(f13460,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1506,f1524,f12026,f1328])).
fof(f15623,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13461,f14004])).
fof(f13461,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1506,f1524,f12026,f1328])).
fof(f15622,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13462,f14004])).
fof(f13462,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1506,f1524,f12026,f1328])).
fof(f15621,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13482,f14004])).
fof(f13482,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1509,f1524,f12026,f1328])).
fof(f15620,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f13490,f14004])).
fof(f13490,plain,(
  ~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1508,f1524,f12026,f1328])).
fof(f15619,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13498,f14004])).
fof(f13498,plain,(
  ( ! [X0] : (~sP4('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6359,f1524,f12026,f1328])).
fof(f15618,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13501,f6345])).
fof(f13501,plain,(
  ( ! [X0] : (sP6('T_0','T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1514,f6359,f1524,f12026,f1328])).
fof(f15617,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13511,f14004])).
fof(f13511,plain,(
  ~sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1510,f1524,f12026,f1328])).
fof(f15616,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13552,f14004])).
fof(f13552,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1585,f1588,f12026,f1328])).
fof(f15615,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13590,f14004])).
fof(f13590,plain,(
  ~sP4('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1588,f12026,f1328])).
fof(f15614,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13617,f14004])).
fof(f13617,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1601,f1604,f12026,f1328])).
fof(f15613,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13655,f14004])).
fof(f13655,plain,(
  ( ! [X0] : (~sP4('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1604,f12026,f1328])).
fof(f15612,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13683,f14004])).
fof(f13683,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1507,f1525,f12026,f1328])).
fof(f15611,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13722,f14004])).
fof(f13722,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3555,f1512,f1525,f12026,f1328])).
fof(f15610,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13823,f14004])).
fof(f13823,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1589,f12026,f1328])).
fof(f15609,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13868,f14004])).
fof(f13868,plain,(
  ( ! [X0] : (~sP4('const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))),'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1605,f12026,f1328])).
fof(f15608,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13897,f14281])).
fof(f14281,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1639,f12026,f1344])).
fof(f13897,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f1515,f1506,f1527,f12026,f1328])).
fof(f15607,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13898,f14281])).
fof(f13898,plain,(
  sP6('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f1586,f1506,f1527,f12026,f1328])).
fof(f15606,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13899,f14281])).
fof(f13899,plain,(
  ( ! [X0] : (sP6('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1602,f1506,f1527,f12026,f1328])).
fof(f15605,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13915,f14004])).
fof(f13915,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3539,f1512,f1527,f12026,f1328])).
fof(f15604,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13976,f13980])).
fof(f13976,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f928,f12026,f1329])).
fof(f15603,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14000,f14004])).
fof(f14000,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f928,f12026,f1333])).
fof(f15602,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14007,f12026])).
fof(f14007,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f1464,f12026,f1337])).
fof(f15601,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14008,f12026])).
fof(f14008,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f1464,f12026,f1337])).
fof(f15600,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14009,f12026])).
fof(f14009,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f1464,f12026,f1337])).
fof(f15599,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14010,f12026])).
fof(f14010,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f1464,f12026,f1337])).
fof(f1526,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1340])).
fof(f15598,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14011,f12026])).
fof(f14011,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f1464,f12026,f1337])).
fof(f15597,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14012,f12026])).
fof(f14012,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f1464,f12026,f1337])).
fof(f15596,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14013,f12026])).
fof(f14013,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f1464,f12026,f1337])).
fof(f15595,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14014,f12026])).
fof(f14014,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f1337])).
fof(f15594,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14015,f12026])).
fof(f14015,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15593,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14016,f12026])).
fof(f14016,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15592,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14017,f12026])).
fof(f14017,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15591,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14019,f12026])).
fof(f14019,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15590,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14020,f12026])).
fof(f14020,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15589,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14021,f12026])).
fof(f14021,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15588,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14022,f12026])).
fof(f14022,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f12026,f12026,f1337])).
fof(f15587,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14023,f12026])).
fof(f14023,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15586,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14024,f12026])).
fof(f14024,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15585,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14025,f12026])).
fof(f14025,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15584,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14027,f12026])).
fof(f14027,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f1464,f12026,f1337])).
fof(f15583,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14028,f12026])).
fof(f14028,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f1464,f12026,f1337])).
fof(f15582,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14029,f12026])).
fof(f14029,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f1464,f12026,f1337])).
fof(f15581,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14030,f12026])).
fof(f14030,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f1464,f12026,f1337])).
fof(f15580,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14031,f12026])).
fof(f14031,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f1464,f12026,f1337])).
fof(f15579,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14032,f12026])).
fof(f14032,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f1464,f12026,f1337])).
fof(f15578,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14033,f12026])).
fof(f14033,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f1464,f12026,f1337])).
fof(f15577,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14039,f12026])).
fof(f14039,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15576,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14040,f12026])).
fof(f14040,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15575,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14041,f12026])).
fof(f14041,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15574,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14042,f12026])).
fof(f14042,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f12026,f12026,f1337])).
fof(f15573,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14043,f12026])).
fof(f14043,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15572,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14044,f12026])).
fof(f14044,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15571,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14045,f12026])).
fof(f14045,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15570,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14046,f12026])).
fof(f14046,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f12026,f12026,f1337])).
fof(f15569,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14047,f12026])).
fof(f14047,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15568,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14048,f12026])).
fof(f14048,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15567,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14049,f12026])).
fof(f14049,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15566,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14055,f12026])).
fof(f14055,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f1464,f12026,f1337])).
fof(f15565,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14056,f12026])).
fof(f14056,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f1464,f12026,f1337])).
fof(f15564,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14057,f12026])).
fof(f14057,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f1464,f12026,f1337])).
fof(f15563,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14058,f12026])).
fof(f14058,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f1464,f12026,f1337])).
fof(f15562,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14059,f12026])).
fof(f14059,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f1464,f12026,f1337])).
fof(f15561,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14060,f12026])).
fof(f14060,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f1464,f12026,f1337])).
fof(f15560,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14061,f12026])).
fof(f14061,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f1464,f12026,f1337])).
fof(f15559,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14063,f12026])).
fof(f14063,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15558,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14064,f12026])).
fof(f14064,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15557,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14065,f12026])).
fof(f14065,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15556,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14066,f12026])).
fof(f14066,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f12026,f12026,f1337])).
fof(f15555,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14067,f12026])).
fof(f14067,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15554,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14068,f12026])).
fof(f14068,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15553,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14069,f12026])).
fof(f14069,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15552,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14070,f12026])).
fof(f14070,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f12026,f12026,f1337])).
fof(f15551,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14071,f12026])).
fof(f14071,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15550,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14072,f12026])).
fof(f14072,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15549,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14073,f12026])).
fof(f14073,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15548,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14074,f12026])).
fof(f14074,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f1464,f12026,f1337])).
fof(f15547,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14075,f12026])).
fof(f14075,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f1464,f12026,f1337])).
fof(f15546,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14076,f12026])).
fof(f14076,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f1464,f12026,f1337])).
fof(f15545,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14077,f12026])).
fof(f14077,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f1464,f12026,f1337])).
fof(f15544,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14079,f12026])).
fof(f14079,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f1464,f12026,f1337])).
fof(f15543,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14080,f12026])).
fof(f14080,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f1464,f12026,f1337])).
fof(f15542,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14081,f12026])).
fof(f14081,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f1464,f12026,f1337])).
fof(f15541,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14086,f12026])).
fof(f14086,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f1337])).
fof(f15540,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14087,f12026])).
fof(f14087,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15539,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14088,f12026])).
fof(f14088,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15538,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14089,f12026])).
fof(f14089,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15537,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14090,f12026])).
fof(f14090,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f12026,f12026,f1337])).
fof(f15536,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14091,f12026])).
fof(f14091,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15535,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14092,f12026])).
fof(f14092,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15534,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14093,f12026])).
fof(f14093,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15533,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14095,f12026])).
fof(f14095,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15532,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14096,f12026])).
fof(f14096,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15531,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14097,f12026])).
fof(f14097,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15530,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14102,f12026])).
fof(f14102,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f1464,f12026,f1337])).
fof(f15529,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14103,f12026])).
fof(f14103,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f1464,f12026,f1337])).
fof(f15528,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14104,f12026])).
fof(f14104,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f1464,f12026,f1337])).
fof(f15527,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14105,f12026])).
fof(f14105,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f1464,f12026,f1337])).
fof(f15526,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14107,f12026])).
fof(f14107,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f1464,f12026,f1337])).
fof(f15525,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14108,f12026])).
fof(f14108,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f1464,f12026,f1337])).
fof(f15524,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14109,f12026])).
fof(f14109,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f1464,f12026,f1337])).
fof(f15523,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14110,f12026])).
fof(f14110,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f1337])).
fof(f15522,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14111,f12026])).
fof(f14111,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15521,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14112,f12026])).
fof(f14112,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15520,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14113,f12026])).
fof(f14113,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15519,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14115,f12026])).
fof(f14115,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15518,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14116,f12026])).
fof(f14116,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15517,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14117,f12026])).
fof(f14117,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15516,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14118,f12026])).
fof(f14118,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f12026,f12026,f1337])).
fof(f15515,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14119,f12026])).
fof(f14119,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15514,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14120,f12026])).
fof(f14120,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15513,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14121,f12026])).
fof(f14121,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15512,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14122,f12026])).
fof(f14122,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f1464,f12026,f1337])).
fof(f15511,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14123,f12026])).
fof(f14123,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f1464,f12026,f1337])).
fof(f15510,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14124,f12026])).
fof(f14124,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f1464,f12026,f1337])).
fof(f15509,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14125,f12026])).
fof(f14125,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f1464,f12026,f1337])).
fof(f15508,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14127,f12026])).
fof(f14127,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f1464,f12026,f1337])).
fof(f15507,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14128,f12026])).
fof(f14128,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f1464,f12026,f1337])).
fof(f15506,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14129,f12026])).
fof(f14129,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f1464,f12026,f1337])).
fof(f15505,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14134,f12026])).
fof(f14134,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f1337])).
fof(f15504,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14135,f12026])).
fof(f14135,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1523,f12026,f12026,f1337])).
fof(f15503,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14136,f12026])).
fof(f14136,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1587,f12026,f12026,f1337])).
fof(f15502,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14137,f12026])).
fof(f14137,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1603,f12026,f12026,f1337])).
fof(f15501,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14138,f12026])).
fof(f14138,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f12026,f12026,f1337])).
fof(f15500,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14139,f12026])).
fof(f14139,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1588,f12026,f12026,f1337])).
fof(f15499,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14140,f12026])).
fof(f14140,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1604,f12026,f12026,f1337])).
fof(f15498,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14141,f12026])).
fof(f14141,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1525,f12026,f12026,f1337])).
fof(f15497,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14143,f12026])).
fof(f14143,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1589,f12026,f12026,f1337])).
fof(f15496,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14144,f12026])).
fof(f14144,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1605,f12026,f12026,f1337])).
fof(f15495,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14145,f12026])).
fof(f14145,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1527,f12026,f12026,f1337])).
fof(f15494,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14150,f1621])).
fof(f14150,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f12026,f1341])).
fof(f15493,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14151,f6620])).
fof(f6620,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f6300,f1444])).
fof(f14151,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f12026,f1341])).
fof(f15492,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14152,f1622])).
fof(f14152,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f12026,f1341])).
fof(f15491,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14153,f1623])).
fof(f14153,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f12026,f1341])).
fof(f15490,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14154,f1622])).
fof(f14154,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f3516,f12026,f1341])).
fof(f3516,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f928,f1622,f1343])).
fof(f15489,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14155,f1621])).
fof(f14155,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3514,f12026,f1341])).
fof(f3514,plain,(
  ~sP6('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1621,f1343])).
fof(f15488,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14156,f1623])).
fof(f14156,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f3517,f12026,f1341])).
fof(f3517,plain,(
  ~sP6('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f928,f1623,f1343])).
fof(f15481,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14197,f1103])).
fof(f14197,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK31(X0,bool))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3505,f12026,f1344])).
fof(f15480,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14198,f1059])).
fof(f14198,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3501,f12026,f1344])).
fof(f15479,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14203,f1672])).
fof(f14203,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3496,f12026,f1344])).
fof(f3496,plain,(
  sP6('T_0','F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1672,f1342])).
fof(f15478,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14204,f1674])).
fof(f14204,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f3497,f12026,f1344])).
fof(f3497,plain,(
  sP6('T_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f928,f1674,f1342])).
fof(f15477,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14205,f5343])).
fof(f14205,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5369,f12026,f1344])).
fof(f15476,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14206,f5716])).
fof(f14206,plain,(
  s(bool,'F_0') != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5808,f12026,f1344])).
fof(f15475,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14207,f1672])).
fof(f14207,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3481,f12026,f1344])).
fof(f15474,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15473,f1059])).
fof(f15473,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14208,f1672])).
fof(f14208,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f4304,f12026,f1344])).
fof(f4304,plain,(
  sP6('T_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3501,f1672])).
fof(f15472,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14210,f1059])).
fof(f14210,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3486,f12026,f1344])).
fof(f15471,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15470,f1059])).
fof(f15470,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14211,f1672])).
fof(f14211,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f4085,f12026,f1344])).
fof(f4085,plain,(
  sP6('T_0','const.combin.I_1'(s(bool,'F_0')),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(superposition,[],[f3486,f1672])).
fof(f15469,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15468,f1103])).
fof(f15468,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK31(X0,bool))) ) | $spl30),
  inference(forward_demodulation,[],[f14212,f1059])).
fof(f14212,plain,(
  ( ! [X0] : (s(bool,sK31(X0,bool)) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f4102,f12026,f1344])).
fof(f15467,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15466,f1059])).
fof(f15466,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14213,f5343])).
fof(f14213,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5270,f12026,f1344])).
fof(f15465,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15464,f1059])).
fof(f15464,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14214,f5716])).
fof(f14214,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5390,f12026,f1344])).
fof(f15463,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15462,f1059])).
fof(f15462,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14215,f6279])).
fof(f6279,plain,(
  s(bool,'F_0') = s(bool,sK31(sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),bool)) | $spl1),
  inference(forward_demodulation,[],[f6273,f1059])).
fof(f6273,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,sK31(sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),bool)) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f5829,f1343])).
fof(f5829,plain,(
  sP6('T_0','const.combin.I_1'(s(bool,'F_0')),sK31(sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),bool)) | $spl1),
  inference(superposition,[],[f4102,f5716])).
fof(f14215,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31(sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5829,f12026,f1344])).
fof(f15461,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14216,f1674])).
fof(f14216,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(unit_resulting_resolution,[],[f3482,f12026,f1344])).
fof(f15460,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14218,f1103])).
fof(f14218,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK31(X0,bool))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f3490,f12026,f1344])).
fof(f15459,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14219,f5343])).
fof(f14219,plain,(
  s(bool,'F_0') != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5366,f12026,f1344])).
fof(f15458,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14220,f5716])).
fof(f14220,plain,(
  s(bool,'F_0') != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5805,f12026,f1344])).
fof(f15457,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15456,f1059])).
fof(f15456,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14221,f5343])).
fof(f14221,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5386,f12026,f1344])).
fof(f5386,plain,(
  sP6('T_0',sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3501,f5343])).
fof(f15455,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f15454,f1059])).
fof(f15454,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f14222,f5716])).
fof(f14222,plain,(
  s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool)) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5825,f12026,f1344])).
fof(f5825,plain,(
  sP6('T_0',sK31(sK31('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),bool),bool),'const.combin.I_1'(s(bool,'F_0'))) | $spl1),
  inference(superposition,[],[f3501,f5716])).
fof(f15448,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14239,f1672])).
fof(f14239,plain,(
  s(bool,'F_0') != s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f11172,f12026,f1344])).
fof(f11172,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0','const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f1520,f1508,f1523,f1464,f1327])).
fof(f15447,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14326,f15127])).
fof(f15127,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1528,f12026,f1447])).
fof(f14326,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f1346])).
fof(f15446,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14340,f15127])).
fof(f14340,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f1347])).
fof(f15445,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14360,f14364])).
fof(f14364,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1355])).
fof(f14360,plain,(
  ( ! [X0] : (~sP7('T_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f928,f12026,f1352])).
fof(f15444,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14370,f1464])).
fof(f14370,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5135,f12026,f1356])).
fof(f15443,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14371,f12026])).
fof(f14371,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6493,f12026,f1356])).
fof(f15442,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14376,f1464])).
fof(f14376,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f1356])).
fof(f15441,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14377,f12026])).
fof(f14377,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6464,f12026,f1356])).
fof(f15440,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14378,f12026])).
fof(f14378,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1934,f12026,f1356])).
fof(f15439,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14379,f12026])).
fof(f14379,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6468,f12026,f1356])).
fof(f15438,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14390,f1464])).
fof(f14390,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5135,f12026,f1357])).
fof(f15437,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14391,f12026])).
fof(f14391,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6493,f12026,f1357])).
fof(f15436,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14396,f12026])).
fof(f14396,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f1357])).
fof(f15435,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14397,f12026])).
fof(f14397,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6464,f12026,f1357])).
fof(f15434,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14398,f1464])).
fof(f14398,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1934,f12026,f1357])).
fof(f15433,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14399,f12026])).
fof(f14399,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6468,f12026,f1357])).
fof(f15432,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14439,f14791])).
fof(f14791,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f12026,f12026,f1373])).
fof(f1373,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f915])).
fof(f14439,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f1362])).
fof(f15431,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14458,f14791])).
fof(f14458,plain,(
  ( ! [X0,X1] : (~sP10(X0,'T_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f1363])).
fof(f15430,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14520,f12026])).
fof(f14520,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1541,f1948,f12026,f1368])).
fof(f15429,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14521,f12026])).
fof(f14521,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1538,f1948,f12026,f1368])).
fof(f15428,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14523,f12026])).
fof(f14523,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1537,f1948,f12026,f1368])).
fof(f15427,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14524,f12026])).
fof(f14524,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1591,f1948,f12026,f1368])).
fof(f15426,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14525,f12026])).
fof(f14525,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1607,f1948,f12026,f1368])).
fof(f15425,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14527,f12026])).
fof(f14527,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6550,f6495,f12026,f1368])).
fof(f15424,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14528,f12026])).
fof(f14528,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1538,f6495,f12026,f1368])).
fof(f15423,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14530,f12026])).
fof(f14530,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1537,f6495,f12026,f1368])).
fof(f15422,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14531,f12026])).
fof(f14531,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1591,f6495,f12026,f1368])).
fof(f15421,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14532,f12026])).
fof(f14532,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1607,f6495,f12026,f1368])).
fof(f15420,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14542,f14705])).
fof(f14705,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1372])).
fof(f14542,plain,(
  ~sP9('F_0','T_0','F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1464,f12026,f1368])).
fof(f15419,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14546,f14705])).
fof(f14546,plain,(
  ( ! [X0] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f1464,f12026,f1368])).
fof(f15418,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14561,f14705])).
fof(f14561,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1368])).
fof(f15417,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14563,f14705])).
fof(f14563,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1368])).
fof(f15416,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14568,f14705])).
fof(f14568,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1368])).
fof(f15415,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14570,f14705])).
fof(f14570,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1368])).
fof(f15414,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14572,f12026])).
fof(f14572,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1536,f12026,f1368])).
fof(f15413,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14574,f12026])).
fof(f14574,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f1536,f12026,f1368])).
fof(f15412,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14583,f12026])).
fof(f14583,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1592,f12026,f1368])).
fof(f15411,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14584,f12026])).
fof(f14584,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1608,f12026,f1368])).
fof(f15410,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14585,f12026])).
fof(f14585,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f1608,f12026,f1368])).
fof(f15409,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14586,f12026])).
fof(f14586,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1539,f12026,f1368])).
fof(f15408,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14587,f12026])).
fof(f14587,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1541,f12026,f1368])).
fof(f15407,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14589,f12026])).
fof(f14589,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f6550,f12026,f1368])).
fof(f15406,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14591,f14705])).
fof(f14591,plain,(
  ~sP9('F_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f12026,f1368])).
fof(f15405,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14594,f14705])).
fof(f14594,plain,(
  ( ! [X0] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'F_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f1464,f12026,f1368])).
fof(f15404,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14607,f14705])).
fof(f14607,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1368])).
fof(f15403,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14608,f14705])).
fof(f14608,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1368])).
fof(f15402,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14612,f14705])).
fof(f14612,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1368])).
fof(f15401,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14613,f14705])).
fof(f14613,plain,(
  ( ! [X0,X1] : (~sP9('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1368])).
fof(f15400,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14638,f12026])).
fof(f14638,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1537,f928,f12026,f1369])).
fof(f15399,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14639,f12026])).
fof(f14639,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1591,f928,f12026,f1369])).
fof(f15398,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14640,f12026])).
fof(f14640,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1607,f928,f12026,f1369])).
fof(f15397,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14641,f12026])).
fof(f14641,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1538,f928,f12026,f1369])).
fof(f15396,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14642,f12026])).
fof(f14642,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1592,f928,f12026,f1369])).
fof(f15395,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14643,f12026])).
fof(f14643,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1608,f928,f12026,f1369])).
fof(f15394,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14644,f12026])).
fof(f14644,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1539,f928,f12026,f1369])).
fof(f15393,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14654,f14705])).
fof(f14654,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1464,f12026,f1369])).
fof(f15392,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14669,f14705])).
fof(f14669,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f12026,f1369])).
fof(f15391,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14671,f12026])).
fof(f14671,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1536,f928,f12026,f1369])).
fof(f15390,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14672,f12026])).
fof(f14672,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1537,f928,f12026,f1369])).
fof(f15389,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14673,f12026])).
fof(f14673,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1591,f928,f12026,f1369])).
fof(f15388,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14674,f12026])).
fof(f14674,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1607,f928,f12026,f1369])).
fof(f15387,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14676,f12026])).
fof(f14676,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1592,f928,f12026,f1369])).
fof(f15386,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14677,f12026])).
fof(f14677,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1608,f928,f12026,f1369])).
fof(f15385,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14678,f12026])).
fof(f14678,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1539,f928,f12026,f1369])).
fof(f15384,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14688,f14705])).
fof(f14688,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1464,f12026,f1369])).
fof(f15383,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14703,f14705])).
fof(f14703,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026,f12026,f1369])).
fof(f15382,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14708,f12026])).
fof(f14708,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1464,f12026,f1373])).
fof(f15381,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14710,f12026])).
fof(f14710,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1373])).
fof(f15380,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14711,f1464])).
fof(f14711,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f12026,f12026,f1373])).
fof(f15379,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14712,f1464])).
fof(f14712,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f6565,f12026,f12026,f1373])).
fof(f15378,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14713,f12026])).
fof(f14713,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6571,f12026,f12026,f1373])).
fof(f15377,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14715,f12026])).
fof(f14715,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1373])).
fof(f15376,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14716,f1464])).
fof(f14716,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1464,f12026,f1373])).
fof(f15375,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14719,f12026])).
fof(f14719,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f1464,f12026,f1373])).
fof(f15374,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14724,f1464])).
fof(f14724,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1373])).
fof(f15373,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14725,f1464])).
fof(f14725,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1373])).
fof(f15372,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14729,f12026])).
fof(f14729,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1373])).
fof(f15371,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14730,f12026])).
fof(f14730,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1373])).
fof(f15370,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14735,f12026])).
fof(f14735,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f12026,f1373])).
fof(f15369,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14737,f12026])).
fof(f14737,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1373])).
fof(f15368,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14739,f1464])).
fof(f14739,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f12026,f12026,f1373])).
fof(f15367,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14740,f12026])).
fof(f14740,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6565,f12026,f12026,f1373])).
fof(f15366,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14741,f1464])).
fof(f14741,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f6571,f12026,f12026,f1373])).
fof(f15365,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14742,f12026])).
fof(f14742,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1373])).
fof(f15364,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14744,f1464])).
fof(f14744,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f12026,f1373])).
fof(f15363,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14747,f12026])).
fof(f14747,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f1464,f12026,f1373])).
fof(f15362,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14752,f1464])).
fof(f14752,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1373])).
fof(f15361,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14753,f1464])).
fof(f14753,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1373])).
fof(f15360,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14757,f12026])).
fof(f14757,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1373])).
fof(f15359,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14758,f12026])).
fof(f14758,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1373])).
fof(f15358,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14764,f1464])).
fof(f14764,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f1464,f12026,f1373])).
fof(f15357,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14765,f12026])).
fof(f14765,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6565,f1464,f12026,f1373])).
fof(f15356,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14767,f12026])).
fof(f14767,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f1373])).
fof(f15355,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14769,f1464])).
fof(f14769,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f12026,f12026,f1373])).
fof(f15354,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14770,f12026])).
fof(f14770,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6565,f12026,f12026,f1373])).
fof(f15353,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14771,f1464])).
fof(f14771,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f6571,f12026,f12026,f1373])).
fof(f15352,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14772,f12026])).
fof(f14772,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f12026,f12026,f1373])).
fof(f15351,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14775,f1464])).
fof(f14775,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f1464,f12026,f1373])).
fof(f15350,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14776,f12026])).
fof(f14776,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6571,f1464,f12026,f1373])).
fof(f15349,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14783,f12026])).
fof(f14783,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f12026,f1373])).
fof(f15348,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14784,f1464])).
fof(f14784,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f12026,f12026,f1373])).
fof(f15347,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14785,f1464])).
fof(f14785,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f6565,f12026,f12026,f1373])).
fof(f15346,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14786,f12026])).
fof(f14786,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6571,f12026,f12026,f1373])).
fof(f15345,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14788,f12026])).
fof(f14788,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6495,f12026,f12026,f1373])).
fof(f15344,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14794,f12026])).
fof(f14794,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f2008,f928,f12026,f1376])).
fof(f2008,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1537,f928,f1446])).
fof(f15343,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14795,f12026])).
fof(f14795,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f2009,f928,f12026,f1376])).
fof(f2009,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f1591,f928,f1446])).
fof(f15342,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14796,f12026])).
fof(f14796,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2088,f928,f12026,f1376])).
fof(f2088,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f1607,f1446])).
fof(f15341,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14797,f12026])).
fof(f14797,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1957,f928,f12026,f1376])).
fof(f1957,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f928,f1375])).
fof(f15340,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14798,f12026])).
fof(f14798,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2011,f928,f12026,f1376])).
fof(f2011,plain,(
  ( ! [X0] : (sP10('T_0','const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1592,f928,f1446])).
fof(f15339,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14799,f12026])).
fof(f14799,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2090,f928,f12026,f1376])).
fof(f2090,plain,(
  ( ! [X0,X1] : (sP10('T_0','const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1608,f1446])).
fof(f15338,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14800,f12026])).
fof(f14800,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2012,f928,f12026,f1376])).
fof(f2012,plain,(
  ( ! [X0] : (sP10('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1539,f928,f1446])).
fof(f15337,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14859,f12026])).
fof(f14859,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1951,f928,f12026,f1376])).
fof(f1951,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f928,f1374])).
fof(f15336,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14860,f12026])).
fof(f14860,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2008,f928,f12026,f1376])).
fof(f15335,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14861,f12026])).
fof(f14861,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2009,f928,f12026,f1376])).
fof(f15334,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14862,f12026])).
fof(f14862,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2088,f928,f12026,f1376])).
fof(f15333,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14864,f12026])).
fof(f14864,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f2011,f928,f12026,f1376])).
fof(f15332,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14865,f12026])).
fof(f14865,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2090,f928,f12026,f1376])).
fof(f15331,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14866,f12026])).
fof(f14866,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f2012,f928,f12026,f1376])).
fof(f15330,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14871,f1464])).
fof(f14871,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1969,f12026,f1379])).
fof(f15329,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14872,f12026])).
fof(f14872,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f1988,f12026,f1379])).
fof(f15328,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14873,f12026])).
fof(f14873,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1990,f12026,f1379])).
fof(f15327,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14874,f12026])).
fof(f14874,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1975,f12026,f1379])).
fof(f1975,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f928,f1392])).
fof(f15326,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14875,f12026])).
fof(f14875,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1992,f12026,f1379])).
fof(f1992,plain,(
  ( ! [X0] : (sP12('T_0','const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1593,f928,f1445])).
fof(f15325,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14876,f12026])).
fof(f14876,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2121,f12026,f1379])).
fof(f2121,plain,(
  ( ! [X0,X1] : (sP12('T_0','const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f928,f1609,f1445])).
fof(f15324,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14877,f12026])).
fof(f14877,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1993,f12026,f1379])).
fof(f1993,plain,(
  ( ! [X0] : (sP12('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1547,f928,f1445])).
fof(f15323,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14881,f12026])).
fof(f14881,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f2712,f12026,f1379])).
fof(f15322,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14894,f1464])).
fof(f14894,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5465,f12026,f1379])).
fof(f15321,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14895,f12026])).
fof(f14895,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5466,f12026,f1379])).
fof(f15320,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14896,f12026])).
fof(f14896,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f6595,f12026,f1379])).
fof(f6595,plain,(
  ( ! [X0] : (sP12('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)),'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f928,f1548,f6300,f1385])).
fof(f15319,plain,(
  $false | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f14897,f12026])).
fof(f14897,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f2712,f12026,f1379])).
fof(f15318,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14898,f12026])).
fof(f14898,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f5468,f12026,f1379])).
fof(f15317,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14965,f14990])).
fof(f14990,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1389])).
fof(f14965,plain,(
  ~sP11('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1966,f928,f12026,f1385])).
fof(f15316,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14968,f14990])).
fof(f14968,plain,(
  ( ! [X0] : (~sP11('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),'T_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6585,f928,f12026,f1385])).
fof(f15315,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f14986,f14990])).
fof(f14986,plain,(
  ( ! [X0] : (~sP11('T_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f928,f928,f12026,f1386])).
fof(f15314,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15001,f1464])).
fof(f15001,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1966,f928,f12026,f1390])).
fof(f15313,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15003,f12026])).
fof(f15003,plain,(
  ( ! [X0] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6585,f928,f12026,f1390])).
fof(f15312,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15009,f1464])).
fof(f15009,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f5498,f928,f12026,f1390])).
fof(f15311,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15141,f1528])).
fof(f15141,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f1447])).
fof(f15310,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15143,f1528])).
fof(f15143,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6464,f12026,f1447])).
fof(f15309,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15145,f1528])).
fof(f15145,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1934,f12026,f1447])).
fof(f15308,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15147,f1528])).
fof(f15147,plain,(
  ( ! [X0,X1] : (~sP7('const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')),X1,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f6468,f12026,f1447])).
fof(f15307,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15152,f12026])).
fof(f15152,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1572,f12026,f1478])).
fof(f1572,plain,(
  ( ! [X4,X5,X3] : (sP0(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(subsumption_resolution,[],[f1569,f928])).
fof(f1569,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | sP0(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(superposition,[],[f1201,f1060])).
fof(f15306,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15153,f12026])).
fof(f15153,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1583,f12026,f1478])).
fof(f1583,plain,(
  ( ! [X0] : (sP0(X0,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) )),
  inference(unit_resulting_resolution,[],[f932,f1201])).
fof(f15305,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15154,f12026])).
fof(f15154,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1599,f12026,f1478])).
fof(f1599,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f943,f1201])).
fof(f15304,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15155,f12026])).
fof(f15155,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1655,f12026,f1478])).
fof(f1655,plain,(
  ( ! [X0] : (sP0(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f928,f1562])).
fof(f1562,plain,(
  ( ! [X2,X3] : (sP0(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1201,f1059])).
fof(f15303,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15156,f12026])).
fof(f15156,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2221,f12026,f1478])).
fof(f2221,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X4)))) )),
  inference(subsumption_resolution,[],[f2200,f928])).
fof(f2200,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f1201,f949])).
fof(f949,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X55] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_lte|split|0')).
fof(f15302,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15157,f12026])).
fof(f15157,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2521,f12026,f1478])).
fof(f2521,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(subsumption_resolution,[],[f2496,f928])).
fof(f2496,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f1201,f958])).
fof(f15301,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15158,f12026])).
fof(f15158,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1503,f12026,f1478])).
fof(f1503,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f930,f1201])).
fof(f15300,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15159,f12026])).
fof(f15159,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2276,f12026,f1478])).
fof(f2276,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2255,f928])).
fof(f2255,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1201,f950])).
fof(f950,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X55] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|29')).
fof(f15299,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15160,f12026])).
fof(f15160,plain,(
  p(s(bool,sK31('T_0',bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1596,f12026,f1478])).
fof(f1596,plain,(
  ( ! [X0] : (sP0(X0,sK31('T_0',bool))) )),
  inference(unit_resulting_resolution,[],[f928,f1554])).
fof(f1554,plain,(
  ( ! [X2,X3] : (sP0(X3,sK31(X2,bool)) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1201,f1103])).
fof(f15298,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15161,f12026])).
fof(f15161,plain,(
  p(s(bool,sK31('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1597,f12026,f1478])).
fof(f1597,plain,(
  ( ! [X0] : (sP0(X0,sK31('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) )),
  inference(unit_resulting_resolution,[],[f930,f1554])).
fof(f15297,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15162,f12026])).
fof(f15162,plain,(
  p(s(bool,sK31('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1598,f12026,f1478])).
fof(f1598,plain,(
  ( ! [X0] : (sP0(X0,sK31('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),bool))) )),
  inference(unit_resulting_resolution,[],[f932,f1554])).
fof(f15296,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15163,f12026])).
fof(f15163,plain,(
  ( ! [X0] : (p(s(bool,sK31('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),bool)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1612,f12026,f1478])).
fof(f1612,plain,(
  ( ! [X0,X1] : (sP0(X0,sK31('const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0')),bool))) )),
  inference(unit_resulting_resolution,[],[f943,f1554])).
fof(f15295,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15164,f12026])).
fof(f15164,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1504,f12026,f1478])).
fof(f1504,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1202])).
fof(f15294,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15165,f12026])).
fof(f15165,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1571,f12026,f1478])).
fof(f1571,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(subsumption_resolution,[],[f1568,f928])).
fof(f1568,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP0('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f1202,f1060])).
fof(f15293,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15166,f12026])).
fof(f15166,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1584,f12026,f1478])).
fof(f1584,plain,(
  ( ! [X0] : (sP0('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f932,f1202])).
fof(f15292,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15167,f12026])).
fof(f15167,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1600,f12026,f1478])).
fof(f1600,plain,(
  ( ! [X0,X1] : (sP0('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f943,f1202])).
fof(f15291,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15168,f12026])).
fof(f15168,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1613,f12026,f1478])).
fof(f1613,plain,(
  ( ! [X0] : (sP0('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1561])).
fof(f1561,plain,(
  ( ! [X0,X1] : (sP0('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1202,f1059])).
fof(f15290,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15169,f12026])).
fof(f15169,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2222,f12026,f1478])).
fof(f2222,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)),X7)) )),
  inference(subsumption_resolution,[],[f2201,f928])).
fof(f2201,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)),X7)) )),
  inference(superposition,[],[f1202,f949])).
fof(f15289,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15170,f12026])).
fof(f15170,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2522,f12026,f1478])).
fof(f2522,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(subsumption_resolution,[],[f2497,f928])).
fof(f2497,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(superposition,[],[f1202,f958])).
fof(f15288,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15171,f12026])).
fof(f15171,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1505,f12026,f1478])).
fof(f1505,plain,(
  ( ! [X0] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f930,f1202])).
fof(f15287,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15172,f12026])).
fof(f15172,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2277,f12026,f1478])).
fof(f2277,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(subsumption_resolution,[],[f2256,f928])).
fof(f2256,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(superposition,[],[f1202,f950])).
fof(f15286,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15173,f12026])).
fof(f15173,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1555,f12026,f1478])).
fof(f1555,plain,(
  ( ! [X0] : (sP0(sK31('T_0',bool),X0)) )),
  inference(unit_resulting_resolution,[],[f928,f1553])).
fof(f1553,plain,(
  ( ! [X0,X1] : (sP0(sK31(X0,bool),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1202,f1103])).
fof(f15285,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15174,f12026])).
fof(f15174,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1556,f12026,f1478])).
fof(f1556,plain,(
  ( ! [X0] : (sP0(sK31('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f930,f1553])).
fof(f15284,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15175,f12026])).
fof(f15175,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1595,f12026,f1478])).
fof(f1595,plain,(
  ( ! [X0] : (sP0(sK31('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f932,f1553])).
fof(f15283,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15176,f12026])).
fof(f15176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1611,f12026,f1478])).
fof(f1611,plain,(
  ( ! [X0,X1] : (sP0(sK31('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),bool),X1)) )),
  inference(unit_resulting_resolution,[],[f943,f1553])).
fof(f15282,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15181,f12026])).
fof(f15181,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1502,f12026,f1478])).
fof(f1502,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1201])).
fof(f15281,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15182,f12026])).
fof(f15182,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1572,f12026,f1478])).
fof(f15280,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15183,f12026])).
fof(f15183,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1583,f12026,f1478])).
fof(f15279,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15184,f12026])).
fof(f15184,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1599,f12026,f1478])).
fof(f15278,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15185,f12026])).
fof(f15185,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1655,f12026,f1478])).
fof(f15277,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15186,f12026])).
fof(f15186,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2221,f12026,f1478])).
fof(f15276,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15187,f12026])).
fof(f15187,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2521,f12026,f1478])).
fof(f15275,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15188,f12026])).
fof(f15188,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1503,f12026,f1478])).
fof(f15274,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15189,f12026])).
fof(f15189,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2276,f12026,f1478])).
fof(f15273,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15190,f12026])).
fof(f15190,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1596,f12026,f1478])).
fof(f15272,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15191,f12026])).
fof(f15191,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1597,f12026,f1478])).
fof(f15271,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15192,f12026])).
fof(f15192,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1598,f12026,f1478])).
fof(f15270,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15193,f12026])).
fof(f15193,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1612,f12026,f1478])).
fof(f15269,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15195,f12026])).
fof(f15195,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1571,f12026,f1478])).
fof(f15268,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15196,f12026])).
fof(f15196,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1584,f12026,f1478])).
fof(f15267,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15197,f12026])).
fof(f15197,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1600,f12026,f1478])).
fof(f15266,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15198,f12026])).
fof(f15198,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1613,f12026,f1478])).
fof(f15265,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15199,f12026])).
fof(f15199,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2222,f12026,f1478])).
fof(f15264,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15200,f12026])).
fof(f15200,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2522,f12026,f1478])).
fof(f15263,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15201,f12026])).
fof(f15201,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl30),
  inference(unit_resulting_resolution,[],[f1505,f12026,f1478])).
fof(f15262,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15202,f12026])).
fof(f15202,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2277,f12026,f1478])).
fof(f15261,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15203,f12026])).
fof(f15203,plain,(
  p(s(bool,sK31('T_0',bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1555,f12026,f1478])).
fof(f15260,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15204,f12026])).
fof(f15204,plain,(
  p(s(bool,sK31('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1556,f12026,f1478])).
fof(f15259,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15205,f12026])).
fof(f15205,plain,(
  p(s(bool,sK31('const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')),bool))) | $spl30),
  inference(unit_resulting_resolution,[],[f1595,f12026,f1478])).
fof(f15258,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15206,f12026])).
fof(f15206,plain,(
  ( ! [X0] : (p(s(bool,sK31('const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')),bool)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1611,f12026,f1478])).
fof(f15257,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15217,f12026])).
fof(f15217,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1225])).
fof(f1225,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f853])).
fof(f15256,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f15219,f12026])).
fof(f15219,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1225])).
fof(f15222,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f945,f12026,f1008])).
fof(f1008,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f755])).
fof(f945,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f150])).
fof(f150,axiom,(
  ! [X69] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X69)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f15220,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f12026,f1225])).
fof(f15218,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f12026,f1225])).
fof(f15210,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1502,f12026,f12026,f1478])).
fof(f15208,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1502,f1464,f12026,f1478])).
fof(f15180,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1502,f12026,f12026,f1478])).
fof(f15178,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1504,f1464,f12026,f1478])).
fof(f15148,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6482,f6468,f12026,f1447])).
fof(f15146,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1534,f1934,f12026,f1447])).
fof(f15144,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6481,f6464,f12026,f1447])).
fof(f15142,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1531,f1933,f12026,f1447])).
fof(f15097,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1103,f12026,f1444])).
fof(f15095,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1245,f12026,f1444])).
fof(f15093,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1247,f12026,f1444])).
fof(f15091,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1246,f12026,f1444])).
fof(f15089,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1236,f12026,f1444])).
fof(f15087,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1235,f12026,f1444])).
fof(f15085,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1104,f12026,f1444])).
fof(f15083,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1059,f12026,f1444])).
fof(f15081,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1269,f12026,f1444])).
fof(f15079,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1479,f12026,f1444])).
fof(f15077,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1470,f12026,f1444])).
fof(f15075,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1061,f12026,f1444])).
fof(f15073,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1105,f12026,f1444])).
fof(f15060,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f950,f12026,f1444])).
fof(f15058,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1639,f12026,f1444])).
fof(f15056,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f958,f12026,f1444])).
fof(f15054,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f949,f12026,f1444])).
fof(f15052,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f1060,f12026,f1444])).
fof(f15049,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1975,f928,f928,f12026,f1393])).
fof(f15032,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5498,f928,f12026,f1391])).
fof(f15024,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1966,f928,f12026,f12026,f1390])).
fof(f15016,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5498,f928,f1464,f12026,f1390])).
fof(f15007,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1966,f12026,f928,f12026,f1390])).
fof(f15005,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1966,f1464,f928,f12026,f1390])).
fof(f14991,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5481,f12026,f1389])).
fof(f14989,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5481,f12026,f1387])).
fof(f14987,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1546,f928,f928,f12026,f1386])).
fof(f14969,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1546,f6585,f928,f12026,f1385])).
fof(f14966,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1549,f1966,f928,f12026,f1385])).
fof(f14916,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1975,f928,f12026,f1380])).
fof(f14902,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1969,f12026,f12026,f1379])).
fof(f14900,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5465,f1464,f12026,f1379])).
fof(f14886,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1969,f12026,f12026,f1379])).
fof(f14884,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1969,f1464,f12026,f1379])).
fof(f14870,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1951,f12026,f928,f12026,f1376])).
fof(f14868,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1951,f1464,f928,f12026,f1376])).
fof(f14831,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1951,f928,f12026,f12026,f1376])).
fof(f14812,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1957,f928,f1464,f12026,f1376])).
fof(f14792,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f12026,f1373])).
fof(f14790,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1464,f12026,f12026,f1373])).
fof(f14781,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f1464,f12026,f1373])).
fof(f14779,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5326,f1464,f1464,f12026,f1373])).
fof(f14762,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f12026,f1373])).
fof(f14760,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f12026,f12026,f1373])).
fof(f14751,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f1464,f12026,f1373])).
fof(f14749,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f1464,f12026,f1373])).
fof(f14734,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f12026,f12026,f12026,f1373])).
fof(f14732,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1464,f12026,f12026,f1373])).
fof(f14723,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f12026,f1464,f12026,f1373])).
fof(f14721,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1464,f1464,f12026,f1373])).
fof(f14706,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5309,f12026,f1372])).
fof(f14704,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1536,f928,f12026,f12026,f1369])).
fof(f14689,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1536,f928,f1464,f12026,f1369])).
fof(f14670,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1536,f928,f12026,f12026,f1369])).
fof(f14655,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1538,f928,f1464,f12026,f1369])).
fof(f14634,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f6550,f12026,f12026,f1368])).
fof(f14631,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1541,f12026,f12026,f1368])).
fof(f14629,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1539,f12026,f12026,f1368])).
fof(f14627,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1608,f12026,f12026,f1368])).
fof(f14625,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1592,f12026,f12026,f1368])).
fof(f14623,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1538,f12026,f12026,f1368])).
fof(f14621,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1607,f12026,f12026,f1368])).
fof(f14619,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1591,f12026,f12026,f1368])).
fof(f14617,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1948,f1537,f12026,f12026,f1368])).
fof(f14615,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1536,f12026,f12026,f1368])).
fof(f14604,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6494,f6550,f1464,f12026,f1368])).
fof(f14601,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1541,f1464,f12026,f1368])).
fof(f14596,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1947,f1536,f1464,f12026,f1368])).
fof(f14571,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6550,f6495,f12026,f12026,f1368])).
fof(f14569,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6550,f6494,f12026,f12026,f1368])).
fof(f14564,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1541,f1948,f12026,f12026,f1368])).
fof(f14562,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1541,f1947,f12026,f12026,f1368])).
fof(f14547,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6550,f6495,f1464,f12026,f1368])).
fof(f14543,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1541,f1948,f1464,f12026,f1368])).
fof(f14459,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1957,f928,f12026,f1363])).
fof(f14440,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1951,f928,f12026,f1362])).
fof(f14421,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5101,f928,f928,f12026,f1361])).
fof(f5101,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1528,f928,f928,f1348])).
fof(f14405,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f12026,f1357])).
fof(f14403,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1934,f1464,f12026,f1357])).
fof(f14395,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f12026,f1357])).
fof(f14393,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5135,f1464,f12026,f1357])).
fof(f14385,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f12026,f1356])).
fof(f14383,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f1464,f12026,f1356])).
fof(f14375,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1933,f12026,f12026,f1356])).
fof(f14373,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5135,f1464,f12026,f1356])).
fof(f14365,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5118,f12026,f1355])).
fof(f14363,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5118,f12026,f1354])).
fof(f14361,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1528,f928,f928,f12026,f1352])).
fof(f14341,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5101,f928,f12026,f1347])).
fof(f14327,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5101,f928,f12026,f1346])).
fof(f14313,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3505,f1103,f12026,f1344])).
fof(f14305,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3501,f1059,f12026,f1344])).
fof(f14296,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5808,f5716,f12026,f1344])).
fof(f14294,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5369,f5343,f12026,f1344])).
fof(f14289,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3497,f1674,f12026,f1344])).
fof(f14285,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3496,f1672,f12026,f1344])).
fof(f14282,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3494,f1639,f12026,f1344])).
fof(f3494,plain,(
  sP6('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f928,f1639,f1342])).
fof(f14278,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3491,f1060,f12026,f1344])).
fof(f3491,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f928,f1060,f1342])).
fof(f14276,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3490,f1103,f12026,f1344])).
fof(f14268,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3486,f1059,f12026,f1344])).
fof(f14259,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5805,f5716,f12026,f1344])).
fof(f14257,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f5366,f5343,f12026,f1344])).
fof(f14252,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3482,f1674,f12026,f1344])).
fof(f14248,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3481,f1672,f12026,f1344])).
fof(f14245,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3479,f1639,f12026,f1344])).
fof(f3479,plain,(
  sP6('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f928,f1639,f1342])).
fof(f14241,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3476,f1060,f12026,f1344])).
fof(f3476,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.DATATYPE_1'(s(X0,X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f928,f1060,f1342])).
fof(f15255,plain,(
  $false | $spl30),
  inference(trivial_inequality_removal,[],[f14202])).
fof(f14202,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f11162,f12026,f1344])).
fof(f15254,plain,(
  $false | $spl30),
  inference(trivial_inequality_removal,[],[f14201])).
fof(f14201,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f11056,f12026,f1344])).
fof(f11056,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1514,f1506,f1524,f928,f1326])).
fof(f14196,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1623,f12026,f1341])).
fof(f14194,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1622,f12026,f1341])).
fof(f14192,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1621,f12026,f1341])).
fof(f14190,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3517,f1623,f12026,f1341])).
fof(f14188,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3516,f1622,f12026,f1341])).
fof(f14186,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3514,f1621,f12026,f1341])).
fof(f14149,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f12026,f1337])).
fof(f14147,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f1464,f12026,f12026,f1337])).
fof(f14133,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f1464,f12026,f1337])).
fof(f14131,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f1464,f1464,f12026,f1337])).
fof(f14101,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f12026,f1337])).
fof(f14099,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f1464,f12026,f12026,f1337])).
fof(f14085,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f12026,f1464,f12026,f1337])).
fof(f14083,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f1464,f1464,f12026,f1337])).
fof(f14053,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f12026,f12026,f1337])).
fof(f14051,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1524,f1464,f12026,f12026,f1337])).
fof(f14037,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1522,f12026,f1464,f12026,f1337])).
fof(f14035,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1526,f1464,f1464,f12026,f1337])).
fof(f14005,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5053,f12026,f1336])).
fof(f14003,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5053,f12026,f1335])).
fof(f14001,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1514,f928,f928,f12026,f1333])).
fof(f13981,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5023,f12026,f1332])).
fof(f13979,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5023,f12026,f1331])).
fof(f13977,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f4043,f928,f928,f12026,f1329])).
fof(f4043,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3476,f1324])).
fof(f13515,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1524,f12026,f1328])).
fof(f13513,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1524,f12026,f1328])).
fof(f13502,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1524,f12026,f1328])).
fof(f13500,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1524,f12026,f1328])).
fof(f13494,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1524,f12026,f1328])).
fof(f13492,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1524,f12026,f1328])).
fof(f13486,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1524,f12026,f1328])).
fof(f13484,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1524,f12026,f1328])).
fof(f13481,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1524,f12026,f1328])).
fof(f13478,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1524,f12026,f1328])).
fof(f13476,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1524,f12026,f1328])).
fof(f13474,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1524,f12026,f1328])).
fof(f13469,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1524,f12026,f1328])).
fof(f13268,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1522,f12026,f1328])).
fof(f13266,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1522,f12026,f1328])).
fof(f13251,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1522,f12026,f1328])).
fof(f13249,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1522,f12026,f1328])).
fof(f13241,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1522,f12026,f1328])).
fof(f13239,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1522,f12026,f1328])).
fof(f13231,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1522,f12026,f1328])).
fof(f13229,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1522,f12026,f1328])).
fof(f13208,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1522,f12026,f1328])).
fof(f13205,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1522,f12026,f1328])).
fof(f13203,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1522,f12026,f1328])).
fof(f13201,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1522,f12026,f1328])).
fof(f13199,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1522,f12026,f1328])).
fof(f12634,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1524,f12026,f1327])).
fof(f12632,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1524,f12026,f1327])).
fof(f12621,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1524,f12026,f1327])).
fof(f12619,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1524,f12026,f1327])).
fof(f12613,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1524,f12026,f1327])).
fof(f12611,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1524,f12026,f1327])).
fof(f12605,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1524,f12026,f1327])).
fof(f12603,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1524,f12026,f1327])).
fof(f12600,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1524,f12026,f1327])).
fof(f12597,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1524,f12026,f1327])).
fof(f12595,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1524,f12026,f1327])).
fof(f12593,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1524,f12026,f1327])).
fof(f12588,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1524,f12026,f1327])).
fof(f12387,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1514,f1510,f1522,f12026,f1327])).
fof(f12385,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1510,f1522,f12026,f1327])).
fof(f12370,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f1514,f6359,f1522,f12026,f1327])).
fof(f12368,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f6359,f1522,f12026,f1327])).
fof(f12360,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1508,f1522,f12026,f1327])).
fof(f12358,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1508,f1522,f12026,f1327])).
fof(f12350,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1514,f1509,f1522,f12026,f1327])).
fof(f12348,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1509,f1522,f12026,f1327])).
fof(f12327,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3521,f1518,f1506,f1522,f12026,f1327])).
fof(f12324,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f6345,f6370,f1506,f1522,f12026,f1327])).
fof(f12322,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1516,f1506,f1522,f12026,f1327])).
fof(f12320,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3520,f1517,f1506,f1522,f12026,f1327])).
fof(f12318,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f3518,f1514,f1506,f1522,f12026,f1327])).
fof(f12179,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f3491,f928,f928,f12026,f1325])).
fof(f12141,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f12026,f1307])).
fof(f12139,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f1464,f12026,f1307])).
fof(f12132,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f12026,f1307])).
fof(f12130,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1930,f1464,f12026,f1307])).
fof(f12123,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f12026,f1306])).
fof(f12121,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1930,f1464,f12026,f1306])).
fof(f12114,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f12026,f12026,f1306])).
fof(f12112,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1551,f1464,f12026,f1306])).
fof(f12105,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1701,f928,f12026,f1220])).
fof(f1701,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f928,f928,f1216])).
fof(f12097,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1701,f928,f12026,f1217])).
fof(f12089,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1704,f12026,f12026,f1215])).
fof(f12087,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1716,f1464,f12026,f1215])).
fof(f12075,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1704,f12026,f12026,f1215])).
fof(f12073,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1704,f1464,f12026,f1215])).
fof(f12061,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1623,f12026,f1186])).
fof(f12059,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1464,f1622,f12026,f1186])).
fof(f12057,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f12026,f1621,f12026,f1186])).
fof(f12055,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1623,f12026,f1186])).
fof(f12053,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1622,f12026,f1186])).
fof(f12051,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f12026,f1621,f12026,f1186])).
fof(f12047,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1623,f12026,f1186])).
fof(f12045,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1464,f1622,f12026,f1186])).
fof(f12043,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f12026,f1621,f12026,f1186])).
fof(f12041,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1623,f12026,f1186])).
fof(f12039,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f12026,f1622,f12026,f1186])).
fof(f12037,plain,(
  $false | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f12026,f1621,f12026,f1186])).
fof(f12033,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1622,f12026,f1000])).
fof(f12031,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1060,f12026,f995])).
fof(f12029,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f928,f12026])).
fof(f12028,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl32),
  inference(cnf_transformation,[],[f12028_D])).
fof(f12028_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f12026,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl30),
  inference(cnf_transformation,[],[f12026_D])).
fof(f12026_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12006,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f1474,f1470])).
fof(f12025,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl28),
  inference(cnf_transformation,[],[f12025_D])).
fof(f12025_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f12005,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f1474,f1061])).
fof(f5917,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl22),
  inference(cnf_transformation,[],[f5917_D])).
fof(f5917_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f8169,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl27),
  inference(cnf_transformation,[],[f8169_D])).
fof(f8169_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f8094,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1003,f8073])).
fof(f5919,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl25),
  inference(cnf_transformation,[],[f5919_D])).
fof(f5919_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f5900,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1003,f934])).
fof(f5914,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.list_size_2'(s(fun(X14,'type.num.num'),X15),s('type.list.list'(X14),'const.list.NIL_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f5914_D])).
fof(f5914_D,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.list_size_2'(s(fun(X14,'type.num.num'),X15),s('type.list.list'(X14),'const.list.NIL_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5905,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl7),
  inference(cnf_transformation,[],[f5905_D])).
fof(f5905_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f5892,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.list_size_2'(s(fun(X14,'type.num.num'),X15),s('type.list.list'(X14),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1003,f1063])).
fof(f1063,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ! [X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f214])).
fof(f214,axiom,(
  ! [X6] : ! [X27] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.list_size_2'(s(fun(X6,'type.num.num'),X27),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.list.list_size_def|split|0')).
fof(f5913,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X12,'type.num.num'),X13),s('type.PairingHeap.heap'(X12),'const.PairingHeap.Empty_0')))) ) | $spl18),
  inference(cnf_transformation,[],[f5913_D])).
fof(f5913_D,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X12,'type.num.num'),X13),s('type.PairingHeap.heap'(X12),'const.PairingHeap.Empty_0')))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5891,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X12,'type.num.num'),X13),s('type.PairingHeap.heap'(X12),'const.PairingHeap.Empty_0')))) )),
  inference(superposition,[],[f1003,f1062])).
fof(f1062,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X0,'type.num.num'),X1),s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))) )),
  inference(cnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X0,'type.num.num'),X1),s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0] : ! [X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X0,'type.num.num'),X1),s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0')))),
  inference(rectify,[],[f188])).
fof(f188,axiom,(
  ! [X6] : ! [X27] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.PairingHeap.heap_size_2'(s(fun(X6,'type.num.num'),X27),s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.heap_size_def|split|0')).
fof(f5912,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) ) | $spl16),
  inference(cnf_transformation,[],[f5912_D])).
fof(f5912_D,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5890,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1003,f957])).
fof(f957,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X55] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X55)))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|6')).
fof(f5909,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f5909_D])).
fof(f5909_D,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5911,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl15),
  inference(cnf_transformation,[],[f5911_D])).
fof(f5911_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f5888,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1003,f1468])).
fof(f1468,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f975,f934])).
fof(f975,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f141])).
fof(f141,axiom,(
  ! [X55] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|11')).
fof(f5908,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(cnf_transformation,[],[f5908_D])).
fof(f5908_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5887,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1003,f931])).
fof(f931,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f111])).
fof(f111,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.prim_rec.PRE|split|0')).
fof(f5906,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(cnf_transformation,[],[f5906_D])).
fof(f5906_D,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5884,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1003,f946])).
fof(f946,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X55] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.numeral.numeral_distrib|split|4')).
fof(f5903,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) | $spl4),
  inference(cnf_transformation,[],[f5903_D])).
fof(f5903_D,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5883,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) )),
  inference(superposition,[],[f1003,f945])).
fof(f1460,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | $spl0),
  inference(cnf_transformation,[],[f1460_D])).
fof(f1460_D,plain,(
  p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1462,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK13),sK14) | $spl3),
  inference(cnf_transformation,[],[f1462_D])).
fof(f1462_D,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK13),sK14) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f926,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK13),sK14) | p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))))),
  inference(cnf_transformation,[],[f746])).
fof(f746,plain,(
  (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK13),sK14)) & (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f745])).
fof(f745,plain,(
  ? [X0,X1] : ((p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1)) & (~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1)))),
  inference(nnf_transformation,[],[f654])).
fof(f654,plain,(
  ? [X0,X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) <~> s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) <=> s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X0),X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X0),X1)))) <=> ~s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X6),X7)))) <=> ~s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(X6),X7)))) <=> ~s('type.PairingHeap.heap'(X6),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X6),X7))),
  file('thm.PairingHeap.PairingHeap_96_prove_ltb','thm.PairingHeap.PairingHeap_96_prove_ltb')).
fof(f1461,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f1461_D])).
fof(f1461_D,plain,(
  ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1463,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14) | $spl2),
  inference(cnf_transformation,[],[f1463_D])).
fof(f1463_D,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f927,plain,(
  s('type.PairingHeap.heap'(sK13),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK13),sK14) | ~p(s(bool,'const.PairingHeap.find_min_side_1'(s('type.PairingHeap.heap'(sK13),sK14))))),
  inference(cnf_transformation,[],[f746])).
% SZS output end Proof for thm.PairingHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 904
% Passive clauses: 5644
% Generated clauses: 53819
% Final active clauses: 523
% Final passive clauses: 2024
% Input formulas: 261
% Initial clauses: 507
% 
% Duplicate literals: 9
% Trivial inequalities: 41
% Fw subsumption resolutions: 1309
% Fw demodulations: 8813
% Bw demodulations: 769
% Condensations: 28
% 
% Simple tautologies: 114
% Equational tautologies: 34
% Forward subsumptions: 35910
% Fw demodulations to eq. taut.: 308
% Bw demodulations to eq. taut.: 29
% 
% Unit resulting resolution: 12678
% Factoring: 1
% Forward superposition: 14447
% Backward superposition: 14303
% Self superposition: 147
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 7272
% SAT solver unit clauses: 3144
% SAT solver binary clauses: 3087
% 
% Sat splits: 17
% Sat splitting refutations: 1592
% 
% Memory used [KB]: 41577
% Time elapsed: 1.614 s
% ------------------------------
% ------------------------------
% Success in time 2.96 s
