Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (47443)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 3136
% Generated clauses: 12785
% Final active clauses: 94
% Final passive clauses: 2923
% Input formulas: 36
% Initial clauses: 52
% 
% Duplicate literals: 492
% Trivial inequalities: 216
% Fw subsumption resolutions: 1592
% Fw demodulations: 810
% 
% Simple tautologies: 275
% Equational tautologies: 65
% Forward subsumptions: 5533
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 213
% Factoring: 2
% Forward superposition: 4056
% Backward superposition: 4826
% Self superposition: 268
% Equality factoring: 164
% Equality resolution: 24
% 
% Split clauses: 255
% Split components: 521
% SAT solver clauses: 650
% SAT solver unit clauses: 88
% SAT solver binary clauses: 168
% 
% Sat splits: 255
% 
% TWLsolver clauses: 648
% TWLsolver calls for satisfiability: 122
% 
% Memory used [KB]: 12792
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (47464)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1039
% Generated clauses: 9646
% Final active clauses: 70
% Final passive clauses: 888
% Input formulas: 36
% Initial clauses: 52
% 
% Duplicate literals: 252
% Trivial inequalities: 52
% Fw subsumption resolutions: 330
% Fw demodulations: 943
% Bw demodulations: 55
% Condensations: 552
% 
% Simple tautologies: 103
% Equational tautologies: 34
% Forward subsumptions: 4546
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 65
% Factoring: 1
% Forward superposition: 3362
% Backward superposition: 2553
% Self superposition: 147
% Equality factoring: 58
% Equality resolution: 10
% 
% Split clauses: 45
% Split components: 94
% SAT solver clauses: 1743
% SAT solver unit clauses: 69
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 45
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 1742
% TWLsolver calls for satisfiability: 378
% 
% Memory used [KB]: 24050
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% (47490)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 703
% Passive clauses: 15381
% Generated clauses: 141730
% Final active clauses: 84
% Final passive clauses: 825
% Input formulas: 36
% Initial clauses: 52
% Discarded non-redundant clauses: 6159
% 
% Duplicate literals: 1517
% Trivial inequalities: 728
% Fw subsumption resolutions: 5914
% Fw demodulations: 31190
% Bw demodulations: 8539
% Condensations: 2168
% Global subsumptions: 281
% 
% Simple tautologies: 600
% Equational tautologies: 1165
% Forward subsumptions: 81952
% Fw demodulations to eq. taut.: 232
% Bw demodulations to eq. taut.: 599
% 
% Unit resulting resolution: 10323
% Factoring: 17
% Forward superposition: 47824
% Backward superposition: 23758
% Self superposition: 311
% Equality factoring: 76
% Equality resolution: 7
% 
% Split clauses: 96
% Split components: 201
% SAT solver clauses: 16275
% SAT solver unit clauses: 400
% SAT solver binary clauses: 5651
% 
% Sat splits: 96
% Sat splitting refutations: 418
% 
% Memory used [KB]: 142129
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f33262,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f75,f94,f74,f95,f93,f73,f72,f99,f71,f97,f70,f103,f101,f92,f280,f504,f507,f508,f511,f513,f521,f520,f563,f562,f519,f518,f510,f509,f517,f516,f565,f515,f514,f506,f505,f614,f666,f667,f837,f1046,f1047,f30748,f30781,f30814,f30816,f30818,f30815,f30819,f30817,f29964,f29965,f29966,f29967,f30820,f29968,f30822,f30825,f30827,f30829,f30826,f30830,f30828,f29973,f29974,f29975,f29976,f30833,f29977,f30835,f30836,f30838,f30840,f30837,f30841,f30839,f29982,f30842,f29983,f30843,f30101,f30845,f30844,f30102,f30847,f30103,f30848,f30846,f30110,f30849,f30111,f30850,f30851,f30118,f30852,f30119,f30853,f30854,f30120,f30855,f30121,f30857,f30122,f30858,f30123,f30859,f30124,f30860,f30125,f30861,f30126,f30856,f30862,f30127,f30863,f30128,f30823,f30129,f30864,f30130,f30865,f30131,f30832,f30867,f30869,f30868,f30780,f30333,f30886,f30334,f30887,f30347,f30901,f30348,f30902,f30353,f30908,f30354,f30909,f30355,f30911,f30356,f30912,f30357,f30913,f30358,f30914,f30359,f30915,f30360,f30910,f30917,f30486,f30927,f30926,f30487,f30929,f30488,f30930,f30928,f30495,f30931,f30496,f30932,f30933,f30503,f30934,f30504,f30935,f30936,f30505,f30937,f30506,f30939,f30507,f30940,f30508,f30941,f30509,f30942,f30510,f30943,f30511,f30938,f30944,f30512,f30824,f30945,f30513,f30888,f30821,f30514,f30946,f30515,f30831,f30947,f30516,f30903,f30834,f31436,f31438,f31441,f31442,f31443,f31444,f31402,f31445,f31446,f31447,f31422,f31449,f564,f32411,f32412,f31437,f32911,f32912,f32931,f33098,f32932,f33101,f33122,f33123,f33124,f33100,f33125,f33256])).
fof(f33256,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f33255,f32504])).
fof(f32504,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | ($spl0 | $spl4)),
  inference(equality_resolution,[],[f32482])).
fof(f32482,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X0,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(duplicate_literal_removal,[],[f32481])).
fof(f32481,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X0,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X0,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f32480,f81])).
fof(f81,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK12(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK11(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK11(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK12(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12])],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X7] : ! [X46] : (? [X13] : s('type.semanticPrimitives.result'(X6,X7),X46) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X13))) | ? [X14] : s('type.semanticPrimitives.result'(X6,X7),X46) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X14))))),
  file('thm.closProps.eqs|split|5','thm.semanticPrimitives.result_nchotomy')).
fof(f32480,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK11(X0,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X0,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f32471])).
fof(f32471,plain,(
  ( ! [X0] : (s(sK2,sK4) != s(sK2,sK4) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK11(X0,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X0,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f95,f32439])).
fof(f32439,plain,(
  ( ! [X1] : (s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK11(X1,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X1) | s('type.semanticPrimitives.result'(sK0,sK1),X1) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X1,sK1,sK0))))) ) | $spl0),
  inference(superposition,[],[f27644,f85])).
fof(f85,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X8,X6] : ! [X13,X18,X19] : s(X8,happ(s(fun(X6,X8),X18),s(X6,X13))) = s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X13))),s(fun(X6,X8),X18),s(fun('type.semanticPrimitives.error_result'(X7),X8),X19)))),
  file('thm.closProps.eqs|split|5','thm.semanticPrimitives.result_case_def|split|0')).
fof(f27644,plain,(
  ( ! [X14] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK11(X14,sK1,sK0)))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X14) | s('type.semanticPrimitives.result'(sK0,sK1),X14) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(X14,sK1,sK0))))) ) | $spl0),
  inference(superposition,[],[f27612,f81])).
fof(f27612,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl0),
  inference(equality_resolution,[],[f7366])).
fof(f7366,plain,(
  ( ! [X0,X1] : (s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X0,sK3,sK6,X1,sK2,sK1,sK0)))) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X0,sK3,sK6,X1,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0)) ) | $spl0),
  inference(equality_resolution,[],[f651])).
fof(f651,plain,(
  ( ! [X19,X20,X18] : (s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK9(sK5,sK6,X18,sK3,X19,X20,sK2,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),X19),s(sK0,sK9(sK5,sK6,X18,sK3,X19,X20,sK2,sK1,sK0)))) | s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X18,sK3,X19,X20,sK2,sK1,sK0)))) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X20),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X18,sK3,X19,X20,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X18),s(fun(sK0,sK2),X19),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X20))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X18)) ) | $spl0),
  inference(superposition,[],[f92,f79])).
fof(f79,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))) | s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,happ(s(fun(X0,X2),X4),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6) | (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6) | ? [X10] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X10))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) | ? [X9] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X9))) & s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),X9))) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),X9)))) | s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6) | ? [X10] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X10))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) | ? [X9] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X9))) & s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),X9))) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),X9))))) | s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semanticPrimitives.result'(X0,X1),X5) = s('type.semanticPrimitives.result'(X0,X1),X6) & ! [X10] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X9))) => s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),X9))) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),X9))))) => s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8] : ((s('type.semanticPrimitives.result'(X0,X1),X5) = s('type.semanticPrimitives.result'(X0,X1),X6) & ! [X10] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X9))) => s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),X9))) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),X9))))) => s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7,X8] : ! [X20,X12,X21,X22,X18,X19] : ((s('type.semanticPrimitives.result'(X6,X7),X21) = s('type.semanticPrimitives.result'(X6,X7),X22) & ! [X13] : (s('type.semanticPrimitives.result'(X6,X7),X22) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X13))) => s(X8,happ(s(fun(X6,X8),X12),s(X6,X13))) = s(X8,happ(s(fun(X6,X8),X18),s(X6,X13)))) & ! [X13] : (s('type.semanticPrimitives.result'(X6,X7),X22) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X13))) => s(X8,happ(s(fun('type.semanticPrimitives.error_result'(X7),X8),X19),s('type.semanticPrimitives.error_result'(X7),X13))) = s(X8,happ(s(fun('type.semanticPrimitives.error_result'(X7),X8),X20),s('type.semanticPrimitives.error_result'(X7),X13))))) => s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X21),s(fun(X6,X8),X18),s(fun('type.semanticPrimitives.error_result'(X7),X8),X19))) = s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X22),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.error_result'(X7),X8),X20))))),
  file('thm.closProps.eqs|split|5','thm.semanticPrimitives.result_case_cong')).
fof(f33255,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | ($spl0 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f33246])).
fof(f33246,plain,(
  s(sK2,sK4) != s(sK2,sK4) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f94,f32910])).
fof(f32910,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | ($spl0 | $spl4)),
  inference(superposition,[],[f32512,f92])).
fof(f32512,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),X2),X3),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.error_result'(sK1),X2),X3)))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f80,f32504])).
fof(f80,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X8,X7] : ! [X13,X18,X19] : s(X8,happ(s(fun('type.semanticPrimitives.error_result'(X7),X8),X19),s('type.semanticPrimitives.error_result'(X7),X13))) = s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X13))),s(fun(X6,X8),X18),s(fun('type.semanticPrimitives.error_result'(X7),X8),X19)))),
  file('thm.closProps.eqs|split|5','thm.semanticPrimitives.result_case_def|split|1')).
fof(f33125,plain,(
  ( ! [X159,X160] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X159),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X160)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl208),
  inference(cnf_transformation,[],[f33125_D])).
fof(f33125_D,plain,(
  ( ! [X159,X160] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X159),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X160)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f33100,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | $spl202),
  inference(cnf_transformation,[],[f33100_D])).
fof(f33100_D,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f33124,plain,(
  ( ! [X159,X160] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X159),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X160)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f32957,f32512])).
fof(f32957,plain,(
  ( ! [X159,X160] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X160),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X159),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X160)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f20358,f32512])).
fof(f20358,plain,(
  ( ! [X0] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),X0),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X0,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f6737])).
fof(f6737,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1,sK3,X0,sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),X1),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl0),
  inference(equality_resolution,[],[f650])).
fof(f650,plain,(
  ( ! [X17,X15,X16] : (s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK10(X17,X16,sK3,X15,sK6,sK5,sK2,sK1,sK0)))) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X17),s('type.semanticPrimitives.error_result'(sK1),sK10(X17,X16,sK3,X15,sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X15),s(fun(sK0,sK2),X16),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X17))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X17,X16,sK3,X15,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X15)) ) | $spl0),
  inference(superposition,[],[f92,f78])).
fof(f78,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))) | s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X3),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X8),s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f68])).
fof(f33123,plain,(
  ( ! [X158,X157] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X157),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X158)),sK5,sK2,sK1,sK0))))) ) | $spl206),
  inference(cnf_transformation,[],[f33123_D])).
fof(f33123_D,plain,(
  ( ! [X158,X157] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X157),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X158)),sK5,sK2,sK1,sK0))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f33122,plain,(
  ( ! [X158,X157] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X157),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X158)),sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f32956,f32512])).
fof(f32956,plain,(
  ( ! [X158,X157] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X158),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,fun(sK0,sK2)),X157),s(fun('type.semanticPrimitives.error_result'(sK1),fun(sK0,sK2)),X158)),sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f19170,f32512])).
fof(f19170,plain,(
  ( ! [X0] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),X0),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,X0,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f6537])).
fof(f6537,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,X0,sK3,X1,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),X1),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl0),
  inference(equality_resolution,[],[f649])).
fof(f649,plain,(
  ( ! [X14,X12,X13] : (s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X12,sK3,X13,X14,sK2,sK1,sK0)))) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X14),s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X12,sK3,X13,X14,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X12),s(fun(sK0,sK2),X13),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X14))) | s('type.semanticPrimitives.result'(sK0,sK1),X12) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,X12,sK3,X13,X14,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X12)) ) | $spl0),
  inference(superposition,[],[f92,f78])).
fof(f33101,plain,(
  ( ! [X69] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,X69,sK5,sK2,sK1,sK0))))) ) | $spl204),
  inference(cnf_transformation,[],[f33101_D])).
fof(f33101_D,plain,(
  ( ! [X69] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,X69,sK5,sK2,sK1,sK0))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f32932,plain,(
  ( ! [X74] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,X74,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f19170,f32512])).
fof(f33098,plain,(
  ( ! [X68] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X68,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl200),
  inference(cnf_transformation,[],[f33098_D])).
fof(f33098_D,plain,(
  ( ! [X68] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X68,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f32931,plain,(
  ( ! [X73] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X73,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f20358,f32512])).
fof(f32912,plain,(
  ( ! [X69] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,X69,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f32512,f19170])).
fof(f32911,plain,(
  ( ! [X68] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK12(sK3,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X68,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f32512,f20358])).
fof(f31437,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK7)))) ) | $spl190),
  inference(cnf_transformation,[],[f31437_D])).
fof(f31437_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK7)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f32412,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X23),bool),X24),s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X21),bool),X22),s('type.semanticPrimitives.result'(sK0,X21),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))))))))))) ) | $spl198),
  inference(cnf_transformation,[],[f32412_D])).
fof(f32412_D,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X23),bool),X24),s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X21),bool),X22),s('type.semanticPrimitives.result'(sK0,X21),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))))))))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f32411,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X23),bool),X24),s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X21),bool),X22),s('type.semanticPrimitives.result'(sK0,X21),'const.semanticPrimitives.Rval_1'(s(sK0,sK7)))))))))) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X20),s(sK0,sK7)))) ) | $spl16),
  inference(subsumption_resolution,[],[f32359,f566])).
fof(f566,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7)))) ) | $spl16),
  inference(resolution,[],[f506,f90])).
fof(f90,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X23,X24] : ((p(s(bool,X23)) => p(s(bool,X24))) => ((p(s(bool,X24)) => p(s(bool,X23))) => s(bool,X23) = s(bool,X24)))),
  file('thm.closProps.eqs|split|5','thm.bool.IMP_ANTISYM_AX')).
fof(f32359,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X23),bool),X24),s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X21),bool),X22),s('type.semanticPrimitives.result'(sK0,X21),'const.semanticPrimitives.Rval_1'(s(sK0,sK7)))))))))) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X20),s(sK0,sK7)))) ) | $spl16),
  inference(superposition,[],[f31894,f31145])).
fof(f31145,plain,(
  ( ! [X4,X5,X3] : (s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X3),bool),X4),s('type.semanticPrimitives.result'(sK0,X3),'const.semanticPrimitives.Rval_1'(s(sK0,sK7)))))) ) | $spl16),
  inference(resolution,[],[f31052,f566])).
fof(f31052,plain,(
  ( ! [X17,X18] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,X17),bool),X18),s('type.semanticPrimitives.result'(sK0,X17),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))))))) ) | $spl16),
  inference(superposition,[],[f587,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),happ(s(fun(X1,'type.semanticPrimitives.result'(X1,X0)),'const.semanticPrimitives.Rval_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),happ(s(fun(X1,'type.semanticPrimitives.result'(X1,X0)),'const.semanticPrimitives.Rval_0'),s(X1,X2)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2] : s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),happ(s(fun(X1,'type.semanticPrimitives.result'(X1,X0)),'const.semanticPrimitives.Rval_0'),s(X1,X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X6] : ! [X15] : s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X15))) = s('type.semanticPrimitives.result'(X6,X7),happ(s(fun(X6,'type.semanticPrimitives.result'(X6,X7)),'const.semanticPrimitives.Rval_0'),s(X6,X15)))),
  file('thm.closProps.eqs|split|5','ARITY_0_1_const.semanticPrimitives.Rval')).
fof(f587,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,happ(s(fun(sK0,X8),X10),s(sK0,sK7))))))) ) | $spl16),
  inference(superposition,[],[f506,f84])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X28,X29,X30] : ! [X31,X32,X33] : s(X28,happ(s(fun(X29,X28),X31),s(X29,happ(s(fun(X30,X29),X32),s(X30,X33))))) = s(X28,happ(s(fun(X30,X28),'_dst_x0x1_2'(s(fun(X29,X28),X31),s(fun(X30,X29),X32))),s(X30,X33)))),
  file('thm.closProps.eqs|split|5','_dst_x0x1_ax')).
fof(f31894,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X10),bool),X11),s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,X12)))))) | p(s(bool,X13)) | s(bool,X12) = s(bool,X13)) ) | $spl16),
  inference(resolution,[],[f31890,f87])).
fof(f87,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f63])).
fof(f31890,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X24),bool),X25),s('type.semanticPrimitives.result'(bool,X24),'const.semanticPrimitives.Rval_1'(s(bool,X23))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f31880,f506])).
fof(f31880,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X24),bool),X25),s('type.semanticPrimitives.result'(bool,X24),'const.semanticPrimitives.Rval_1'(s(bool,X23)))))) | ~p(s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK7)))) | ~p(s(bool,X23))) ) | $spl16),
  inference(superposition,[],[f31792,f90])).
fof(f31792,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X36,X37),bool),X39),s('type.semanticPrimitives.result'(X36,X37),'const.semanticPrimitives.Rval_1'(s(X36,happ(s(fun(sK0,X36),X38),s(sK0,sK7))))))))) ) | $spl16),
  inference(superposition,[],[f31047,f86])).
fof(f31047,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X4),s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(sK0,X1),X3),s(sK0,sK7))))))))) ) | $spl16),
  inference(superposition,[],[f587,f84])).
fof(f564,plain,(
  ( ! [X66] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | $spl26),
  inference(cnf_transformation,[],[f564_D])).
fof(f564_D,plain,(
  ( ! [X66] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f31449,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X40,sK0),X41),s(fun(bool,X40),X42))),s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK7))))))) | p(s(bool,X43)) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X40,sK0),X41),s(X40,happ(s(fun(bool,X40),X42),s(bool,X43)))))))) ) | $spl196),
  inference(cnf_transformation,[],[f31449_D])).
fof(f31449_D,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X40,sK0),X41),s(fun(bool,X40),X42))),s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK7))))))) | p(s(bool,X43)) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X40,sK0),X41),s(X40,happ(s(fun(bool,X40),X42),s(bool,X43)))))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f31422,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X40,sK0),X41),s(fun(bool,X40),X42))),s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK7))))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X40,sK0),X41),s(X40,happ(s(fun(bool,X40),X42),s(bool,X43))))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X38),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,X43))) ) | ($spl4 | $spl16 | $spl18)),
  inference(superposition,[],[f1087,f30969])).
fof(f30969,plain,(
  ( ! [X6,X7] : (s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X7),s('type.semanticPrimitives.error_result'(sK1),sK8)))) ) | ($spl16 | $spl18)),
  inference(resolution,[],[f509,f506])).
fof(f1087,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X145,sK0),X146),s(fun(bool,X145),X147))),s(bool,X144))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X145,sK0),X146),s(X145,happ(s(fun(bool,X145),X147),s(bool,X143))))))) | p(s(bool,X144)) | p(s(bool,X143))) ) | $spl4),
  inference(superposition,[],[f1020,f87])).
fof(f1020,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X85,sK0),'_dst_x0x1_2'(s(fun(X83,sK0),X84),s(fun(X85,X83),X86))),s(X85,X87))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X83,sK0),X84),s(X83,happ(s(fun(X85,X83),X86),s(X85,X87)))))))) ) | $spl4),
  inference(superposition,[],[f95,f84])).
fof(f31447,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X11),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31413,f566])).
fof(f31413,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X11),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,X13))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f87,f30969])).
fof(f31446,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X8),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31412,f566])).
fof(f31412,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK7))) | p(s(bool,X10)) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X8),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f87,f30969])).
fof(f31445,plain,(
  ( ! [X64,X62,X65] : (p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK7)))) ) | $spl194),
  inference(cnf_transformation,[],[f31445_D])).
fof(f31445_D,plain,(
  ( ! [X64,X62,X65] : (p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK7)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f31402,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK7)))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f31063,f30969])).
fof(f31063,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2)) ) | $spl16),
  inference(resolution,[],[f31061,f87])).
fof(f31061,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12))))) ) | $spl16),
  inference(subsumption_resolution,[],[f31050,f506])).
fof(f31050,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | ~p(s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK7)))) | ~p(s(bool,X12))) ) | $spl16),
  inference(superposition,[],[f587,f90])).
fof(f31444,plain,(
  ( ! [X59,X60] : (p(s(bool,happ(s(fun(bool,bool),X60),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X59),s('type.semanticPrimitives.error_result'(sK1),sK8))))))) ) | $spl192),
  inference(cnf_transformation,[],[f31444_D])).
fof(f31444_D,plain,(
  ( ! [X59,X60] : (p(s(bool,happ(s(fun(bool,bool),X60),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X59),s('type.semanticPrimitives.error_result'(sK1),sK8))))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f31443,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(bool,bool),X60),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X59),s('type.semanticPrimitives.error_result'(sK1),sK8)))))) | s(bool,X61) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK7)))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31401,f566])).
fof(f31401,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(bool,bool),X60),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X59),s('type.semanticPrimitives.error_result'(sK1),sK8)))))) | p(s(bool,X61)) | s(bool,X61) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK7)))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f31063,f30969])).
fof(f31442,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X29),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK7)))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31392,f566])).
fof(f31392,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X29),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK7))) | p(s(bool,X30))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f87,f30969])).
fof(f31441,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X26),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X27) = s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK7)))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31391,f566])).
fof(f31391,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X26),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X27) = s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK7))) | p(s(bool,X27))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f87,f30969])).
fof(f31438,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X6),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31380,f566])).
fof(f31380,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X6),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,X7))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f30969,f87])).
fof(f31436,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X3),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f31379,f566])).
fof(f31379,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK7))) | p(s(bool,X4)) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X3),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f30969,f87])).
fof(f30834,plain,(
  ( ! [X557,X555,X556] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X555),sK0),X556),s('type.semanticPrimitives.error_result'(X555),X557))))) | s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557))) = s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557)),X555,bool))))) ) | $spl58),
  inference(cnf_transformation,[],[f30834_D])).
fof(f30834_D,plain,(
  ( ! [X557,X555,X556] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X555),sK0),X556),s('type.semanticPrimitives.error_result'(X555),X557))))) | s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557))) = s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557)),X555,bool))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f30903,plain,(
  ( ! [X405,X404] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X404),s(bool,X405)))))) ) | $spl126),
  inference(cnf_transformation,[],[f30903_D])).
fof(f30903_D,plain,(
  ( ! [X405,X404] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X404),s(bool,X405)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f30516,plain,(
  ( ! [X1734,X1736,X1738,X1735,X1737] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X1737),s(bool,X1738))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1734),sK0),X1735),s('type.semanticPrimitives.error_result'(X1734),X1736))))) | s('type.semanticPrimitives.result'(bool,X1734),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1734),X1736))) = s('type.semanticPrimitives.result'(bool,X1734),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1734),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1734),X1736)),X1734,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f95,f21143])).
fof(f21143,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (s(X120,happ(s(fun(bool,X120),X123),s(bool,X125))) = s(X120,happ(s(fun('type.semanticPrimitives.error_result'(X121),X120),X124),s('type.semanticPrimitives.error_result'(X121),X122))) | s('type.semanticPrimitives.result'(bool,X121),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X121),X122))) = s('type.semanticPrimitives.result'(bool,X121),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X121),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X121),X122)),X121,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1234,f80])).
fof(f1234,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(X24,happ(s(fun(bool,X24),X25),s(bool,X22))) = s(X24,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X21),X23),s(fun(bool,X24),X25),s(fun('type.semanticPrimitives.error_result'(X21),X24),X26))) | s('type.semanticPrimitives.result'(bool,X21),X23) = s('type.semanticPrimitives.result'(bool,X21),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X21),sK12(X23,X21,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f85,f964])).
fof(f964,plain,(
  ( ! [X273,X275,X274] : (s('type.semanticPrimitives.result'(bool,X274),X273) = s('type.semanticPrimitives.result'(bool,X274),'const.semanticPrimitives.Rval_1'(s(bool,X275))) | s('type.semanticPrimitives.result'(bool,X274),X273) = s('type.semanticPrimitives.result'(bool,X274),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X274),sK12(X273,X274,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f81,f771])).
fof(f771,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2)) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f736,f736])).
fof(f736,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK7)))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f566,f565])).
fof(f30947,plain,(
  ( ! [X1728,X1730,X1729,X1733] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1733,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1728),sK0),X1729),s('type.semanticPrimitives.error_result'(X1728),X1730)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730))) = s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730)),X1728,bool))))) ) | $spl188),
  inference(cnf_transformation,[],[f30947_D])).
fof(f30947_D,plain,(
  ( ! [X1728,X1730,X1729,X1733] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1733,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1728),sK0),X1729),s('type.semanticPrimitives.error_result'(X1728),X1730)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730))) = s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730)),X1728,bool))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f30831,plain,(
  ( ! [X545,X550] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X545),s(bool,X550)))))) ) | $spl52),
  inference(cnf_transformation,[],[f30831_D])).
fof(f30831_D,plain,(
  ( ! [X545,X550] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X545),s(bool,X550)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f30515,plain,(
  ( ! [X1728,X1730,X1732,X1729,X1731,X1733] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1731),s(bool,X1732))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1733,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1728),sK0),X1729),s('type.semanticPrimitives.error_result'(X1728),X1730)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730))) = s('type.semanticPrimitives.result'(bool,X1728),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1728),X1730)),X1728,bool))))) ) | ($spl0 | $spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f20708,f21143])).
fof(f20708,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X71,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f20707,f95])).
fof(f20707,plain,(
  ( ! [X70,X71] : (s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X71,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20706,f85])).
fof(f20706,plain,(
  ( ! [X70,X71] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X71))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X71,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20661,f86])).
fof(f20661,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X71))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X71,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(superposition,[],[f4658,f86])).
fof(f4658,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X1) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X1),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X0))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X0,sK6,sK3,X1,sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f648])).
fof(f648,plain,(
  ( ! [X10,X11,X9] : (s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK9(X11,X10,sK3,X9,sK6,sK5,sK2,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),X10),s(sK0,sK9(X11,X10,sK3,X9,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X11,X10,sK3,X9,sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X9),s(fun(sK0,sK2),X10),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X11))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X9)) ) | $spl0),
  inference(superposition,[],[f92,f77])).
fof(f77,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))) | s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,happ(s(fun(X0,X2),X4),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f68])).
fof(f30946,plain,(
  ( ! [X1727,X1722,X1724,X1723] : (s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724)))),sK3,sK6,X1727,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724))) = s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724)),X1722,bool))))) ) | $spl186),
  inference(cnf_transformation,[],[f30946_D])).
fof(f30946_D,plain,(
  ( ! [X1727,X1722,X1724,X1723] : (s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724)))),sK3,sK6,X1727,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724))) = s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724)),X1722,bool))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f30514,plain,(
  ( ! [X1725,X1727,X1722,X1724,X1726,X1723] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1725),s(bool,X1726))))) | s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1722),sK0),X1723),s('type.semanticPrimitives.error_result'(X1722),X1724)))),sK3,sK6,X1727,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724))) = s('type.semanticPrimitives.result'(bool,X1722),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1722),X1724)),X1722,bool))))) ) | ($spl0 | $spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f20525,f21143])).
fof(f20525,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK3,sK6,X71,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20524,f86])).
fof(f20524,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK3,sK6,X71,sK2,sK1,sK0))))) ) | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f20523,f95])).
fof(f20523,plain,(
  ( ! [X70,X71] : (s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK3,sK6,X71,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20522,f85])).
fof(f20522,plain,(
  ( ! [X70,X71] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X71))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK3,sK6,X71,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20476,f86])).
fof(f20476,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X70))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X71))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,X70)),sK3,sK6,X71,sK2,sK1,sK0))))) ) | $spl0),
  inference(superposition,[],[f4628,f86])).
fof(f4628,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1))) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X0,sK3,sK6,X1,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f647])).
fof(f647,plain,(
  ( ! [X6,X8,X7] : (s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK9(sK5,sK6,X6,sK3,X7,X8,sK2,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),X7),s(sK0,sK9(sK5,sK6,X6,sK3,X7,X8,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),X6) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X6,sK3,X7,X8,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X6),s(fun(sK0,sK2),X7),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X6)) ) | $spl0),
  inference(superposition,[],[f92,f77])).
fof(f30821,plain,(
  ( ! [X474,X475,X473] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X473),'type.semanticPrimitives.error_result'(sK1)),X474),s('type.semanticPrimitives.error_result'(X473),X475))))) | s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475))) = s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475)),X473,bool))))) ) | $spl40),
  inference(cnf_transformation,[],[f30821_D])).
fof(f30821_D,plain,(
  ( ! [X474,X475,X473] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X473),'type.semanticPrimitives.error_result'(sK1)),X474),s('type.semanticPrimitives.error_result'(X473),X475))))) | s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475))) = s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475)),X473,bool))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f30888,plain,(
  ( ! [X271,X270] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X270),s(bool,X271)))))) ) | $spl120),
  inference(cnf_transformation,[],[f30888_D])).
fof(f30888_D,plain,(
  ( ! [X271,X270] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X270),s(bool,X271)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f30513,plain,(
  ( ! [X1718,X1720,X1717,X1719,X1721] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1720),s(bool,X1721))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1717),'type.semanticPrimitives.error_result'(sK1)),X1718),s('type.semanticPrimitives.error_result'(X1717),X1719))))) | s('type.semanticPrimitives.result'(bool,X1717),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1717),X1719))) = s('type.semanticPrimitives.result'(bool,X1717),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1717),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1717),X1719)),X1717,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f94,f21143])).
fof(f30945,plain,(
  ( ! [X1711,X1712,X1716,X1713] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1716,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1711),'type.semanticPrimitives.error_result'(sK1)),X1712),s('type.semanticPrimitives.error_result'(X1711),X1713)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713))) = s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713)),X1711,bool))))) ) | $spl184),
  inference(cnf_transformation,[],[f30945_D])).
fof(f30945_D,plain,(
  ( ! [X1711,X1712,X1716,X1713] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1716,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1711),'type.semanticPrimitives.error_result'(sK1)),X1712),s('type.semanticPrimitives.error_result'(X1711),X1713)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713))) = s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713)),X1711,bool))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f30824,plain,(
  ( ! [X490,X485] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X485),s(bool,X490)))))) ) | $spl46),
  inference(cnf_transformation,[],[f30824_D])).
fof(f30824_D,plain,(
  ( ! [X490,X485] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X485),s(bool,X490)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f30512,plain,(
  ( ! [X1711,X1712,X1714,X1716,X1713,X1715] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1714),s(bool,X1715))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1716,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1711),'type.semanticPrimitives.error_result'(sK1)),X1712),s('type.semanticPrimitives.error_result'(X1711),X1713)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713))) = s('type.semanticPrimitives.result'(bool,X1711),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1711),X1713)),X1711,bool))))) ) | ($spl0 | $spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f20361,f21143])).
fof(f20361,plain,(
  ( ! [X68,X69] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X68))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X69,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),X68)),sK6,sK5,sK2,sK1,sK0))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f20360,f94])).
fof(f20360,plain,(
  ( ! [X68,X69] : (s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),X68))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X68))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X69,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),X68)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20359,f80])).
fof(f20359,plain,(
  ( ! [X68,X69] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X68))),s(fun(sK0,sK2),X69),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X68))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X69,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),X68)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(forward_demodulation,[],[f20315,f82])).
fof(f82,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),happ(s(fun('type.semanticPrimitives.error_result'(X1),'type.semanticPrimitives.result'(X0,X1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(X1),X2)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),happ(s(fun('type.semanticPrimitives.error_result'(X1),'type.semanticPrimitives.result'(X0,X1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2] : s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),happ(s(fun('type.semanticPrimitives.error_result'(X1),'type.semanticPrimitives.result'(X0,X1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : ! [X15] : s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X15))) = s('type.semanticPrimitives.result'(X6,X7),happ(s(fun('type.semanticPrimitives.error_result'(X7),'type.semanticPrimitives.result'(X6,X7)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(X7),X15)))),
  file('thm.closProps.eqs|split|5','ARITY_0_1_const.semanticPrimitives.Rerr')).
fof(f20315,plain,(
  ( ! [X68,X69] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X68))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X69,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),X68)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),X68))),s(fun(sK0,sK2),X69),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl0),
  inference(superposition,[],[f6737,f82])).
fof(f30944,plain,(
  ( ! [X1707,X1706,X1708] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1706),'type.semanticPrimitives.result'(sK0,sK1)),X1707),s('type.semanticPrimitives.error_result'(X1706),X1708))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708))) = s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708)),X1706,bool))))) ) | $spl182),
  inference(cnf_transformation,[],[f30944_D])).
fof(f30944_D,plain,(
  ( ! [X1707,X1706,X1708] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1706),'type.semanticPrimitives.result'(sK0,sK1)),X1707),s('type.semanticPrimitives.error_result'(X1706),X1708))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708))) = s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708)),X1706,bool))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f30938,plain,(
  ( ! [X1672,X1671] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1671),s(bool,X1672)))) ) | $spl170),
  inference(cnf_transformation,[],[f30938_D])).
fof(f30938_D,plain,(
  ( ! [X1672,X1671] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1671),s(bool,X1672)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f30511,plain,(
  ( ! [X1707,X1709,X1706,X1708,X1710] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1709),s(bool,X1710))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1706),'type.semanticPrimitives.result'(sK0,sK1)),X1707),s('type.semanticPrimitives.error_result'(X1706),X1708))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708))) = s('type.semanticPrimitives.result'(bool,X1706),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1706),X1708)),X1706,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f27612,f21143])).
fof(f30943,plain,(
  ( ! [X1701,X1705,X1700,X1702] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1705,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702))) = s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702)),X1700,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702))),s(fun(sK0,sK2),X1705),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl180),
  inference(cnf_transformation,[],[f30943_D])).
fof(f30943_D,plain,(
  ( ! [X1701,X1705,X1700,X1702] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1705,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702))) = s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702)),X1700,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702))),s(fun(sK0,sK2),X1705),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f30510,plain,(
  ( ! [X1701,X1703,X1705,X1700,X1702,X1704] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1703),s(bool,X1704))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1705,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1700),'type.semanticPrimitives.result'(sK0,sK1)),X1701),s('type.semanticPrimitives.error_result'(X1700),X1702))),s(fun(sK0,sK2),X1705),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702))) = s('type.semanticPrimitives.result'(bool,X1700),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1700),X1702)),X1700,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f6737,f21143])).
fof(f30942,plain,(
  ( ! [X1694,X1699,X1695,X1696] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696)),sK3,X1699,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696)),X1694,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))),s(fun(sK0,sK2),X1699),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl178),
  inference(cnf_transformation,[],[f30942_D])).
fof(f30942_D,plain,(
  ( ! [X1694,X1699,X1695,X1696] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696)),sK3,X1699,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696)),X1694,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))),s(fun(sK0,sK2),X1699),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f30509,plain,(
  ( ! [X1694,X1697,X1699,X1695,X1696,X1698] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1697),s(bool,X1698))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696)),sK3,X1699,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1694),'type.semanticPrimitives.result'(sK0,sK1)),X1695),s('type.semanticPrimitives.error_result'(X1694),X1696))),s(fun(sK0,sK2),X1699),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696))) = s('type.semanticPrimitives.result'(bool,X1694),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1694),X1696)),X1694,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f6537,f21143])).
fof(f30941,plain,(
  ( ! [X1689,X1693,X1688,X1690] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1693))) | s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690))) = s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690)),X1688,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1693,sK6,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl176),
  inference(cnf_transformation,[],[f30941_D])).
fof(f30941_D,plain,(
  ( ! [X1689,X1693,X1688,X1690] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1693))) | s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690))) = s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690)),X1688,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1693,sK6,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690)),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f30508,plain,(
  ( ! [X1692,X1689,X1691,X1693,X1688,X1690] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1691),s(bool,X1692))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1693))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1693,sK6,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1688),'type.semanticPrimitives.result'(sK0,sK1)),X1689),s('type.semanticPrimitives.error_result'(X1688),X1690)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690))) = s('type.semanticPrimitives.result'(bool,X1688),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1688),X1690)),X1688,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f4658,f21143])).
fof(f30940,plain,(
  ( ! [X1683,X1687,X1682,X1684] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1687))) | s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684)),X1682,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684)),sK3,sK6,X1687,sK2,sK1,sK0))))) ) | $spl174),
  inference(cnf_transformation,[],[f30940_D])).
fof(f30940_D,plain,(
  ( ! [X1683,X1687,X1682,X1684] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1687))) | s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684)),X1682,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684)),sK3,sK6,X1687,sK2,sK1,sK0))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f30507,plain,(
  ( ! [X1683,X1685,X1687,X1682,X1684,X1686] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1685),s(bool,X1686))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1687))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1682),'type.semanticPrimitives.result'(sK0,sK1)),X1683),s('type.semanticPrimitives.error_result'(X1682),X1684)),sK3,sK6,X1687,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684))) = s('type.semanticPrimitives.result'(bool,X1682),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1682),X1684)),X1682,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f4628,f21143])).
fof(f30939,plain,(
  ( ! [X1676,X1681,X1675,X1677,X1680] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677))) = s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677)),X1675,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677))),s(fun(sK0,sK2),X1681),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1680))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl172),
  inference(cnf_transformation,[],[f30939_D])).
fof(f30939_D,plain,(
  ( ! [X1676,X1681,X1675,X1677,X1680] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677))) = s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677)),X1675,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677))),s(fun(sK0,sK2),X1681),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1680))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f30506,plain,(
  ( ! [X1676,X1678,X1681,X1675,X1677,X1679,X1680] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1678),s(bool,X1679))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1680,X1681,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1675),'type.semanticPrimitives.result'(sK0,sK1)),X1676),s('type.semanticPrimitives.error_result'(X1675),X1677))),s(fun(sK0,sK2),X1681),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1680))) | s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677))) = s('type.semanticPrimitives.result'(bool,X1675),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1675),X1677)),X1675,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f646,f21143])).
fof(f646,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X3) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X5,X4,sK3,X3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X5,X4,sK3,X3,sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X3),s(fun(sK0,sK2),X4),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X5)))) ) | $spl0),
  inference(superposition,[],[f92,f76])).
fof(f76,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,X2),X7),s(fun('type.semanticPrimitives.error_result'(X1),X2),X8))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X6),s(fun(X0,X2),X4),s(fun('type.semanticPrimitives.error_result'(X1),X2),X3))) | s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK10(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X6) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK9(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X5) != s('type.semanticPrimitives.result'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f68])).
fof(f30937,plain,(
  ( ! [X1668,X1670,X1674,X1669,X1673] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670)),X1668,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))),s(fun(sK0,sK2),X1673),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1674))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0))))) ) | $spl168),
  inference(cnf_transformation,[],[f30937_D])).
fof(f30937_D,plain,(
  ( ! [X1668,X1670,X1674,X1669,X1673] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670)),X1668,bool)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))),s(fun(sK0,sK2),X1673),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1674))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f30505,plain,(
  ( ! [X1668,X1670,X1672,X1674,X1669,X1671,X1673] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1671),s(bool,X1672))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670)),sK3,X1673,X1674,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1668),'type.semanticPrimitives.result'(sK0,sK1)),X1669),s('type.semanticPrimitives.error_result'(X1668),X1670))),s(fun(sK0,sK2),X1673),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1674))) | s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670))) = s('type.semanticPrimitives.result'(bool,X1668),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1668),X1670)),X1668,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f645,f21143])).
fof(f645,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),X0) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,X0,sK3,X1,X2,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),X0) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,X0,sK3,X1,X2,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),X0),s(fun(sK0,sK2),X1),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X2)))) ) | $spl0),
  inference(superposition,[],[f92,f76])).
fof(f30936,plain,(
  ( ! [X1663,X1664,X1665] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun('type.semanticPrimitives.error_result'(X1663),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1664),s('type.semanticPrimitives.error_result'(X1663),X1665)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665))) = s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665)),X1663,bool))))) ) | $spl166),
  inference(cnf_transformation,[],[f30936_D])).
fof(f30936_D,plain,(
  ( ! [X1663,X1664,X1665] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun('type.semanticPrimitives.error_result'(X1663),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1664),s('type.semanticPrimitives.error_result'(X1663),X1665)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665))) = s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665)),X1663,bool))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f30935,plain,(
  ( ! [X1662,X1661] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1661),s(bool,X1662)))))) ) | $spl164),
  inference(cnf_transformation,[],[f30935_D])).
fof(f30935_D,plain,(
  ( ! [X1662,X1661] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1661),s(bool,X1662)))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f30504,plain,(
  ( ! [X1663,X1664,X1666,X1665,X1667] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1666),s(bool,X1667))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun('type.semanticPrimitives.error_result'(X1663),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1664),s('type.semanticPrimitives.error_result'(X1663),X1665)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665))) = s('type.semanticPrimitives.result'(bool,X1663),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1663),X1665)),X1663,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20703,f21143])).
fof(f20703,plain,(
  ( ! [X0] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X0))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X0,sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f4658])).
fof(f30934,plain,(
  ( ! [X1658,X1660,X1659] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1658),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1659),s('type.semanticPrimitives.error_result'(X1658),X1660)),sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660))) = s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660)),X1658,bool))))) ) | $spl162),
  inference(cnf_transformation,[],[f30934_D])).
fof(f30934_D,plain,(
  ( ! [X1658,X1660,X1659] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1658),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1659),s('type.semanticPrimitives.error_result'(X1658),X1660)),sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660))) = s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660)),X1658,bool))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f30503,plain,(
  ( ! [X1658,X1660,X1662,X1659,X1661] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1661),s(bool,X1662))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun('type.semanticPrimitives.error_result'(X1658),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1659),s('type.semanticPrimitives.error_result'(X1658),X1660)),sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660))) = s('type.semanticPrimitives.result'(bool,X1658),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1658),X1660)),X1658,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20518,f21143])).
fof(f20518,plain,(
  ( ! [X0] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X0))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,X0,sK2,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f4628])).
fof(f30933,plain,(
  ( ! [X1617,X1615,X1616] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun('type.semanticPrimitives.error_result'(X1615),fun(sK0,sK2)),X1616),s('type.semanticPrimitives.error_result'(X1615),X1617)),sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617))) = s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617)),X1615,bool))))) ) | $spl160),
  inference(cnf_transformation,[],[f30933_D])).
fof(f30933_D,plain,(
  ( ! [X1617,X1615,X1616] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun('type.semanticPrimitives.error_result'(X1615),fun(sK0,sK2)),X1616),s('type.semanticPrimitives.error_result'(X1615),X1617)),sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617))) = s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617)),X1615,bool))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f30932,plain,(
  ( ! [X1614,X1613] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun(bool,fun(sK0,sK2)),X1613),s(bool,X1614))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl158),
  inference(cnf_transformation,[],[f30932_D])).
fof(f30932_D,plain,(
  ( ! [X1614,X1613] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun(bool,fun(sK0,sK2)),X1613),s(bool,X1614))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f30496,plain,(
  ( ! [X1617,X1619,X1615,X1616,X1618] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun(bool,fun(sK0,sK2)),X1618),s(bool,X1619))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun('type.semanticPrimitives.error_result'(X1615),fun(sK0,sK2)),X1616),s('type.semanticPrimitives.error_result'(X1615),X1617)),sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617))) = s('type.semanticPrimitives.result'(bool,X1615),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1615),X1617)),X1615,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20358,f21143])).
fof(f30931,plain,(
  ( ! [X1610,X1612,X1611] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1610),fun(sK0,sK2)),X1611),s('type.semanticPrimitives.error_result'(X1610),X1612)),sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612))) = s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612)),X1610,bool))))) ) | $spl156),
  inference(cnf_transformation,[],[f30931_D])).
fof(f30931_D,plain,(
  ( ! [X1610,X1612,X1611] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1610),fun(sK0,sK2)),X1611),s('type.semanticPrimitives.error_result'(X1610),X1612)),sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612))) = s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612)),X1610,bool))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f30495,plain,(
  ( ! [X1610,X1612,X1614,X1611,X1613] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun(bool,fun(sK0,sK2)),X1613),s(bool,X1614))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun('type.semanticPrimitives.error_result'(X1610),fun(sK0,sK2)),X1611),s('type.semanticPrimitives.error_result'(X1610),X1612)),sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612))) = s('type.semanticPrimitives.result'(bool,X1610),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1610),X1612)),X1610,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f19170,f21143])).
fof(f30928,plain,(
  ( ! [X1562,X1557,X1556,X1558] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1556),fun(bool,sK0)),X1557),s('type.semanticPrimitives.error_result'(X1556),X1558))),s(bool,X1562))))) | s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558))) = s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558)),X1556,bool))))) ) | $spl150),
  inference(cnf_transformation,[],[f30928_D])).
fof(f30928_D,plain,(
  ( ! [X1562,X1557,X1556,X1558] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1556),fun(bool,sK0)),X1557),s('type.semanticPrimitives.error_result'(X1556),X1558))),s(bool,X1562))))) | s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558))) = s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558)),X1556,bool))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f30930,plain,(
  ( ! [X1566,X1568,X1567] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1566),s(bool,X1567))),s(bool,X1568)))))) ) | $spl154),
  inference(cnf_transformation,[],[f30930_D])).
fof(f30930_D,plain,(
  ( ! [X1566,X1568,X1567] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1566),s(bool,X1567))),s(bool,X1568)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f30488,plain,(
  ( ! [X1569,X1564,X1566,X1568,X1563,X1567,X1565] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1566),s(bool,X1567))),s(bool,X1568))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1563),fun(bool,sK0)),X1564),s('type.semanticPrimitives.error_result'(X1563),X1565))),s(bool,X1569))))) | s('type.semanticPrimitives.result'(bool,X1563),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1563),X1565))) = s('type.semanticPrimitives.result'(bool,X1563),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1563),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1563),X1565)),X1563,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f2363,f21143])).
fof(f2363,plain,(
  ( ! [X263,X264,X266] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X263),s(bool,X266))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X263),s(bool,X264)))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(forward_demodulation,[],[f2178,f85])).
fof(f2178,plain,(
  ( ! [X263,X265,X267,X264,X266] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rval_1'(s(bool,X266))),s(fun(bool,sK0),X263),s(fun('type.semanticPrimitives.error_result'(X265),sK0),X267))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X263),s(bool,X264)))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f95,f1874])).
fof(f1874,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(X58,happ(s(fun(bool,X58),X59),s(bool,X56))) = s(X58,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,X57))),s(fun(bool,X58),X59),s(fun('type.semanticPrimitives.error_result'(X55),X58),X60)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f85,f1801])).
fof(f1801,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X1))) = s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) ) | ($spl16 | $spl28)),
  inference(equality_resolution,[],[f1742])).
fof(f1742,plain,(
  ( ! [X35,X33,X38,X34] : (s('type.semanticPrimitives.result'(bool,X33),X34) != s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s('type.semanticPrimitives.result'(bool,X33),X34) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X38)))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1741,f1303])).
fof(f1303,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.semanticPrimitives.result'(bool,X22),'const.semanticPrimitives.Rval_1'(s(bool,X21))) != s('type.semanticPrimitives.result'(bool,X22),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X22),sK12(X23,X22,bool)))) | s('type.semanticPrimitives.result'(bool,X22),X23) = s('type.semanticPrimitives.result'(bool,X22),'const.semanticPrimitives.Rval_1'(s(bool,X20)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1278,f771])).
fof(f1278,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X2))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X0),sK12(X1,X0,bool)))) | s('type.semanticPrimitives.result'(bool,X0),X1) = s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) ) | ($spl16 | $spl28)),
  inference(equality_factoring,[],[f964])).
fof(f1741,plain,(
  ( ! [X37,X35,X33,X38,X34] : (s('type.semanticPrimitives.result'(bool,X33),X34) != s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X37))) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X33),sK12(X34,X33,bool)))) | s('type.semanticPrimitives.result'(bool,X33),X34) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X38)))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1689,f1303])).
fof(f1689,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s('type.semanticPrimitives.result'(bool,X33),X34) != s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X36))) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X33),sK12(X34,X33,bool)))) | s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X37))) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X33),sK12(X34,X33,bool)))) | s('type.semanticPrimitives.result'(bool,X33),X34) = s('type.semanticPrimitives.result'(bool,X33),'const.semanticPrimitives.Rval_1'(s(bool,X38)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1600,f964])).
fof(f1600,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s('type.semanticPrimitives.result'(bool,X6),X7) != s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X8))) | s('type.semanticPrimitives.result'(bool,X6),X7) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X9))) | s('type.semanticPrimitives.result'(bool,X6),X7) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X10)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1303,f964])).
fof(f30929,plain,(
  ( ! [X1560,X1559,X1561] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1559),s(bool,X1560))))),s(bool,X1561)))) ) | $spl152),
  inference(cnf_transformation,[],[f30929_D])).
fof(f30929_D,plain,(
  ( ! [X1560,X1559,X1561] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1559),s(bool,X1560))))),s(bool,X1561)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f30487,plain,(
  ( ! [X1560,X1562,X1557,X1559,X1561,X1556,X1558] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1559),s(bool,X1560))))),s(bool,X1561))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1556),fun(bool,sK0)),X1557),s('type.semanticPrimitives.error_result'(X1556),X1558))),s(bool,X1562))))) | s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558))) = s('type.semanticPrimitives.result'(bool,X1556),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1556),X1558)),X1556,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f3828,f21143])).
fof(f3828,plain,(
  ( ! [X103,X101,X102] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),X101))),s(bool,X103))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X101),s(bool,X102)))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f95,f2355])).
fof(f2355,plain,(
  ( ! [X92,X90,X88,X87,X85,X89] : (s(X85,happ(s(fun(X88,X85),X89),s(X88,happ(s(fun(bool,X88),X90),s(bool,X87))))) = s(X85,happ(s(fun(bool,X85),'_dst_x0x1_2'(s(fun(X88,X85),X89),s(fun(bool,X88),X90))),s(bool,X92)))) ) | ($spl16 | $spl28)),
  inference(forward_demodulation,[],[f2133,f85])).
fof(f2133,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (s(X85,happ(s(fun(bool,X85),'_dst_x0x1_2'(s(fun(X88,X85),X89),s(fun(bool,X88),X90))),s(bool,X92))) = s(X85,happ(s(fun(X88,X85),X89),s(X88,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X86),'const.semanticPrimitives.Rval_1'(s(bool,X87))),s(fun(bool,X88),X90),s(fun('type.semanticPrimitives.error_result'(X86),X88),X91)))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1874,f83])).
fof(f83,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X37,X7,X8] : ! [X27,X19,X12,X18] : s(X37,happ(s(fun(X8,X37),X12),s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X27),s(fun(X6,X8),X18),s(fun('type.semanticPrimitives.error_result'(X7),X8),X19))))) = s(X37,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X27),s(fun(X6,X37),'_dst_x0x1_2'(s(fun(X8,X37),X12),s(fun(X6,X8),X18))),s(fun('type.semanticPrimitives.error_result'(X7),X37),'_dst_x0x1_2'(s(fun(X8,X37),X12),s(fun('type.semanticPrimitives.error_result'(X7),X8),X19)))))),
  file('thm.closProps.eqs|split|5','thm.closProps.closProps_76_prove_ltb')).
fof(f30926,plain,(
  ( ! [X1550,X1555,X1549,X1551] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1549),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1550),s('type.semanticPrimitives.error_result'(X1549),X1551))),s(bool,X1555))))) | s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551))) = s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551)),X1549,bool))))) ) | $spl146),
  inference(cnf_transformation,[],[f30926_D])).
fof(f30926_D,plain,(
  ( ! [X1550,X1555,X1549,X1551] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1549),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1550),s('type.semanticPrimitives.error_result'(X1549),X1551))),s(bool,X1555))))) | s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551))) = s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551)),X1549,bool))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f30927,plain,(
  ( ! [X1553,X1552,X1554] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1552),s(bool,X1553))),s(bool,X1554)))))) ) | $spl148),
  inference(cnf_transformation,[],[f30927_D])).
fof(f30927_D,plain,(
  ( ! [X1553,X1552,X1554] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1552),s(bool,X1553))),s(bool,X1554)))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f30486,plain,(
  ( ! [X1550,X1553,X1555,X1549,X1551,X1552,X1554] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1552),s(bool,X1553))),s(bool,X1554))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1549),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1550),s('type.semanticPrimitives.error_result'(X1549),X1551))),s(bool,X1555))))) | s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551))) = s('type.semanticPrimitives.result'(bool,X1549),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1549),X1551)),X1549,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f2362,f21143])).
fof(f2362,plain,(
  ( ! [X261,X259,X258] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X258),s(bool,X261))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X258),s(bool,X259)))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(forward_demodulation,[],[f2177,f85])).
fof(f2177,plain,(
  ( ! [X261,X259,X260,X262,X258] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X260),'const.semanticPrimitives.Rval_1'(s(bool,X261))),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X258),s(fun('type.semanticPrimitives.error_result'(X260),'type.semanticPrimitives.error_result'(sK1)),X262))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X258),s(bool,X259)))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f94,f1874])).
fof(f30917,plain,(
  s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8)),sK1,bool)))) | $spl144),
  inference(cnf_transformation,[],[f30917_D])).
fof(f30917_D,plain,(
  s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8)),sK1,bool)))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f30910,plain,(
  ( ! [X457,X456] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X456),s(bool,X457)))) ) | $spl132),
  inference(cnf_transformation,[],[f30910_D])).
fof(f30910_D,plain,(
  ( ! [X457,X456] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X456),s(bool,X457)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f30360,plain,(
  ( ! [X498,X497] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X497),s(bool,X498))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8)),sK1,bool))))) ) | ($spl9 | $spl16 | $spl28)),
  inference(superposition,[],[f98,f21143])).
fof(f98,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | $spl9),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f30915,plain,(
  ( ! [X492,X496,X493] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X496))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493)))),sK1,bool))))) ) | $spl142),
  inference(cnf_transformation,[],[f30915_D])).
fof(f30915_D,plain,(
  ( ! [X492,X496,X493] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X496))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493)))),sK1,bool))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f30359,plain,(
  ( ! [X492,X494,X496,X493,X495] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X494),s(bool,X495))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X496))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X492),s(bool,X493)))),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f2362,f21143])).
fof(f30914,plain,(
  ( ! [X486,X488,X485,X487,X491] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(X485,happ(s(fun(bool,X485),X487),s(bool,X491))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488)))),sK1,bool))))) ) | $spl140),
  inference(cnf_transformation,[],[f30914_D])).
fof(f30914_D,plain,(
  ( ! [X486,X488,X485,X487,X491] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(X485,happ(s(fun(bool,X485),X487),s(bool,X491))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488)))),sK1,bool))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f30358,plain,(
  ( ! [X486,X488,X490,X485,X487,X489,X491] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X489),s(bool,X490))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(X485,happ(s(fun(bool,X485),X487),s(bool,X491))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X485,'type.semanticPrimitives.error_result'(sK1)),X486),s(fun(bool,X485),X487))),s(bool,X488)))),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1678,f21143])).
fof(f1678,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153)))))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1677,f819])).
fof(f819,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f506,f736])).
fof(f1677,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1645,f819])).
fof(f1645,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153)) | ~p(s(bool,X154))) ) | $spl2),
  inference(superposition,[],[f1038,f90])).
fof(f1038,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X85,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X83,'type.semanticPrimitives.error_result'(sK1)),X84),s(fun(X85,X83),X86))),s(X85,X87))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X83,'type.semanticPrimitives.error_result'(sK1)),X84),s(X83,happ(s(fun(X85,X83),X86),s(X85,X87)))))))) ) | $spl2),
  inference(superposition,[],[f94,f84])).
fof(f30913,plain,(
  ( ! [X478,X480,X482,X479,X481] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(X479,happ(s(fun(X478,X479),X481),s(X478,X482))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482)))),sK1,bool))))) ) | $spl138),
  inference(cnf_transformation,[],[f30913_D])).
fof(f30913_D,plain,(
  ( ! [X478,X480,X482,X479,X481] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(X479,happ(s(fun(X478,X479),X481),s(X478,X482))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482)))),sK1,bool))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f30357,plain,(
  ( ! [X478,X484,X480,X482,X479,X481,X483] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X483),s(bool,X484))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(X479,happ(s(fun(X478,X479),X481),s(X478,X482))))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X478,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X479,'type.semanticPrimitives.error_result'(sK1)),X480),s(fun(X478,X479),X481))),s(X478,X482)))),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1038,f21143])).
fof(f30912,plain,(
  ( ! [X476,X472,X469,X471,X475,X477,X468,X470] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X475),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X476),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X477))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472)))),sK1,bool)))) | s('type.semanticPrimitives.result'(X468,X469),X470) != s('type.semanticPrimitives.result'(X468,X469),X475) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rval_1'(s(X468,sK9(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468)))) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X469),sK10(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468))))) ) | $spl136),
  inference(cnf_transformation,[],[f30912_D])).
fof(f30912_D,plain,(
  ( ! [X476,X472,X469,X471,X475,X477,X468,X470] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X475),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X476),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X477))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472)))),sK1,bool)))) | s('type.semanticPrimitives.result'(X468,X469),X470) != s('type.semanticPrimitives.result'(X468,X469),X475) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rval_1'(s(X468,sK9(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468)))) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X469),sK10(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f30356,plain,(
  ( ! [X476,X472,X474,X469,X471,X475,X477,X473,X468,X470] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X473),s(bool,X474))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X475),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X476),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X477))))) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X469),sK10(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468)))) | s('type.semanticPrimitives.result'(X468,X469),X470) = s('type.semanticPrimitives.result'(X468,X469),'const.semanticPrimitives.Rval_1'(s(X468,sK9(X477,X476,X470,X475,X471,X472,'type.semanticPrimitives.error_result'(sK1),X469,X468)))) | s('type.semanticPrimitives.result'(X468,X469),X470) != s('type.semanticPrimitives.result'(X468,X469),X475) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X468,X469),X470),s(fun(X468,'type.semanticPrimitives.error_result'(sK1)),X471),s(fun('type.semanticPrimitives.error_result'(X469),'type.semanticPrimitives.error_result'(sK1)),X472)))),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1026,f21143])).
fof(f1026,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X5),s(fun(X0,'type.semanticPrimitives.error_result'(sK1)),X6),s(fun('type.semanticPrimitives.error_result'(X1),'type.semanticPrimitives.error_result'(sK1)),X7))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X2),s(fun(X0,'type.semanticPrimitives.error_result'(sK1)),X3),s(fun('type.semanticPrimitives.error_result'(X1),'type.semanticPrimitives.error_result'(sK1)),X4))))) | s('type.semanticPrimitives.result'(X0,X1),X5) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK10(X4,X3,X5,X2,X6,X7,'type.semanticPrimitives.error_result'(sK1),X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X5) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK9(X4,X3,X5,X2,X6,X7,'type.semanticPrimitives.error_result'(sK1),X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) != s('type.semanticPrimitives.result'(X0,X1),X5)) ) | $spl2),
  inference(superposition,[],[f94,f76])).
fof(f30911,plain,(
  ( ! [X466,X461,X459,X465,X467,X460,X462,X458] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X465),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X466),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X467))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462)))),sK1,bool)))) | s('type.semanticPrimitives.result'(X458,X459),X460) != s('type.semanticPrimitives.result'(X458,X459),X465) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rval_1'(s(X458,sK9(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458)))) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X459),sK10(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458))))) ) | $spl134),
  inference(cnf_transformation,[],[f30911_D])).
fof(f30911_D,plain,(
  ( ! [X466,X461,X459,X465,X467,X460,X462,X458] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X465),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X466),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X467))))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462)))),sK1,bool)))) | s('type.semanticPrimitives.result'(X458,X459),X460) != s('type.semanticPrimitives.result'(X458,X459),X465) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rval_1'(s(X458,sK9(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458)))) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X459),sK10(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f30355,plain,(
  ( ! [X466,X461,X463,X459,X465,X467,X460,X462,X458,X464] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X463),s(bool,X464))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X465),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X466),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X467))))) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X459),sK10(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458)))) | s('type.semanticPrimitives.result'(X458,X459),X465) = s('type.semanticPrimitives.result'(X458,X459),'const.semanticPrimitives.Rval_1'(s(X458,sK9(X462,X461,X465,X460,X466,X467,'type.semanticPrimitives.error_result'(sK1),X459,X458)))) | s('type.semanticPrimitives.result'(X458,X459),X460) != s('type.semanticPrimitives.result'(X458,X459),X465) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462))))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X458,X459),X460),s(fun(X458,'type.semanticPrimitives.error_result'(sK1)),X461),s(fun('type.semanticPrimitives.error_result'(X459),'type.semanticPrimitives.error_result'(sK1)),X462)))),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1027,f21143])).
fof(f1027,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X8,X9),X13),s(fun(X8,'type.semanticPrimitives.error_result'(sK1)),X14),s(fun('type.semanticPrimitives.error_result'(X9),'type.semanticPrimitives.error_result'(sK1)),X15))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X8,X9),X10),s(fun(X8,'type.semanticPrimitives.error_result'(sK1)),X11),s(fun('type.semanticPrimitives.error_result'(X9),'type.semanticPrimitives.error_result'(sK1)),X12))))) | s('type.semanticPrimitives.result'(X8,X9),X10) = s('type.semanticPrimitives.result'(X8,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK10(X15,X14,X10,X13,X11,X12,'type.semanticPrimitives.error_result'(sK1),X9,X8)))) | s('type.semanticPrimitives.result'(X8,X9),X10) = s('type.semanticPrimitives.result'(X8,X9),'const.semanticPrimitives.Rval_1'(s(X8,sK9(X15,X14,X10,X13,X11,X12,'type.semanticPrimitives.error_result'(sK1),X9,X8)))) | s('type.semanticPrimitives.result'(X8,X9),X10) != s('type.semanticPrimitives.result'(X8,X9),X13)) ) | $spl2),
  inference(superposition,[],[f94,f76])).
fof(f30909,plain,(
  ( ! [X455] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455)),sK1,bool))))) ) | $spl130),
  inference(cnf_transformation,[],[f30909_D])).
fof(f30909_D,plain,(
  ( ! [X455] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455)),sK1,bool))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f30354,plain,(
  ( ! [X455,X457,X456] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),X456),s(bool,X457))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) | s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455))) = s('type.semanticPrimitives.result'(bool,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X455)),sK1,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f94,f21143])).
fof(f30908,plain,(
  ( ! [X447,X449,X451,X446,X452,X448,X450] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X448,sK0),X449),s(X448,happ(s(fun(X447,X448),X450),s(X447,happ(s(fun('type.semanticPrimitives.error_result'(X446),X447),X451),s('type.semanticPrimitives.error_result'(X446),X452))))))))) | s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452))) = s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452)),X446,bool))))) ) | $spl128),
  inference(cnf_transformation,[],[f30908_D])).
fof(f30908_D,plain,(
  ( ! [X447,X449,X451,X446,X452,X448,X450] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X448,sK0),X449),s(X448,happ(s(fun(X447,X448),X450),s(X447,happ(s(fun('type.semanticPrimitives.error_result'(X446),X447),X451),s('type.semanticPrimitives.error_result'(X446),X452))))))))) | s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452))) = s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452)),X446,bool))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f30353,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X453),s(bool,X454))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X448,sK0),X449),s(X448,happ(s(fun(X447,X448),X450),s(X447,happ(s(fun('type.semanticPrimitives.error_result'(X446),X447),X451),s('type.semanticPrimitives.error_result'(X446),X452))))))))) | s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452))) = s('type.semanticPrimitives.result'(bool,X446),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X446),X452)),X446,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1105,f21143])).
fof(f1105,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X9,sK0),'_dst_x0x1_2'(s(fun(X5,sK0),'_dst_x0x1_2'(s(fun(X6,sK0),X7),s(fun(X5,X6),X8))),s(fun(X9,X5),X10))),s(X9,X11))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X6,sK0),X7),s(X6,happ(s(fun(X5,X6),X8),s(X5,happ(s(fun(X9,X5),X10),s(X9,X11)))))))))) ) | $spl4),
  inference(superposition,[],[f1020,f84])).
fof(f30902,plain,(
  ( ! [X399,X401,X403,X400,X402] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X400,sK0),X401),s(X400,happ(s(fun('type.semanticPrimitives.error_result'(X399),X400),X402),s('type.semanticPrimitives.error_result'(X399),X403))))))) | s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403))) = s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403)),X399,bool))))) ) | $spl124),
  inference(cnf_transformation,[],[f30902_D])).
fof(f30902_D,plain,(
  ( ! [X399,X401,X403,X400,X402] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X400,sK0),X401),s(X400,happ(s(fun('type.semanticPrimitives.error_result'(X399),X400),X402),s('type.semanticPrimitives.error_result'(X399),X403))))))) | s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403))) = s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403)),X399,bool))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f30348,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),X404),s(bool,X405))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X400,sK0),X401),s(X400,happ(s(fun('type.semanticPrimitives.error_result'(X399),X400),X402),s('type.semanticPrimitives.error_result'(X399),X403))))))) | s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403))) = s('type.semanticPrimitives.result'(bool,X399),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X399),X403)),X399,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1020,f21143])).
fof(f30901,plain,(
  ( ! [X393,X395,X396,X392,X394] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X393,sK0),X394),s(X393,happ(s(fun('type.semanticPrimitives.error_result'(X392),X393),X395),s('type.semanticPrimitives.error_result'(X392),X396))))))) | s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396))) = s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396)),X392,bool))))) ) | $spl122),
  inference(cnf_transformation,[],[f30901_D])).
fof(f30901_D,plain,(
  ( ! [X393,X395,X396,X392,X394] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X393,sK0),X394),s(X393,happ(s(fun('type.semanticPrimitives.error_result'(X392),X393),X395),s('type.semanticPrimitives.error_result'(X392),X396))))))) | s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396))) = s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396)),X392,bool))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f30347,plain,(
  ( ! [X397,X393,X395,X396,X398,X392,X394] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X397),s(bool,X398))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X393,sK0),X394),s(X393,happ(s(fun('type.semanticPrimitives.error_result'(X392),X393),X395),s('type.semanticPrimitives.error_result'(X392),X396))))))) | s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396))) = s('type.semanticPrimitives.result'(bool,X392),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X392),X396)),X392,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1021,f21143])).
fof(f1021,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X88,sK0),'_dst_x0x1_2'(s(fun(X89,sK0),X90),s(fun(X88,X89),X91))),s(X88,X92))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X89,sK0),X90),s(X89,happ(s(fun(X88,X89),X91),s(X88,X92)))))))) ) | $spl4),
  inference(superposition,[],[f95,f84])).
fof(f30887,plain,(
  ( ! [X269,X265,X267,X268,X266] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X266,'type.semanticPrimitives.error_result'(sK1)),X267),s(X266,happ(s(fun('type.semanticPrimitives.error_result'(X265),X266),X268),s('type.semanticPrimitives.error_result'(X265),X269))))))) | s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269))) = s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269)),X265,bool))))) ) | $spl118),
  inference(cnf_transformation,[],[f30887_D])).
fof(f30887_D,plain,(
  ( ! [X269,X265,X267,X268,X266] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X266,'type.semanticPrimitives.error_result'(sK1)),X267),s(X266,happ(s(fun('type.semanticPrimitives.error_result'(X265),X266),X268),s('type.semanticPrimitives.error_result'(X265),X269))))))) | s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269))) = s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269)),X265,bool))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f30334,plain,(
  ( ! [X269,X271,X265,X267,X268,X270,X266] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X270),s(bool,X271))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X266,'type.semanticPrimitives.error_result'(sK1)),X267),s(X266,happ(s(fun('type.semanticPrimitives.error_result'(X265),X266),X268),s('type.semanticPrimitives.error_result'(X265),X269))))))) | s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269))) = s('type.semanticPrimitives.result'(bool,X265),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X265),X269)),X265,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1038,f21143])).
fof(f30886,plain,(
  ( ! [X261,X259,X260,X262,X258] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X259,'type.semanticPrimitives.error_result'(sK1)),X260),s(X259,happ(s(fun('type.semanticPrimitives.error_result'(X258),X259),X261),s('type.semanticPrimitives.error_result'(X258),X262))))))) | s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262))) = s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262)),X258,bool))))) ) | $spl116),
  inference(cnf_transformation,[],[f30886_D])).
fof(f30886_D,plain,(
  ( ! [X261,X259,X260,X262,X258] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X259,'type.semanticPrimitives.error_result'(sK1)),X260),s(X259,happ(s(fun('type.semanticPrimitives.error_result'(X258),X259),X261),s('type.semanticPrimitives.error_result'(X258),X262))))))) | s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262))) = s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262)),X258,bool))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f30333,plain,(
  ( ! [X261,X263,X259,X260,X262,X258,X264] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X263),s(bool,X264))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X259,'type.semanticPrimitives.error_result'(sK1)),X260),s(X259,happ(s(fun('type.semanticPrimitives.error_result'(X258),X259),X261),s('type.semanticPrimitives.error_result'(X258),X262))))))) | s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262))) = s('type.semanticPrimitives.result'(bool,X258),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X258),X262)),X258,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1039,f21143])).
fof(f1039,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X88,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X89,'type.semanticPrimitives.error_result'(sK1)),X90),s(fun(X88,X89),X91))),s(X88,X92))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X89,'type.semanticPrimitives.error_result'(sK1)),X90),s(X89,happ(s(fun(X88,X89),X91),s(X88,X92)))))))) ) | $spl2),
  inference(superposition,[],[f94,f84])).
fof(f30780,plain,(
  ( ! [X11,X9] : (s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)),X9,bool))))) ) | $spl30),
  inference(cnf_transformation,[],[f30780_D])).
fof(f30780_D,plain,(
  ( ! [X11,X9] : (s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)),X9,bool))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f30868,plain,(
  ( ! [X2732,X2730,X2731,X2737] : (s('type.semanticPrimitives.result'(bool,X2730),'const.semanticPrimitives.Rval_1'(s(bool,X2737))) = s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2730)),X2731),s(bool,X2732)))) ) | $spl114),
  inference(cnf_transformation,[],[f30868_D])).
fof(f30868_D,plain,(
  ( ! [X2732,X2730,X2731,X2737] : (s('type.semanticPrimitives.result'(bool,X2730),'const.semanticPrimitives.Rval_1'(s(bool,X2737))) = s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2730)),X2731),s(bool,X2732)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f30869,plain,(
  ( ! [X2757,X2759,X2761,X2756,X2758,X2764] : (s('type.semanticPrimitives.result'(bool,X2756),'const.semanticPrimitives.Rval_1'(s(bool,X2764))) = s('type.semanticPrimitives.result'(bool,X2756),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2756)),X2757),s(bool,X2758))) | s('type.semanticPrimitives.result'(bool,X2759),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),X2761))) = s('type.semanticPrimitives.result'(bool,X2759),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),X2761)),X2759,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f30222,f21143])).
fof(f30222,plain,(
  ( ! [X2757,X2759,X2761,X2763,X2756,X2758,X2764,X2760,X2762] : (s('type.semanticPrimitives.result'(bool,X2756),happ(s(fun('type.semanticPrimitives.error_result'(X2759),'type.semanticPrimitives.result'(bool,X2756)),X2760),s('type.semanticPrimitives.error_result'(X2759),X2761))) != s('type.semanticPrimitives.result'(bool,X2756),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2756)),'_dst_x0x1_2'(s(fun(bool,'type.semanticPrimitives.result'(bool,X2756)),'const.semanticPrimitives.Rval_0'),s(fun(bool,bool),X2762))),s(bool,X2763))) | s('type.semanticPrimitives.result'(bool,X2756),'const.semanticPrimitives.Rval_1'(s(bool,X2764))) = s('type.semanticPrimitives.result'(bool,X2756),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2756)),X2757),s(bool,X2758))) | s('type.semanticPrimitives.result'(bool,X2759),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),X2761))) = s('type.semanticPrimitives.result'(bool,X2759),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2759),X2761)),X2759,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f11064,f21143])).
fof(f11064,plain,(
  ( ! [X61,X66,X64,X62,X65] : (s('type.semanticPrimitives.result'(bool,X61),X65) != s('type.semanticPrimitives.result'(bool,X61),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X61)),'_dst_x0x1_2'(s(fun(bool,'type.semanticPrimitives.result'(bool,X61)),'const.semanticPrimitives.Rval_0'),s(fun(bool,bool),X62))),s(bool,X64))) | s('type.semanticPrimitives.result'(bool,X61),X65) = s('type.semanticPrimitives.result'(bool,X61),'const.semanticPrimitives.Rval_1'(s(bool,X66)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1742,f3727])).
fof(f3727,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s('type.semanticPrimitives.result'(X96,X97),'const.semanticPrimitives.Rval_1'(s(X96,happ(s(fun(bool,X96),X98),s(bool,X99))))) = s('type.semanticPrimitives.result'(X96,X97),happ(s(fun(bool,'type.semanticPrimitives.result'(X96,X97)),'_dst_x0x1_2'(s(fun(X96,'type.semanticPrimitives.result'(X96,X97)),'const.semanticPrimitives.Rval_0'),s(fun(bool,X96),X98))),s(bool,X100)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f2355,f86])).
fof(f30867,plain,(
  ( ! [X2732,X2730,X2733,X2735,X2731,X2737] : (s('type.semanticPrimitives.result'(bool,X2730),'const.semanticPrimitives.Rval_1'(s(bool,X2737))) = s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2730)),X2731),s(bool,X2732))) | s('type.semanticPrimitives.result'(bool,X2733),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),X2735))) = s('type.semanticPrimitives.result'(bool,X2733),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),X2735)),X2733,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f30219,f21143])).
fof(f30219,plain,(
  ( ! [X2732,X2734,X2730,X2736,X2733,X2735,X2731,X2737] : (s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2730)),'const.semanticPrimitives.Rval_0'),s(bool,X2736))) != s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun('type.semanticPrimitives.error_result'(X2733),'type.semanticPrimitives.result'(bool,X2730)),X2734),s('type.semanticPrimitives.error_result'(X2733),X2735))) | s('type.semanticPrimitives.result'(bool,X2730),'const.semanticPrimitives.Rval_1'(s(bool,X2737))) = s('type.semanticPrimitives.result'(bool,X2730),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X2730)),X2731),s(bool,X2732))) | s('type.semanticPrimitives.result'(bool,X2733),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),X2735))) = s('type.semanticPrimitives.result'(bool,X2733),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2733),X2735)),X2733,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f2903,f21143])).
fof(f2903,plain,(
  ( ! [X33,X31,X29,X32] : (s('type.semanticPrimitives.result'(bool,X29),X32) != s('type.semanticPrimitives.result'(bool,X29),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X29)),'const.semanticPrimitives.Rval_0'),s(bool,X31))) | s('type.semanticPrimitives.result'(bool,X29),X32) = s('type.semanticPrimitives.result'(bool,X29),'const.semanticPrimitives.Rval_1'(s(bool,X33)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1742,f2502])).
fof(f2502,plain,(
  ( ! [X57,X58,X56] : (s('type.semanticPrimitives.result'(bool,X56),'const.semanticPrimitives.Rval_1'(s(bool,X57))) = s('type.semanticPrimitives.result'(bool,X56),happ(s(fun(bool,'type.semanticPrimitives.result'(bool,X56)),'const.semanticPrimitives.Rval_0'),s(bool,X58)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f2123,f86])).
fof(f2123,plain,(
  ( ! [X6,X0,X5,X3] : (s(X0,happ(s(fun(bool,X0),X3),s(bool,X5))) = s(X0,happ(s(fun(bool,X0),X3),s(bool,X6)))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f1874,f1874])).
fof(f30832,plain,(
  ( ! [X549,X547,X548] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X547),sK0),X548),s('type.semanticPrimitives.error_result'(X547),X549))))) | s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549))) = s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549)),X547,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f30832_D])).
fof(f30832_D,plain,(
  ( ! [X549,X547,X548] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X547),sK0),X548),s('type.semanticPrimitives.error_result'(X547),X549))))) | s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549))) = s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549)),X547,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f30131,plain,(
  ( ! [X1839,X1837,X1835,X1838,X1836] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1837),sK0),X1838),s('type.semanticPrimitives.error_result'(X1837),X1839))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1835),s(bool,X1836))))) | s('type.semanticPrimitives.result'(bool,X1837),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1837),X1839))) = s('type.semanticPrimitives.result'(bool,X1837),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1837),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1837),X1839)),X1837,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f95,f21143])).
fof(f30865,plain,(
  ( ! [X1829,X1834,X1830] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1834,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1829),s(bool,X1830)))),sK6,sK5,sK2,sK1,sK0))))) ) | $spl112),
  inference(cnf_transformation,[],[f30865_D])).
fof(f30865_D,plain,(
  ( ! [X1829,X1834,X1830] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1834,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1829),s(bool,X1830)))),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f30130,plain,(
  ( ! [X1833,X1831,X1829,X1834,X1832,X1830] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1831),sK0),X1832),s('type.semanticPrimitives.error_result'(X1831),X1833))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1834,sK6,sK3,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1829),s(bool,X1830)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1831),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1831),X1833))) = s('type.semanticPrimitives.result'(bool,X1831),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1831),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1831),X1833)),X1831,bool))))) ) | ($spl0 | $spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f20708,f21143])).
fof(f30864,plain,(
  ( ! [X1824,X1823,X1828] : (s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824)))),sK3,sK6,X1828,sK2,sK1,sK0))))) ) | $spl110),
  inference(cnf_transformation,[],[f30864_D])).
fof(f30864_D,plain,(
  ( ! [X1824,X1823,X1828] : (s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824)))),sK3,sK6,X1828,sK2,sK1,sK0))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f30129,plain,(
  ( ! [X1824,X1827,X1825,X1823,X1828,X1826] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X1825),sK0),X1826),s('type.semanticPrimitives.error_result'(X1825),X1827))))) | s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824))))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(sK0,'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rval_0'),s(sK0,happ(s(fun(bool,sK0),X1823),s(bool,X1824)))),sK3,sK6,X1828,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1825),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1825),X1827))) = s('type.semanticPrimitives.result'(bool,X1825),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1825),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1825),X1827)),X1825,bool))))) ) | ($spl0 | $spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f20525,f21143])).
fof(f30823,plain,(
  ( ! [X482,X481,X483] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X481),'type.semanticPrimitives.error_result'(sK1)),X482),s('type.semanticPrimitives.error_result'(X481),X483))))) | s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483))) = s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483)),X481,bool))))) ) | $spl44),
  inference(cnf_transformation,[],[f30823_D])).
fof(f30823_D,plain,(
  ( ! [X482,X481,X483] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X481),'type.semanticPrimitives.error_result'(sK1)),X482),s('type.semanticPrimitives.error_result'(X481),X483))))) | s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483))) = s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483)),X481,bool))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f30128,plain,(
  ( ! [X1822,X1820,X1818,X1821,X1819] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1820),'type.semanticPrimitives.error_result'(sK1)),X1821),s('type.semanticPrimitives.error_result'(X1820),X1822))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1818),s(bool,X1819))))) | s('type.semanticPrimitives.result'(bool,X1820),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1820),X1822))) = s('type.semanticPrimitives.result'(bool,X1820),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1820),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1820),X1822)),X1820,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f94,f21143])).
fof(f30863,plain,(
  ( ! [X1812,X1817,X1813] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1817,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1812),s(bool,X1813)))),sK6,sK5,sK2,sK1,sK0))))) ) | $spl108),
  inference(cnf_transformation,[],[f30863_D])).
fof(f30863_D,plain,(
  ( ! [X1812,X1817,X1813] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1817,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1812),s(bool,X1813)))),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f30127,plain,(
  ( ! [X1816,X1814,X1812,X1817,X1815,X1813] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1814),'type.semanticPrimitives.error_result'(sK1)),X1815),s('type.semanticPrimitives.error_result'(X1814),X1816))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1817,sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK1),'type.semanticPrimitives.result'(sK0,sK1)),'const.semanticPrimitives.Rerr_0'),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X1812),s(bool,X1813)))),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1814),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1814),X1816))) = s('type.semanticPrimitives.result'(bool,X1814),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1814),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1814),X1816)),X1814,bool))))) ) | ($spl0 | $spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f20361,f21143])).
fof(f30862,plain,(
  ( ! [X1808,X1807] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1807),s(bool,X1808))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl106),
  inference(cnf_transformation,[],[f30862_D])).
fof(f30862_D,plain,(
  ( ! [X1808,X1807] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1807),s(bool,X1808))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f30856,plain,(
  ( ! [X1771,X1773,X1772] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1771),'type.semanticPrimitives.result'(sK0,sK1)),X1772),s('type.semanticPrimitives.error_result'(X1771),X1773))) | s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773))) = s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773)),X1771,bool))))) ) | $spl94),
  inference(cnf_transformation,[],[f30856_D])).
fof(f30856_D,plain,(
  ( ! [X1771,X1773,X1772] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1771),'type.semanticPrimitives.result'(sK0,sK1)),X1772),s('type.semanticPrimitives.error_result'(X1771),X1773))) | s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773))) = s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773)),X1771,bool))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f30126,plain,(
  ( ! [X1810,X1808,X1811,X1809,X1807] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1809),'type.semanticPrimitives.result'(sK0,sK1)),X1810),s('type.semanticPrimitives.error_result'(X1809),X1811))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1807),s(bool,X1808))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1809),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1809),X1811))) = s('type.semanticPrimitives.result'(bool,X1809),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1809),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1809),X1811)),X1809,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f27612,f21143])).
fof(f30861,plain,(
  ( ! [X1806,X1802,X1801] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1806,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802))),s(fun(sK0,sK2),X1806),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl104),
  inference(cnf_transformation,[],[f30861_D])).
fof(f30861_D,plain,(
  ( ! [X1806,X1802,X1801] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1806,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802))),s(fun(sK0,sK2),X1806),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f30125,plain,(
  ( ! [X1806,X1804,X1802,X1805,X1803,X1801] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1803),'type.semanticPrimitives.result'(sK0,sK1)),X1804),s('type.semanticPrimitives.error_result'(X1803),X1805))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,X1806,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1801),s(bool,X1802))),s(fun(sK0,sK2),X1806),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1803),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1803),X1805))) = s('type.semanticPrimitives.result'(bool,X1803),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1803),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1803),X1805)),X1803,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f6737,f21143])).
fof(f30860,plain,(
  ( ! [X1795,X1800,X1796] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796)),sK3,X1800,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))),s(fun(sK0,sK2),X1800),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) | $spl102),
  inference(cnf_transformation,[],[f30860_D])).
fof(f30860_D,plain,(
  ( ! [X1795,X1800,X1796] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796)),sK3,X1800,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))),s(fun(sK0,sK2),X1800),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f30124,plain,(
  ( ! [X1795,X1800,X1798,X1796,X1799,X1797] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1797),'type.semanticPrimitives.result'(sK0,sK1)),X1798),s('type.semanticPrimitives.error_result'(X1797),X1799))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796)),sK3,X1800,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1795),s(bool,X1796))),s(fun(sK0,sK2),X1800),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1797),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1797),X1799))) = s('type.semanticPrimitives.result'(bool,X1797),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1797),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1797),X1799)),X1797,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f6537,f21143])).
fof(f30859,plain,(
  ( ! [X1789,X1794,X1790] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1794))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1794,sK6,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl100),
  inference(cnf_transformation,[],[f30859_D])).
fof(f30859_D,plain,(
  ( ! [X1789,X1794,X1790] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1794))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1794,sK6,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790)),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f30123,plain,(
  ( ! [X1793,X1791,X1789,X1794,X1792,X1790] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1791),'type.semanticPrimitives.result'(sK0,sK1)),X1792),s('type.semanticPrimitives.error_result'(X1791),X1793))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1794))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1794,sK6,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1789),s(bool,X1790)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1791),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1791),X1793))) = s('type.semanticPrimitives.result'(bool,X1791),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1791),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1791),X1793)),X1791,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f4658,f21143])).
fof(f30858,plain,(
  ( ! [X1784,X1788,X1783] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1788))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784)),sK3,sK6,X1788,sK2,sK1,sK0))))) ) | $spl98),
  inference(cnf_transformation,[],[f30858_D])).
fof(f30858_D,plain,(
  ( ! [X1784,X1788,X1783] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1788))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784)),sK3,sK6,X1788,sK2,sK1,sK0))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f30122,plain,(
  ( ! [X1787,X1784,X1786,X1788,X1783,X1785] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1785),'type.semanticPrimitives.result'(sK0,sK1)),X1786),s('type.semanticPrimitives.error_result'(X1785),X1787))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1788))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1783),s(bool,X1784)),sK3,sK6,X1788,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1785),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1785),X1787))) = s('type.semanticPrimitives.result'(bool,X1785),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1785),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1785),X1787)),X1785,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f4628,f21143])).
fof(f30857,plain,(
  ( ! [X1776,X1782,X1777,X1781] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777))),s(fun(sK0,sK2),X1782),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1781))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0))))) ) | $spl96),
  inference(cnf_transformation,[],[f30857_D])).
fof(f30857_D,plain,(
  ( ! [X1776,X1782,X1777,X1781] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777))),s(fun(sK0,sK2),X1782),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1781))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f30121,plain,(
  ( ! [X1776,X1778,X1780,X1782,X1777,X1779,X1781] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1778),'type.semanticPrimitives.result'(sK0,sK1)),X1779),s('type.semanticPrimitives.error_result'(X1778),X1780))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(X1781,X1782,sK3,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777)),sK6,sK5,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1776),s(bool,X1777))),s(fun(sK0,sK2),X1782),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1781))) | s('type.semanticPrimitives.result'(bool,X1778),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1778),X1780))) = s('type.semanticPrimitives.result'(bool,X1778),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1778),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1778),X1780)),X1778,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f646,f21143])).
fof(f30855,plain,(
  ( ! [X1769,X1775,X1770,X1774] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))),s(fun(sK0,sK2),X1774),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1775))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0))))) ) | $spl92),
  inference(cnf_transformation,[],[f30855_D])).
fof(f30855_D,plain,(
  ( ! [X1769,X1775,X1770,X1774] : (s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))),s(fun(sK0,sK2),X1774),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1775))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f30120,plain,(
  ( ! [X1769,X1771,X1773,X1775,X1770,X1772,X1774] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun('type.semanticPrimitives.error_result'(X1771),'type.semanticPrimitives.result'(sK0,sK1)),X1772),s('type.semanticPrimitives.error_result'(X1771),X1773))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770)),sK3,X1774,X1775,sK2,sK1,sK0)))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),happ(s(fun(bool,'type.semanticPrimitives.result'(sK0,sK1)),X1769),s(bool,X1770))),s(fun(sK0,sK2),X1774),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),X1775))) | s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773))) = s('type.semanticPrimitives.result'(bool,X1771),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1771),X1773)),X1771,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f645,f21143])).
fof(f30854,plain,(
  ( ! [X1765,X1764] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1764),s(bool,X1765)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl90),
  inference(cnf_transformation,[],[f30854_D])).
fof(f30854_D,plain,(
  ( ! [X1765,X1764] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1764),s(bool,X1765)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f30853,plain,(
  ( ! [X1761,X1763,X1762] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1761),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1762),s('type.semanticPrimitives.error_result'(X1761),X1763))))) | s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763))) = s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763)),X1761,bool))))) ) | $spl88),
  inference(cnf_transformation,[],[f30853_D])).
fof(f30853_D,plain,(
  ( ! [X1761,X1763,X1762] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1761),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1762),s('type.semanticPrimitives.error_result'(X1761),X1763))))) | s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763))) = s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763)),X1761,bool))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f30119,plain,(
  ( ! [X1765,X1767,X1764,X1766,X1768] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1766),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1767),s('type.semanticPrimitives.error_result'(X1766),X1768))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1764),s(bool,X1765)),sK6,sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1766),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1766),X1768))) = s('type.semanticPrimitives.result'(bool,X1766),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1766),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1766),X1768)),X1766,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20703,f21143])).
fof(f30852,plain,(
  ( ! [X1759,X1760] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1759),s(bool,X1760)),sK2,sK1,sK0))))) ) | $spl86),
  inference(cnf_transformation,[],[f30852_D])).
fof(f30852_D,plain,(
  ( ! [X1759,X1760] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1759),s(bool,X1760)),sK2,sK1,sK0))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f30118,plain,(
  ( ! [X1761,X1763,X1759,X1760,X1762] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1761),fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1762),s('type.semanticPrimitives.error_result'(X1761),X1763))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK10(sK5,sK6,sK3,sK3,sK6,happ(s(fun(bool,fun('type.semanticPrimitives.error_result'(sK1),sK2)),X1759),s(bool,X1760)),sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763))) = s('type.semanticPrimitives.result'(bool,X1761),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1761),X1763)),X1761,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20518,f21143])).
fof(f30851,plain,(
  ( ! [X1716,X1717] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun(bool,fun(sK0,sK2)),X1716),s(bool,X1717)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) | $spl84),
  inference(cnf_transformation,[],[f30851_D])).
fof(f30851_D,plain,(
  ( ! [X1716,X1717] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun(bool,fun(sK0,sK2)),X1716),s(bool,X1717)),sK3,sK3,sK6,sK5,sK2,sK1,sK0))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f30850,plain,(
  ( ! [X1714,X1713,X1715] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1713),fun(sK0,sK2)),X1714),s('type.semanticPrimitives.error_result'(X1713),X1715))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715))) = s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715)),X1713,bool))))) ) | $spl82),
  inference(cnf_transformation,[],[f30850_D])).
fof(f30850_D,plain,(
  ( ! [X1714,X1713,X1715] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1713),fun(sK0,sK2)),X1714),s('type.semanticPrimitives.error_result'(X1713),X1715))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715))) = s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715)),X1713,bool))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f30111,plain,(
  ( ! [X1716,X1718,X1720,X1717,X1719] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1718),fun(sK0,sK2)),X1719),s('type.semanticPrimitives.error_result'(X1718),X1720))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,happ(s(fun(bool,fun(sK0,sK2)),X1716),s(bool,X1717)),sK3,sK3,sK6,sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1718),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1718),X1720))) = s('type.semanticPrimitives.result'(bool,X1718),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1718),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1718),X1720)),X1718,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f20358,f21143])).
fof(f30849,plain,(
  ( ! [X1711,X1712] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun(bool,fun(sK0,sK2)),X1711),s(bool,X1712)),sK5,sK2,sK1,sK0))))) ) | $spl80),
  inference(cnf_transformation,[],[f30849_D])).
fof(f30849_D,plain,(
  ( ! [X1711,X1712] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun(bool,fun(sK0,sK2)),X1711),s(bool,X1712)),sK5,sK2,sK1,sK0))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f30110,plain,(
  ( ! [X1711,X1712,X1714,X1713,X1715] : (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),happ(s(fun('type.semanticPrimitives.error_result'(X1713),fun(sK0,sK2)),X1714),s('type.semanticPrimitives.error_result'(X1713),X1715))),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK9(sK5,sK6,sK3,sK3,happ(s(fun(bool,fun(sK0,sK2)),X1711),s(bool,X1712)),sK5,sK2,sK1,sK0)))) | s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715))) = s('type.semanticPrimitives.result'(bool,X1713),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1713),X1715)),X1713,bool))))) ) | ($spl0 | $spl16 | $spl28)),
  inference(superposition,[],[f19170,f21143])).
fof(f30846,plain,(
  ( ! [X1663,X1658,X1657] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1657),s(bool,X1658))),s(bool,X1663)))))) ) | $spl74),
  inference(cnf_transformation,[],[f30846_D])).
fof(f30846_D,plain,(
  ( ! [X1663,X1658,X1657] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1657),s(bool,X1658))),s(bool,X1663)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f30848,plain,(
  ( ! [X1666,X1668,X1667,X1669] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1666),fun(bool,sK0)),X1667),s('type.semanticPrimitives.error_result'(X1666),X1668))),s(bool,X1669))))) | s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668))) = s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668)),X1666,bool))))) ) | $spl78),
  inference(cnf_transformation,[],[f30848_D])).
fof(f30848_D,plain,(
  ( ! [X1666,X1668,X1667,X1669] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1666),fun(bool,sK0)),X1667),s('type.semanticPrimitives.error_result'(X1666),X1668))),s(bool,X1669))))) | s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668))) = s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668)),X1666,bool))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f30103,plain,(
  ( ! [X1664,X1666,X1668,X1670,X1665,X1667,X1669] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1666),fun(bool,sK0)),X1667),s('type.semanticPrimitives.error_result'(X1666),X1668))),s(bool,X1669))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1664),s(bool,X1665))),s(bool,X1670))))) | s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668))) = s('type.semanticPrimitives.result'(bool,X1666),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1666),X1668)),X1666,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f2363,f21143])).
fof(f30847,plain,(
  ( ! [X1660,X1662,X1659,X1661] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1659),fun(bool,sK0)),X1660),s('type.semanticPrimitives.error_result'(X1659),X1661))))),s(bool,X1662))) | s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661))) = s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661)),X1659,bool))))) ) | $spl76),
  inference(cnf_transformation,[],[f30847_D])).
fof(f30847_D,plain,(
  ( ! [X1660,X1662,X1659,X1661] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1659),fun(bool,sK0)),X1660),s('type.semanticPrimitives.error_result'(X1659),X1661))))),s(bool,X1662))) | s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661))) = s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661)),X1659,bool))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f30102,plain,(
  ( ! [X1663,X1658,X1660,X1662,X1657,X1659,X1661] : (s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1_2'(s(fun(sK0,sK2),sK6),s(fun(bool,sK0),happ(s(fun('type.semanticPrimitives.error_result'(X1659),fun(bool,sK0)),X1660),s('type.semanticPrimitives.error_result'(X1659),X1661))))),s(bool,X1662))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),happ(s(fun(bool,fun(bool,sK0)),X1657),s(bool,X1658))),s(bool,X1663))))) | s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661))) = s('type.semanticPrimitives.result'(bool,X1659),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1659),X1661)),X1659,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f3828,f21143])).
fof(f30844,plain,(
  ( ! [X1650,X1656,X1651] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1650),s(bool,X1651))),s(bool,X1656)))))) ) | $spl70),
  inference(cnf_transformation,[],[f30844_D])).
fof(f30844_D,plain,(
  ( ! [X1650,X1656,X1651] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1650),s(bool,X1651))),s(bool,X1656)))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f30845,plain,(
  ( ! [X1652,X1654,X1653,X1655] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1652),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1653),s('type.semanticPrimitives.error_result'(X1652),X1654))),s(bool,X1655))))) | s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654))) = s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654)),X1652,bool))))) ) | $spl72),
  inference(cnf_transformation,[],[f30845_D])).
fof(f30845_D,plain,(
  ( ! [X1652,X1654,X1653,X1655] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1652),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1653),s('type.semanticPrimitives.error_result'(X1652),X1654))),s(bool,X1655))))) | s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654))) = s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654)),X1652,bool))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f30101,plain,(
  ( ! [X1650,X1652,X1654,X1656,X1651,X1653,X1655] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun('type.semanticPrimitives.error_result'(X1652),fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1653),s('type.semanticPrimitives.error_result'(X1652),X1654))),s(bool,X1655))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),happ(s(fun(bool,fun(bool,'type.semanticPrimitives.error_result'(sK1))),X1650),s(bool,X1651))),s(bool,X1656))))) | s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654))) = s('type.semanticPrimitives.result'(bool,X1652),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1652),X1654)),X1652,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f2362,f21143])).
fof(f30843,plain,(
  ( ! [X633,X632,X634] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(X632),sK2),X633),s('type.semanticPrimitives.error_result'(X632),X634))) | s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634))) = s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634)),X632,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f30843_D])).
fof(f30843_D,plain,(
  ( ! [X633,X632,X634] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(X632),sK2),X633),s('type.semanticPrimitives.error_result'(X632),X634))) | s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634))) = s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634)),X632,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f29983,plain,(
  ( ! [X633,X635,X632,X634,X630] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(X632),sK2),X633),s('type.semanticPrimitives.error_result'(X632),X634))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X630),s(bool,X635))))) | s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634))) = s('type.semanticPrimitives.result'(bool,X632),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X632),X634)),X632,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f3828,f21143])).
fof(f30842,plain,(
  ( ! [X625,X621,X623,X624,X626,X622] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X622,sK0),X623),s(X622,happ(s(fun(X621,X622),X624),s(X621,happ(s(fun(bool,X621),X625),s(bool,X626)))))))))) ) | $spl66),
  inference(cnf_transformation,[],[f30842_D])).
fof(f30842_D,plain,(
  ( ! [X625,X621,X623,X624,X626,X622] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X622,sK0),X623),s(X622,happ(s(fun(X621,X622),X624),s(X621,happ(s(fun(bool,X621),X625),s(bool,X626)))))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f29982,plain,(
  ( ! [X629,X625,X627,X621,X623,X628,X624,X626,X622] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X627),sK0),X628),s('type.semanticPrimitives.error_result'(X627),X629))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X622,sK0),X623),s(X622,happ(s(fun(X621,X622),X624),s(X621,happ(s(fun(bool,X621),X625),s(bool,X626))))))))) | s('type.semanticPrimitives.result'(bool,X627),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X627),X629))) = s('type.semanticPrimitives.result'(bool,X627),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X627),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X627),X629)),X627,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1105,f21143])).
fof(f30839,plain,(
  ( ! [X596,X592,X594,X597,X593,X595,X591] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X594),s(fun(bool,sK0),X591),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X592))))) | s('type.semanticPrimitives.result'(bool,X597),X593) != s('type.semanticPrimitives.result'(bool,X597),X594) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X593),s(fun(bool,sK0),X595),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X596))))) | s('type.semanticPrimitives.result'(bool,X597),X593) = s('type.semanticPrimitives.result'(bool,X597),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X597),sK10(X592,X591,X593,X594,X595,X596,sK0,X597,bool))))) ) | $spl64),
  inference(cnf_transformation,[],[f30839_D])).
fof(f30839_D,plain,(
  ( ! [X596,X592,X594,X597,X593,X595,X591] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X594),s(fun(bool,sK0),X591),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X592))))) | s('type.semanticPrimitives.result'(bool,X597),X593) != s('type.semanticPrimitives.result'(bool,X597),X594) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X593),s(fun(bool,sK0),X595),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X596))))) | s('type.semanticPrimitives.result'(bool,X597),X593) = s('type.semanticPrimitives.result'(bool,X597),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X597),sK10(X592,X591,X593,X594,X595,X596,sK0,X597,bool))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f30841,plain,(
  ( ! [X617,X613,X615,X611,X620,X616,X618,X612,X614] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X617),X615),s(fun(bool,sK0),X613),s(fun('type.semanticPrimitives.error_result'(X617),sK0),X612))))) | s('type.semanticPrimitives.result'(bool,X617),X614) = s('type.semanticPrimitives.result'(bool,X617),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X617),sK10(X612,X613,X614,X615,X611,X616,sK0,X617,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X617),X614),s(fun(bool,sK0),X611),s(fun('type.semanticPrimitives.error_result'(X617),sK0),X616))))) | s('type.semanticPrimitives.result'(bool,X617),X614) != s('type.semanticPrimitives.result'(bool,X617),X615) | s('type.semanticPrimitives.result'(bool,X618),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),X620))) = s('type.semanticPrimitives.result'(bool,X618),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),X620)),X618,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29981,f21143])).
fof(f29981,plain,(
  ( ! [X617,X619,X613,X615,X611,X620,X616,X618,X612,X614] : (s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X618),sK0),X619),s('type.semanticPrimitives.error_result'(X618),X620))) != s(sK0,happ(s(fun(bool,sK0),X613),s(bool,sK9(X612,X613,X614,X615,X611,X616,sK0,X617,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X617),X615),s(fun(bool,sK0),X613),s(fun('type.semanticPrimitives.error_result'(X617),sK0),X612))))) | s('type.semanticPrimitives.result'(bool,X617),X614) = s('type.semanticPrimitives.result'(bool,X617),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X617),sK10(X612,X613,X614,X615,X611,X616,sK0,X617,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X617),X614),s(fun(bool,sK0),X611),s(fun('type.semanticPrimitives.error_result'(X617),sK0),X616))))) | s('type.semanticPrimitives.result'(bool,X617),X614) != s('type.semanticPrimitives.result'(bool,X617),X615) | s('type.semanticPrimitives.result'(bool,X618),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),X620))) = s('type.semanticPrimitives.result'(bool,X618),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X618),X620)),X618,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1010,f21143])).
fof(f1010,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(sK0,happ(s(fun(X16,sK0),X19),s(X16,sK9(X20,X19,X21,X18,X22,X23,sK0,X17,X16)))) != s(sK0,happ(s(fun(X16,sK0),X22),s(X16,sK9(X20,X19,X21,X18,X22,X23,sK0,X17,X16)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X16,X17),X18),s(fun(X16,sK0),X19),s(fun('type.semanticPrimitives.error_result'(X17),sK0),X20))))) | s('type.semanticPrimitives.result'(X16,X17),X21) = s('type.semanticPrimitives.result'(X16,X17),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X17),sK10(X20,X19,X21,X18,X22,X23,sK0,X17,X16)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X16,X17),X21),s(fun(X16,sK0),X22),s(fun('type.semanticPrimitives.error_result'(X17),sK0),X23))))) | s('type.semanticPrimitives.result'(X16,X17),X18) != s('type.semanticPrimitives.result'(X16,X17),X21)) ) | $spl4),
  inference(superposition,[],[f95,f77])).
fof(f30837,plain,(
  ( ! [X584,X586,X582,X585,X587,X581,X583] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X583),s(fun(bool,sK0),X585),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X586))))) | s('type.semanticPrimitives.result'(bool,X587),X583) != s('type.semanticPrimitives.result'(bool,X587),X584) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X584),s(fun(bool,sK0),X581),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X582))))) | s('type.semanticPrimitives.result'(bool,X587),X583) = s('type.semanticPrimitives.result'(bool,X587),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X587),sK10(X582,X581,X583,X584,X585,X586,sK0,X587,bool))))) ) | $spl62),
  inference(cnf_transformation,[],[f30837_D])).
fof(f30837_D,plain,(
  ( ! [X584,X586,X582,X585,X587,X581,X583] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X583),s(fun(bool,sK0),X585),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X586))))) | s('type.semanticPrimitives.result'(bool,X587),X583) != s('type.semanticPrimitives.result'(bool,X587),X584) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X584),s(fun(bool,sK0),X581),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X582))))) | s('type.semanticPrimitives.result'(bool,X587),X583) = s('type.semanticPrimitives.result'(bool,X587),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X587),sK10(X582,X581,X583,X584,X585,X586,sK0,X587,bool))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f30840,plain,(
  ( ! [X606,X610,X602,X605,X607,X601,X603,X608,X604] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X607),X604),s(fun(bool,sK0),X601),s(fun('type.semanticPrimitives.error_result'(X607),sK0),X606))))) | s('type.semanticPrimitives.result'(bool,X607),X604) = s('type.semanticPrimitives.result'(bool,X607),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X607),sK10(X602,X603,X604,X605,X601,X606,sK0,X607,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X607),X605),s(fun(bool,sK0),X603),s(fun('type.semanticPrimitives.error_result'(X607),sK0),X602))))) | s('type.semanticPrimitives.result'(bool,X607),X604) != s('type.semanticPrimitives.result'(bool,X607),X605) | s('type.semanticPrimitives.result'(bool,X608),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),X610))) = s('type.semanticPrimitives.result'(bool,X608),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),X610)),X608,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29980,f21143])).
fof(f29980,plain,(
  ( ! [X604,X606,X602,X609,X605,X607,X601,X603,X608,X610] : (s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X608),sK0),X609),s('type.semanticPrimitives.error_result'(X608),X610))) != s(sK0,happ(s(fun(bool,sK0),X603),s(bool,sK9(X602,X603,X604,X605,X601,X606,sK0,X607,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X607),X604),s(fun(bool,sK0),X601),s(fun('type.semanticPrimitives.error_result'(X607),sK0),X606))))) | s('type.semanticPrimitives.result'(bool,X607),X604) = s('type.semanticPrimitives.result'(bool,X607),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X607),sK10(X602,X603,X604,X605,X601,X606,sK0,X607,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X607),X605),s(fun(bool,sK0),X603),s(fun('type.semanticPrimitives.error_result'(X607),sK0),X602))))) | s('type.semanticPrimitives.result'(bool,X607),X604) != s('type.semanticPrimitives.result'(bool,X607),X605) | s('type.semanticPrimitives.result'(bool,X608),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),X610))) = s('type.semanticPrimitives.result'(bool,X608),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X608),X610)),X608,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1011,f21143])).
fof(f1011,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(sK0,happ(s(fun(X24,sK0),X27),s(X24,sK9(X31,X30,X26,X29,X27,X28,sK0,X25,X24)))) != s(sK0,happ(s(fun(X24,sK0),X30),s(X24,sK9(X31,X30,X26,X29,X27,X28,sK0,X25,X24)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X24,X25),X26),s(fun(X24,sK0),X27),s(fun('type.semanticPrimitives.error_result'(X25),sK0),X28))))) | s('type.semanticPrimitives.result'(X24,X25),X26) = s('type.semanticPrimitives.result'(X24,X25),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X25),sK10(X31,X30,X26,X29,X27,X28,sK0,X25,X24)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X24,X25),X29),s(fun(X24,sK0),X30),s(fun('type.semanticPrimitives.error_result'(X25),sK0),X31))))) | s('type.semanticPrimitives.result'(X24,X25),X26) != s('type.semanticPrimitives.result'(X24,X25),X29)) ) | $spl4),
  inference(superposition,[],[f95,f77])).
fof(f30838,plain,(
  ( ! [X600,X596,X598,X592,X594,X597,X593,X595,X591] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X594),s(fun(bool,sK0),X591),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X592))))) | s('type.semanticPrimitives.result'(bool,X597),X593) = s('type.semanticPrimitives.result'(bool,X597),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X597),sK10(X592,X591,X593,X594,X595,X596,sK0,X597,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X593),s(fun(bool,sK0),X595),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X596))))) | s('type.semanticPrimitives.result'(bool,X597),X593) != s('type.semanticPrimitives.result'(bool,X597),X594) | s('type.semanticPrimitives.result'(bool,X598),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),X600))) = s('type.semanticPrimitives.result'(bool,X598),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),X600)),X598,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29979,f21143])).
fof(f29979,plain,(
  ( ! [X600,X596,X598,X592,X594,X597,X599,X593,X595,X591] : (s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X598),sK0),X599),s('type.semanticPrimitives.error_result'(X598),X600))) != s(sK0,happ(s(fun(bool,sK0),X595),s(bool,sK9(X592,X591,X593,X594,X595,X596,sK0,X597,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X594),s(fun(bool,sK0),X591),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X592))))) | s('type.semanticPrimitives.result'(bool,X597),X593) = s('type.semanticPrimitives.result'(bool,X597),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X597),sK10(X592,X591,X593,X594,X595,X596,sK0,X597,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X597),X593),s(fun(bool,sK0),X595),s(fun('type.semanticPrimitives.error_result'(X597),sK0),X596))))) | s('type.semanticPrimitives.result'(bool,X597),X593) != s('type.semanticPrimitives.result'(bool,X597),X594) | s('type.semanticPrimitives.result'(bool,X598),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),X600))) = s('type.semanticPrimitives.result'(bool,X598),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X598),X600)),X598,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1010,f21143])).
fof(f30836,plain,(
  ( ! [X588,X590,X584,X586,X582,X585,X587,X581,X583] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X583),s(fun(bool,sK0),X585),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X586))))) | s('type.semanticPrimitives.result'(bool,X587),X583) = s('type.semanticPrimitives.result'(bool,X587),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X587),sK10(X582,X581,X583,X584,X585,X586,sK0,X587,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X584),s(fun(bool,sK0),X581),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X582))))) | s('type.semanticPrimitives.result'(bool,X587),X583) != s('type.semanticPrimitives.result'(bool,X587),X584) | s('type.semanticPrimitives.result'(bool,X588),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),X590))) = s('type.semanticPrimitives.result'(bool,X588),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),X590)),X588,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29978,f21143])).
fof(f29978,plain,(
  ( ! [X588,X590,X584,X586,X582,X589,X585,X587,X581,X583] : (s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X588),sK0),X589),s('type.semanticPrimitives.error_result'(X588),X590))) != s(sK0,happ(s(fun(bool,sK0),X585),s(bool,sK9(X582,X581,X583,X584,X585,X586,sK0,X587,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X583),s(fun(bool,sK0),X585),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X586))))) | s('type.semanticPrimitives.result'(bool,X587),X583) = s('type.semanticPrimitives.result'(bool,X587),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X587),sK10(X582,X581,X583,X584,X585,X586,sK0,X587,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X587),X584),s(fun(bool,sK0),X581),s(fun('type.semanticPrimitives.error_result'(X587),sK0),X582))))) | s('type.semanticPrimitives.result'(bool,X587),X583) != s('type.semanticPrimitives.result'(bool,X587),X584) | s('type.semanticPrimitives.result'(bool,X588),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),X590))) = s('type.semanticPrimitives.result'(bool,X588),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X588),X590)),X588,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1011,f21143])).
fof(f30835,plain,(
  ( ! [X561,X559,X566,X560] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X559,sK0),X560),s(X559,happ(s(fun(bool,X559),X561),s(bool,X566)))))))) ) | $spl60),
  inference(cnf_transformation,[],[f30835_D])).
fof(f30835_D,plain,(
  ( ! [X561,X559,X566,X560] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X559,sK0),X560),s(X559,happ(s(fun(bool,X559),X561),s(bool,X566)))))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f29977,plain,(
  ( ! [X575,X580,X576,X578,X574,X577,X579] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X578),sK0),X579),s('type.semanticPrimitives.error_result'(X578),X580))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X574,sK0),X575),s(X574,happ(s(fun(bool,X574),X576),s(bool,X577))))))) | s('type.semanticPrimitives.result'(bool,X578),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X578),X580))) = s('type.semanticPrimitives.result'(bool,X578),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X578),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X578),X580)),X578,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1020,f21143])).
fof(f30833,plain,(
  ( ! [X553,X551,X558,X552] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X551,sK0),X552),s(X551,happ(s(fun(bool,X551),X553),s(bool,X558)))))))) ) | $spl56),
  inference(cnf_transformation,[],[f30833_D])).
fof(f30833_D,plain,(
  ( ! [X553,X551,X558,X552] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X551,sK0),X552),s(X551,happ(s(fun(bool,X551),X553),s(bool,X558)))))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f29976,plain,(
  ( ! [X573,X569,X571,X567,X572,X568,X570] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X571),sK0),X572),s('type.semanticPrimitives.error_result'(X571),X573))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X567,sK0),X568),s(X567,happ(s(fun(bool,X567),X569),s(bool,X570))))))) | s('type.semanticPrimitives.result'(bool,X571),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X571),X573))) = s('type.semanticPrimitives.result'(bool,X571),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X571),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X571),X573)),X571,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1021,f21143])).
fof(f29975,plain,(
  ( ! [X565,X561,X563,X559,X564,X566,X560] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X563),sK0),X564),s('type.semanticPrimitives.error_result'(X563),X565))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X559,sK0),X560),s(X559,happ(s(fun(bool,X559),X561),s(bool,X566))))))) | s('type.semanticPrimitives.result'(bool,X563),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X563),X565))) = s('type.semanticPrimitives.result'(bool,X563),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X563),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X563),X565)),X563,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1116,f21143])).
fof(f1116,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153)))))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1115,f819])).
fof(f1115,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1089,f819])).
fof(f1089,plain,(
  ( ! [X156,X154,X157,X155,X153] : (~p(s(bool,X154)) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153)) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154)))))) ) | $spl4),
  inference(superposition,[],[f1020,f90])).
fof(f29974,plain,(
  ( ! [X557,X553,X555,X551,X556,X558,X552] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X555),sK0),X556),s('type.semanticPrimitives.error_result'(X555),X557))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X551,sK0),X552),s(X551,happ(s(fun(bool,X551),X553),s(bool,X558))))))) | s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557))) = s('type.semanticPrimitives.result'(bool,X555),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X555),X557)),X555,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f1413,f21143])).
fof(f1413,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153)))))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1412,f819])).
fof(f1412,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153))) ) | ($spl4 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1383,f819])).
fof(f1383,plain,(
  ( ! [X156,X154,X157,X155,X153] : (~p(s(bool,X154)) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun(X155,sK0),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153)) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(X155,sK0),X156),s(fun(bool,X155),X157))),s(bool,X154)))))) ) | $spl4),
  inference(superposition,[],[f1021,f90])).
fof(f29973,plain,(
  ( ! [X549,X545,X547,X548,X550] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,happ(s(fun('type.semanticPrimitives.error_result'(X547),sK0),X548),s('type.semanticPrimitives.error_result'(X547),X549))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(bool,sK0),X545),s(bool,X550))))) | s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549))) = s('type.semanticPrimitives.result'(bool,X547),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X547),X549)),X547,bool))))) ) | ($spl4 | $spl16 | $spl28)),
  inference(superposition,[],[f2363,f21143])).
fof(f30828,plain,(
  ( ! [X520,X516,X518,X521,X517,X519,X515] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X518),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X515),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X516))))) | s('type.semanticPrimitives.result'(bool,X521),X517) != s('type.semanticPrimitives.result'(bool,X521),X518) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X517),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X519),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X520))))) | s('type.semanticPrimitives.result'(bool,X521),X517) = s('type.semanticPrimitives.result'(bool,X521),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X521),sK10(X516,X515,X517,X518,X519,X520,'type.semanticPrimitives.error_result'(sK1),X521,bool))))) ) | $spl50),
  inference(cnf_transformation,[],[f30828_D])).
fof(f30828_D,plain,(
  ( ! [X520,X516,X518,X521,X517,X519,X515] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X518),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X515),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X516))))) | s('type.semanticPrimitives.result'(bool,X521),X517) != s('type.semanticPrimitives.result'(bool,X521),X518) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X517),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X519),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X520))))) | s('type.semanticPrimitives.result'(bool,X521),X517) = s('type.semanticPrimitives.result'(bool,X521),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X521),sK10(X516,X515,X517,X518,X519,X520,'type.semanticPrimitives.error_result'(sK1),X521,bool))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f30830,plain,(
  ( ! [X540,X542,X536,X538,X541,X537,X539,X535,X544] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X541),X539),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X537),s(fun('type.semanticPrimitives.error_result'(X541),'type.semanticPrimitives.error_result'(sK1)),X536))))) | s('type.semanticPrimitives.result'(bool,X541),X538) = s('type.semanticPrimitives.result'(bool,X541),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X541),sK10(X536,X537,X538,X539,X535,X540,'type.semanticPrimitives.error_result'(sK1),X541,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X541),X538),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X535),s(fun('type.semanticPrimitives.error_result'(X541),'type.semanticPrimitives.error_result'(sK1)),X540))))) | s('type.semanticPrimitives.result'(bool,X541),X538) != s('type.semanticPrimitives.result'(bool,X541),X539) | s('type.semanticPrimitives.result'(bool,X542),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),X544))) = s('type.semanticPrimitives.result'(bool,X542),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),X544)),X542,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29972,f21143])).
fof(f29972,plain,(
  ( ! [X540,X542,X536,X538,X541,X543,X537,X539,X535,X544] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X542),'type.semanticPrimitives.error_result'(sK1)),X543),s('type.semanticPrimitives.error_result'(X542),X544))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X537),s(bool,sK9(X536,X537,X538,X539,X535,X540,'type.semanticPrimitives.error_result'(sK1),X541,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X541),X539),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X537),s(fun('type.semanticPrimitives.error_result'(X541),'type.semanticPrimitives.error_result'(sK1)),X536))))) | s('type.semanticPrimitives.result'(bool,X541),X538) = s('type.semanticPrimitives.result'(bool,X541),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X541),sK10(X536,X537,X538,X539,X535,X540,'type.semanticPrimitives.error_result'(sK1),X541,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X541),X538),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X535),s(fun('type.semanticPrimitives.error_result'(X541),'type.semanticPrimitives.error_result'(sK1)),X540))))) | s('type.semanticPrimitives.result'(bool,X541),X538) != s('type.semanticPrimitives.result'(bool,X541),X539) | s('type.semanticPrimitives.result'(bool,X542),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),X544))) = s('type.semanticPrimitives.result'(bool,X542),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X542),X544)),X542,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1028,f21143])).
fof(f1028,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X16,'type.semanticPrimitives.error_result'(sK1)),X19),s(X16,sK9(X20,X19,X21,X18,X22,X23,'type.semanticPrimitives.error_result'(sK1),X17,X16)))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X16,'type.semanticPrimitives.error_result'(sK1)),X22),s(X16,sK9(X20,X19,X21,X18,X22,X23,'type.semanticPrimitives.error_result'(sK1),X17,X16)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X16,X17),X18),s(fun(X16,'type.semanticPrimitives.error_result'(sK1)),X19),s(fun('type.semanticPrimitives.error_result'(X17),'type.semanticPrimitives.error_result'(sK1)),X20))))) | s('type.semanticPrimitives.result'(X16,X17),X21) = s('type.semanticPrimitives.result'(X16,X17),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X17),sK10(X20,X19,X21,X18,X22,X23,'type.semanticPrimitives.error_result'(sK1),X17,X16)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X16,X17),X21),s(fun(X16,'type.semanticPrimitives.error_result'(sK1)),X22),s(fun('type.semanticPrimitives.error_result'(X17),'type.semanticPrimitives.error_result'(sK1)),X23))))) | s('type.semanticPrimitives.result'(X16,X17),X18) != s('type.semanticPrimitives.result'(X16,X17),X21)) ) | $spl2),
  inference(superposition,[],[f94,f77])).
fof(f30826,plain,(
  ( ! [X509,X511,X505,X507,X508,X510,X506] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X507),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X509),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X510))))) | s('type.semanticPrimitives.result'(bool,X511),X507) != s('type.semanticPrimitives.result'(bool,X511),X508) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X508),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X505),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X506))))) | s('type.semanticPrimitives.result'(bool,X511),X507) = s('type.semanticPrimitives.result'(bool,X511),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X511),sK10(X506,X505,X507,X508,X509,X510,'type.semanticPrimitives.error_result'(sK1),X511,bool))))) ) | $spl48),
  inference(cnf_transformation,[],[f30826_D])).
fof(f30826_D,plain,(
  ( ! [X509,X511,X505,X507,X508,X510,X506] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X507),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X509),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X510))))) | s('type.semanticPrimitives.result'(bool,X511),X507) != s('type.semanticPrimitives.result'(bool,X511),X508) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X508),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X505),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X506))))) | s('type.semanticPrimitives.result'(bool,X511),X507) = s('type.semanticPrimitives.result'(bool,X511),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X511),sK10(X506,X505,X507,X508,X509,X510,'type.semanticPrimitives.error_result'(sK1),X511,bool))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f30829,plain,(
  ( ! [X532,X534,X528,X530,X526,X529,X531,X525,X527] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X531),X528),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X525),s(fun('type.semanticPrimitives.error_result'(X531),'type.semanticPrimitives.error_result'(sK1)),X530))))) | s('type.semanticPrimitives.result'(bool,X531),X528) = s('type.semanticPrimitives.result'(bool,X531),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X531),sK10(X526,X527,X528,X529,X525,X530,'type.semanticPrimitives.error_result'(sK1),X531,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X531),X529),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X527),s(fun('type.semanticPrimitives.error_result'(X531),'type.semanticPrimitives.error_result'(sK1)),X526))))) | s('type.semanticPrimitives.result'(bool,X531),X528) != s('type.semanticPrimitives.result'(bool,X531),X529) | s('type.semanticPrimitives.result'(bool,X532),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),X534))) = s('type.semanticPrimitives.result'(bool,X532),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),X534)),X532,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29971,f21143])).
fof(f29971,plain,(
  ( ! [X532,X534,X528,X530,X526,X533,X529,X531,X525,X527] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X532),'type.semanticPrimitives.error_result'(sK1)),X533),s('type.semanticPrimitives.error_result'(X532),X534))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X527),s(bool,sK9(X526,X527,X528,X529,X525,X530,'type.semanticPrimitives.error_result'(sK1),X531,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X531),X528),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X525),s(fun('type.semanticPrimitives.error_result'(X531),'type.semanticPrimitives.error_result'(sK1)),X530))))) | s('type.semanticPrimitives.result'(bool,X531),X528) = s('type.semanticPrimitives.result'(bool,X531),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X531),sK10(X526,X527,X528,X529,X525,X530,'type.semanticPrimitives.error_result'(sK1),X531,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X531),X529),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X527),s(fun('type.semanticPrimitives.error_result'(X531),'type.semanticPrimitives.error_result'(sK1)),X526))))) | s('type.semanticPrimitives.result'(bool,X531),X528) != s('type.semanticPrimitives.result'(bool,X531),X529) | s('type.semanticPrimitives.result'(bool,X532),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),X534))) = s('type.semanticPrimitives.result'(bool,X532),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X532),X534)),X532,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1029,f21143])).
fof(f1029,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X24,'type.semanticPrimitives.error_result'(sK1)),X27),s(X24,sK9(X31,X30,X26,X29,X27,X28,'type.semanticPrimitives.error_result'(sK1),X25,X24)))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X24,'type.semanticPrimitives.error_result'(sK1)),X30),s(X24,sK9(X31,X30,X26,X29,X27,X28,'type.semanticPrimitives.error_result'(sK1),X25,X24)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X24,X25),X26),s(fun(X24,'type.semanticPrimitives.error_result'(sK1)),X27),s(fun('type.semanticPrimitives.error_result'(X25),'type.semanticPrimitives.error_result'(sK1)),X28))))) | s('type.semanticPrimitives.result'(X24,X25),X26) = s('type.semanticPrimitives.result'(X24,X25),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X25),sK10(X31,X30,X26,X29,X27,X28,'type.semanticPrimitives.error_result'(sK1),X25,X24)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X24,X25),X29),s(fun(X24,'type.semanticPrimitives.error_result'(sK1)),X30),s(fun('type.semanticPrimitives.error_result'(X25),'type.semanticPrimitives.error_result'(sK1)),X31))))) | s('type.semanticPrimitives.result'(X24,X25),X26) != s('type.semanticPrimitives.result'(X24,X25),X29)) ) | $spl2),
  inference(superposition,[],[f94,f77])).
fof(f30827,plain,(
  ( ! [X524,X520,X522,X516,X518,X521,X517,X519,X515] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X518),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X515),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X516))))) | s('type.semanticPrimitives.result'(bool,X521),X517) = s('type.semanticPrimitives.result'(bool,X521),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X521),sK10(X516,X515,X517,X518,X519,X520,'type.semanticPrimitives.error_result'(sK1),X521,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X517),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X519),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X520))))) | s('type.semanticPrimitives.result'(bool,X521),X517) != s('type.semanticPrimitives.result'(bool,X521),X518) | s('type.semanticPrimitives.result'(bool,X522),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),X524))) = s('type.semanticPrimitives.result'(bool,X522),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),X524)),X522,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29970,f21143])).
fof(f29970,plain,(
  ( ! [X524,X520,X522,X516,X518,X521,X523,X517,X519,X515] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X522),'type.semanticPrimitives.error_result'(sK1)),X523),s('type.semanticPrimitives.error_result'(X522),X524))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X519),s(bool,sK9(X516,X515,X517,X518,X519,X520,'type.semanticPrimitives.error_result'(sK1),X521,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X518),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X515),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X516))))) | s('type.semanticPrimitives.result'(bool,X521),X517) = s('type.semanticPrimitives.result'(bool,X521),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X521),sK10(X516,X515,X517,X518,X519,X520,'type.semanticPrimitives.error_result'(sK1),X521,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X521),X517),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X519),s(fun('type.semanticPrimitives.error_result'(X521),'type.semanticPrimitives.error_result'(sK1)),X520))))) | s('type.semanticPrimitives.result'(bool,X521),X517) != s('type.semanticPrimitives.result'(bool,X521),X518) | s('type.semanticPrimitives.result'(bool,X522),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),X524))) = s('type.semanticPrimitives.result'(bool,X522),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X522),X524)),X522,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1028,f21143])).
fof(f30825,plain,(
  ( ! [X509,X511,X505,X507,X512,X514,X508,X510,X506] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X507),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X509),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X510))))) | s('type.semanticPrimitives.result'(bool,X511),X507) = s('type.semanticPrimitives.result'(bool,X511),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X511),sK10(X506,X505,X507,X508,X509,X510,'type.semanticPrimitives.error_result'(sK1),X511,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X508),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X505),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X506))))) | s('type.semanticPrimitives.result'(bool,X511),X507) != s('type.semanticPrimitives.result'(bool,X511),X508) | s('type.semanticPrimitives.result'(bool,X512),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),X514))) = s('type.semanticPrimitives.result'(bool,X512),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),X514)),X512,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29969,f21143])).
fof(f29969,plain,(
  ( ! [X509,X511,X505,X507,X512,X514,X508,X510,X506,X513] : (s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X512),'type.semanticPrimitives.error_result'(sK1)),X513),s('type.semanticPrimitives.error_result'(X512),X514))) != s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X509),s(bool,sK9(X506,X505,X507,X508,X509,X510,'type.semanticPrimitives.error_result'(sK1),X511,bool)))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X507),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X509),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X510))))) | s('type.semanticPrimitives.result'(bool,X511),X507) = s('type.semanticPrimitives.result'(bool,X511),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X511),sK10(X506,X505,X507,X508,X509,X510,'type.semanticPrimitives.error_result'(sK1),X511,bool)))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X511),X508),s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X505),s(fun('type.semanticPrimitives.error_result'(X511),'type.semanticPrimitives.error_result'(sK1)),X506))))) | s('type.semanticPrimitives.result'(bool,X511),X507) != s('type.semanticPrimitives.result'(bool,X511),X508) | s('type.semanticPrimitives.result'(bool,X512),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),X514))) = s('type.semanticPrimitives.result'(bool,X512),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X512),X514)),X512,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1029,f21143])).
fof(f30822,plain,(
  ( ! [X478,X484,X477,X479] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X477,'type.semanticPrimitives.error_result'(sK1)),X478),s(X477,happ(s(fun(bool,X477),X479),s(bool,X484)))))))) ) | $spl42),
  inference(cnf_transformation,[],[f30822_D])).
fof(f30822_D,plain,(
  ( ! [X478,X484,X477,X479] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X477,'type.semanticPrimitives.error_result'(sK1)),X478),s(X477,happ(s(fun(bool,X477),X479),s(bool,X484)))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f29968,plain,(
  ( ! [X499,X500,X502,X498,X504,X501,X503] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X502),'type.semanticPrimitives.error_result'(sK1)),X503),s('type.semanticPrimitives.error_result'(X502),X504))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X498,'type.semanticPrimitives.error_result'(sK1)),X499),s(X498,happ(s(fun(bool,X498),X500),s(bool,X501))))))) | s('type.semanticPrimitives.result'(bool,X502),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X502),X504))) = s('type.semanticPrimitives.result'(bool,X502),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X502),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X502),X504)),X502,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1038,f21143])).
fof(f30820,plain,(
  ( ! [X476,X469,X471,X470] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X469,'type.semanticPrimitives.error_result'(sK1)),X470),s(X469,happ(s(fun(bool,X469),X471),s(bool,X476)))))))) ) | $spl38),
  inference(cnf_transformation,[],[f30820_D])).
fof(f30820_D,plain,(
  ( ! [X476,X469,X471,X470] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X469,'type.semanticPrimitives.error_result'(sK1)),X470),s(X469,happ(s(fun(bool,X469),X471),s(bool,X476)))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f29967,plain,(
  ( ! [X492,X494,X496,X493,X495,X491,X497] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X495),'type.semanticPrimitives.error_result'(sK1)),X496),s('type.semanticPrimitives.error_result'(X495),X497))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X491,'type.semanticPrimitives.error_result'(sK1)),X492),s(X491,happ(s(fun(bool,X491),X493),s(bool,X494))))))) | s('type.semanticPrimitives.result'(bool,X495),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X495),X497))) = s('type.semanticPrimitives.result'(bool,X495),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X495),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X495),X497)),X495,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1039,f21143])).
fof(f29966,plain,(
  ( ! [X488,X490,X485,X487,X489] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X487),'type.semanticPrimitives.error_result'(sK1)),X488),s('type.semanticPrimitives.error_result'(X487),X489))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),X485),s(bool,X490))))) | s('type.semanticPrimitives.result'(bool,X487),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X487),X489))) = s('type.semanticPrimitives.result'(bool,X487),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X487),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X487),X489)),X487,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f2362,f21143])).
fof(f29965,plain,(
  ( ! [X478,X484,X482,X477,X479,X481,X483] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X481),'type.semanticPrimitives.error_result'(sK1)),X482),s('type.semanticPrimitives.error_result'(X481),X483))))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X477,'type.semanticPrimitives.error_result'(sK1)),X478),s(X477,happ(s(fun(bool,X477),X479),s(bool,X484))))))) | s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483))) = s('type.semanticPrimitives.result'(bool,X481),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X481),X483)),X481,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f1678,f21143])).
fof(f29964,plain,(
  ( ! [X476,X474,X469,X471,X475,X473,X470] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun('type.semanticPrimitives.error_result'(X473),'type.semanticPrimitives.error_result'(sK1)),X474),s('type.semanticPrimitives.error_result'(X473),X475))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X469,'type.semanticPrimitives.error_result'(sK1)),X470),s(X469,happ(s(fun(bool,X469),X471),s(bool,X476))))))) | s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475))) = s('type.semanticPrimitives.result'(bool,X473),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X473),X475)),X473,bool))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(superposition,[],[f2020,f21143])).
fof(f2020,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153)))))))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f2019,f819])).
fof(f2019,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153))) ) | ($spl2 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f1985,f819])).
fof(f1985,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK1)),'_dst_x0x1_2'(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(fun(bool,X155),X157))),s(bool,X154))))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),happ(s(fun(X155,'type.semanticPrimitives.error_result'(sK1)),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,X153))))))) | ~p(s(bool,X153)) | ~p(s(bool,X154))) ) | $spl2),
  inference(superposition,[],[f1039,f90])).
fof(f30817,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X65),s(fun(bool,X61),X62),s(fun('type.semanticPrimitives.error_result'(X68),X61),X67))) = s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X66),s(fun(bool,X61),X64),s(fun('type.semanticPrimitives.error_result'(X68),X61),X63))) | s('type.semanticPrimitives.result'(bool,X68),X65) != s('type.semanticPrimitives.result'(bool,X68),X66) | s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X63),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) != s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X67),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool))))) ) | $spl36),
  inference(cnf_transformation,[],[f30817_D])).
fof(f30817_D,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X65),s(fun(bool,X61),X62),s(fun('type.semanticPrimitives.error_result'(X68),X61),X67))) = s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X66),s(fun(bool,X61),X64),s(fun('type.semanticPrimitives.error_result'(X68),X61),X63))) | s('type.semanticPrimitives.result'(bool,X68),X65) != s('type.semanticPrimitives.result'(bool,X68),X66) | s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X63),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) != s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X67),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f30819,plain,(
  ( ! [X90,X88,X87,X85,X83,X93,X91,X89,X86,X84] : (s(X83,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X90),X86),s(fun(bool,X83),X88),s(fun('type.semanticPrimitives.error_result'(X90),X83),X89))) = s(X83,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X90),X87),s(fun(bool,X83),X84),s(fun('type.semanticPrimitives.error_result'(X90),X83),X85))) | s(X83,happ(s(fun('type.semanticPrimitives.error_result'(X90),X83),X85),s('type.semanticPrimitives.error_result'(X90),sK10(X85,X84,X86,X87,X88,X89,X83,X90,bool)))) != s(X83,happ(s(fun('type.semanticPrimitives.error_result'(X90),X83),X89),s('type.semanticPrimitives.error_result'(X90),sK10(X85,X84,X86,X87,X88,X89,X83,X90,bool)))) | s('type.semanticPrimitives.result'(bool,X90),X86) != s('type.semanticPrimitives.result'(bool,X90),X87) | s('type.semanticPrimitives.result'(bool,X91),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),X93))) = s('type.semanticPrimitives.result'(bool,X91),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),X93)),X91,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29917,f21143])).
fof(f29917,plain,(
  ( ! [X92,X90,X88,X87,X85,X83,X93,X91,X89,X86,X84] : (s(X83,happ(s(fun('type.semanticPrimitives.error_result'(X91),X83),X92),s('type.semanticPrimitives.error_result'(X91),X93))) != s(X83,happ(s(fun(bool,X83),X88),s(bool,sK9(X85,X84,X86,X87,X88,X89,X83,X90,bool)))) | s(X83,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X90),X86),s(fun(bool,X83),X88),s(fun('type.semanticPrimitives.error_result'(X90),X83),X89))) = s(X83,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X90),X87),s(fun(bool,X83),X84),s(fun('type.semanticPrimitives.error_result'(X90),X83),X85))) | s(X83,happ(s(fun('type.semanticPrimitives.error_result'(X90),X83),X85),s('type.semanticPrimitives.error_result'(X90),sK10(X85,X84,X86,X87,X88,X89,X83,X90,bool)))) != s(X83,happ(s(fun('type.semanticPrimitives.error_result'(X90),X83),X89),s('type.semanticPrimitives.error_result'(X90),sK10(X85,X84,X86,X87,X88,X89,X83,X90,bool)))) | s('type.semanticPrimitives.result'(bool,X90),X86) != s('type.semanticPrimitives.result'(bool,X90),X87) | s('type.semanticPrimitives.result'(bool,X91),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),X93))) = s('type.semanticPrimitives.result'(bool,X91),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X91),X93)),X91,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f79,f21143])).
fof(f30815,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51] : (s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X54),s(fun(bool,X50),X51),s(fun('type.semanticPrimitives.error_result'(X57),X50),X56))) = s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X55),s(fun(bool,X50),X53),s(fun('type.semanticPrimitives.error_result'(X57),X50),X52))) | s('type.semanticPrimitives.result'(bool,X57),X54) != s('type.semanticPrimitives.result'(bool,X57),X55) | s('type.semanticPrimitives.result'(bool,X57),X54) = s('type.semanticPrimitives.result'(bool,X57),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X57),sK10(X52,X53,X54,X55,X51,X56,X50,X57,bool))))) ) | $spl34),
  inference(cnf_transformation,[],[f30815_D])).
fof(f30815_D,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51] : (s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X54),s(fun(bool,X50),X51),s(fun('type.semanticPrimitives.error_result'(X57),X50),X56))) = s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X55),s(fun(bool,X50),X53),s(fun('type.semanticPrimitives.error_result'(X57),X50),X52))) | s('type.semanticPrimitives.result'(bool,X57),X54) != s('type.semanticPrimitives.result'(bool,X57),X55) | s('type.semanticPrimitives.result'(bool,X57),X54) = s('type.semanticPrimitives.result'(bool,X57),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X57),sK10(X52,X53,X54,X55,X51,X56,X50,X57,bool))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f30818,plain,(
  ( ! [X80,X78,X76,X74,X72,X79,X77,X75,X73,X82] : (s(X72,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X79),X75),s(fun(bool,X72),X77),s(fun('type.semanticPrimitives.error_result'(X79),X72),X78))) = s(X72,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X79),X76),s(fun(bool,X72),X73),s(fun('type.semanticPrimitives.error_result'(X79),X72),X74))) | s('type.semanticPrimitives.result'(bool,X79),X75) = s('type.semanticPrimitives.result'(bool,X79),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X79),sK10(X74,X73,X75,X76,X77,X78,X72,X79,bool)))) | s('type.semanticPrimitives.result'(bool,X79),X75) != s('type.semanticPrimitives.result'(bool,X79),X76) | s('type.semanticPrimitives.result'(bool,X80),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),X82))) = s('type.semanticPrimitives.result'(bool,X80),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),X82)),X80,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29916,f21143])).
fof(f29916,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X79,X77,X75,X73,X82] : (s(X72,happ(s(fun('type.semanticPrimitives.error_result'(X80),X72),X81),s('type.semanticPrimitives.error_result'(X80),X82))) != s(X72,happ(s(fun(bool,X72),X77),s(bool,sK9(X74,X73,X75,X76,X77,X78,X72,X79,bool)))) | s(X72,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X79),X75),s(fun(bool,X72),X77),s(fun('type.semanticPrimitives.error_result'(X79),X72),X78))) = s(X72,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X79),X76),s(fun(bool,X72),X73),s(fun('type.semanticPrimitives.error_result'(X79),X72),X74))) | s('type.semanticPrimitives.result'(bool,X79),X75) = s('type.semanticPrimitives.result'(bool,X79),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X79),sK10(X74,X73,X75,X76,X77,X78,X72,X79,bool)))) | s('type.semanticPrimitives.result'(bool,X79),X75) != s('type.semanticPrimitives.result'(bool,X79),X76) | s('type.semanticPrimitives.result'(bool,X80),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),X82))) = s('type.semanticPrimitives.result'(bool,X80),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X80),X82)),X80,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f77,f21143])).
fof(f30816,plain,(
  ( ! [X61,X68,X66,X64,X62,X71,X69,X67,X65,X63] : (s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X65),s(fun(bool,X61),X62),s(fun('type.semanticPrimitives.error_result'(X68),X61),X67))) = s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X66),s(fun(bool,X61),X64),s(fun('type.semanticPrimitives.error_result'(X68),X61),X63))) | s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X63),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) != s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X67),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) | s('type.semanticPrimitives.result'(bool,X68),X65) != s('type.semanticPrimitives.result'(bool,X68),X66) | s('type.semanticPrimitives.result'(bool,X69),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),X71))) = s('type.semanticPrimitives.result'(bool,X69),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),X71)),X69,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29915,f21143])).
fof(f29915,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X71,X69,X67,X65,X63] : (s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X69),X61),X70),s('type.semanticPrimitives.error_result'(X69),X71))) != s(X61,happ(s(fun(bool,X61),X64),s(bool,sK9(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) | s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X65),s(fun(bool,X61),X62),s(fun('type.semanticPrimitives.error_result'(X68),X61),X67))) = s(X61,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X68),X66),s(fun(bool,X61),X64),s(fun('type.semanticPrimitives.error_result'(X68),X61),X63))) | s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X63),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) != s(X61,happ(s(fun('type.semanticPrimitives.error_result'(X68),X61),X67),s('type.semanticPrimitives.error_result'(X68),sK10(X63,X64,X65,X66,X62,X67,X61,X68,bool)))) | s('type.semanticPrimitives.result'(bool,X68),X65) != s('type.semanticPrimitives.result'(bool,X68),X66) | s('type.semanticPrimitives.result'(bool,X69),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),X71))) = s('type.semanticPrimitives.result'(bool,X69),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X69),X71)),X69,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f79,f21143])).
fof(f30814,plain,(
  ( ! [X57,X54,X52,X50,X60,X58,X56,X55,X53,X51] : (s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X54),s(fun(bool,X50),X51),s(fun('type.semanticPrimitives.error_result'(X57),X50),X56))) = s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X55),s(fun(bool,X50),X53),s(fun('type.semanticPrimitives.error_result'(X57),X50),X52))) | s('type.semanticPrimitives.result'(bool,X57),X54) = s('type.semanticPrimitives.result'(bool,X57),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X57),sK10(X52,X53,X54,X55,X51,X56,X50,X57,bool)))) | s('type.semanticPrimitives.result'(bool,X57),X54) != s('type.semanticPrimitives.result'(bool,X57),X55) | s('type.semanticPrimitives.result'(bool,X58),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),X60))) = s('type.semanticPrimitives.result'(bool,X58),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),X60)),X58,bool))))) ) | ($spl16 | $spl28)),
  inference(subsumption_resolution,[],[f29914,f21143])).
fof(f29914,plain,(
  ( ! [X59,X57,X54,X52,X50,X60,X58,X56,X55,X53,X51] : (s(X50,happ(s(fun('type.semanticPrimitives.error_result'(X58),X50),X59),s('type.semanticPrimitives.error_result'(X58),X60))) != s(X50,happ(s(fun(bool,X50),X53),s(bool,sK9(X52,X53,X54,X55,X51,X56,X50,X57,bool)))) | s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X54),s(fun(bool,X50),X51),s(fun('type.semanticPrimitives.error_result'(X57),X50),X56))) = s(X50,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(bool,X57),X55),s(fun(bool,X50),X53),s(fun('type.semanticPrimitives.error_result'(X57),X50),X52))) | s('type.semanticPrimitives.result'(bool,X57),X54) = s('type.semanticPrimitives.result'(bool,X57),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X57),sK10(X52,X53,X54,X55,X51,X56,X50,X57,bool)))) | s('type.semanticPrimitives.result'(bool,X57),X54) != s('type.semanticPrimitives.result'(bool,X57),X55) | s('type.semanticPrimitives.result'(bool,X58),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),X60))) = s('type.semanticPrimitives.result'(bool,X58),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),X60)),X58,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f77,f21143])).
fof(f30781,plain,(
  ( ! [X14,X12,X8,X15,X13] : (s(X8,happ(s(fun(bool,X8),X12),s(bool,X13))) = s(X8,happ(s(fun(bool,X8),X14),s(bool,X15)))) ) | $spl32),
  inference(cnf_transformation,[],[f30781_D])).
fof(f30781_D,plain,(
  ( ! [X14,X12,X8,X15,X13] : (s(X8,happ(s(fun(bool,X8),X12),s(bool,X13))) = s(X8,happ(s(fun(bool,X8),X14),s(bool,X15)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f30748,plain,(
  ( ! [X14,X12,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(bool,X8),X12),s(bool,X13))) = s(X8,happ(s(fun(bool,X8),X14),s(bool,X15))) | s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)),X9,bool))))) ) | ($spl16 | $spl28)),
  inference(duplicate_literal_removal,[],[f29832])).
fof(f29832,plain,(
  ( ! [X14,X12,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(bool,X8),X12),s(bool,X13))) = s(X8,happ(s(fun(bool,X8),X14),s(bool,X15))) | s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)),X9,bool)))) | s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK12('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)),X9,bool))))) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f21143,f21143])).
fof(f1047,plain,(
  $false | ($spl0 | $spl9 | $spl12)),
  inference(subsumption_resolution,[],[f643,f98])).
fof(f643,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ($spl0 | $spl12)),
  inference(superposition,[],[f92,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.semanticPrimitives.error_result'(sK1),X0),X1),s('type.semanticPrimitives.error_result'(sK1),sK8))) = s(X0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,X0),X2),s(fun('type.semanticPrimitives.error_result'(sK1),X0),X1)))) ) | $spl12),
  inference(superposition,[],[f80,f103])).
fof(f1046,plain,(
  $false | ($spl2 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f1041,f103])).
fof(f1041,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f1040])).
fof(f1040,plain,(
  s(sK2,sK4) != s(sK2,sK4) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) | ($spl2 | $spl8)),
  inference(superposition,[],[f94,f99])).
fof(f837,plain,(
  $false | ($spl16 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f803,f506])).
fof(f803,plain,(
  ( ! [X140] : (~p(s(bool,happ(s(fun(sK0,bool),X140),s(sK0,sK7))))) ) | ($spl16 | $spl22 | $spl28)),
  inference(superposition,[],[f562,f736])).
fof(f667,plain,(
  $false | ($spl0 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f654,f96])).
fof(f96,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | $spl7),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f654,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | ($spl0 | $spl10)),
  inference(superposition,[],[f104,f92])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK7))) = s(X0,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,X0),X1),s(fun('type.semanticPrimitives.error_result'(sK1),X0),X2)))) ) | $spl10),
  inference(superposition,[],[f85,f101])).
fof(f666,plain,(
  $false | ($spl0 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f644,f96])).
fof(f644,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | ($spl0 | $spl10)),
  inference(superposition,[],[f92,f104])).
fof(f614,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f564,f562])).
fof(f505,plain,(
  ( ! [X61,X60] : (s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X61))) ) | $spl14),
  inference(cnf_transformation,[],[f505_D])).
fof(f505_D,plain,(
  ( ! [X61,X60] : (s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X61))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f506,plain,(
  ( ! [X59] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK7))))) ) | $spl16),
  inference(cnf_transformation,[],[f506_D])).
fof(f506_D,plain,(
  ( ! [X59] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK7))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f514,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK7)))) | s(bool,X58) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X57),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X58))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f438,f104])).
fof(f438,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X57),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X58)) | p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X56),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X57))))) ) | $spl12),
  inference(superposition,[],[f87,f329])).
fof(f515,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK7)))) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X61))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f439,f104])).
fof(f439,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X59),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60)))) | p(s(bool,X61))) ) | $spl12),
  inference(superposition,[],[f87,f329])).
fof(f565,plain,(
  ( ! [X67,X65] : (p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7)))) ) | $spl28),
  inference(cnf_transformation,[],[f565_D])).
fof(f565_D,plain,(
  ( ! [X67,X65] : (p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f516,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,X64))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f440,f104])).
fof(f440,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X64) = s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X62),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63))) | p(s(bool,X64))) ) | $spl12),
  inference(superposition,[],[f87,f329])).
fof(f517,plain,(
  ( ! [X66,X67,X65] : (s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | p(s(bool,X67))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f441,f104])).
fof(f441,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X67) = s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X65),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66))) | p(s(bool,X67))) ) | $spl12),
  inference(superposition,[],[f87,f329])).
fof(f509,plain,(
  ( ! [X66,X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8)))) ) | $spl18),
  inference(cnf_transformation,[],[f509_D])).
fof(f509_D,plain,(
  ( ! [X66,X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f510,plain,(
  ( ! [X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7))))) ) | $spl20),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  ( ! [X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f518,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK7)))) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X70))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f442,f104])).
fof(f442,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X68),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69)))) | ~p(s(bool,X70))) ) | $spl12),
  inference(superposition,[],[f90,f329])).
fof(f519,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK7)))) | s(bool,X73) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X72),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X73))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f443,f104])).
fof(f443,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X72),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X73)) | ~p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X71),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X72))))) ) | $spl12),
  inference(superposition,[],[f90,f329])).
fof(f562,plain,(
  ( ! [X78] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X78),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) | $spl22),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ( ! [X78] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X78),s('type.semanticPrimitives.error_result'(sK1),sK8))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f563,plain,(
  ( ! [X79,X77] : (s(bool,X79) = s(bool,happ(s(fun(sK0,bool),X77),s(sK0,sK7))) | ~p(s(bool,X79))) ) | $spl24),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  ( ! [X79,X77] : (s(bool,X79) = s(bool,happ(s(fun(sK0,bool),X77),s(sK0,sK7))) | ~p(s(bool,X79))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f520,plain,(
  ( ! [X76,X74,X75] : (s(bool,X76) = s(bool,happ(s(fun(sK0,bool),X74),s(sK0,sK7))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X75),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | ~p(s(bool,X76))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f444,f104])).
fof(f444,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X75),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X76) = s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X74),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X75))) | ~p(s(bool,X76))) ) | $spl12),
  inference(superposition,[],[f90,f329])).
fof(f521,plain,(
  ( ! [X78,X79,X77] : (s(bool,X79) = s(bool,happ(s(fun(sK0,bool),X77),s(sK0,sK7))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X78),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | ~p(s(bool,X79))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f445,f104])).
fof(f445,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X78),s('type.semanticPrimitives.error_result'(sK1),sK8)))) | s(bool,X79) = s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X77),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X78))) | ~p(s(bool,X79))) ) | $spl12),
  inference(superposition,[],[f90,f329])).
fof(f513,plain,(
  $false | ($spl1 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f437,f99])).
fof(f437,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ($spl1 | $spl12)),
  inference(superposition,[],[f93,f329])).
fof(f511,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK7)))) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X70))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f426,f104])).
fof(f426,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X68),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X69)))) | ~p(s(bool,X70))) ) | $spl12),
  inference(superposition,[],[f329,f90])).
fof(f508,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK7)))) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X67))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f425,f104])).
fof(f425,plain,(
  ( ! [X66,X67,X65] : (s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66),s('type.semanticPrimitives.error_result'(sK1),sK8))) | ~p(s(bool,X67)) | ~p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X65),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X66))))) ) | $spl12),
  inference(superposition,[],[f329,f90])).
fof(f507,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK7)))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X64))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f424,f104])).
fof(f424,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X64)) | p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X62),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X63))))) ) | $spl12),
  inference(superposition,[],[f329,f87])).
fof(f504,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK7)))) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,X61))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f423,f104])).
fof(f423,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60),s('type.semanticPrimitives.error_result'(sK1),sK8))) | p(s(bool,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,bool),X59),s(fun('type.semanticPrimitives.error_result'(sK1),bool),X60)))) | p(s(bool,X61))) ) | $spl12),
  inference(superposition,[],[f329,f87])).
fof(f280,plain,(
  $false | ($spl1 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f210,f97])).
fof(f210,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | ($spl1 | $spl10)),
  inference(superposition,[],[f93,f104])).
fof(f92,plain,(
  s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | $spl0),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f101,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))) | $spl10),
  inference(cnf_transformation,[],[f101_D])).
fof(f101_D,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f103,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) | $spl12),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f70,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  (s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))) & s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7)))) | (s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) & s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))))) & (s(sK2,sK4) != s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | (! [X8] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X8))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X8)))) & ! [X7] : (s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X7))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),X7))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | ? [X10] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X10))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X10)))) | ? [X9] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X9))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X9))))) & (s(X2,X4) != s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | (! [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) | s(X2,X4) != s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) & ! [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) | s(X2,X4) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7)))))))),
  inference(rectify,[],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | ? [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) | ? [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7))))) & (s(X2,X4) != s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | (! [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) | s(X2,X4) != s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) & ! [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) | s(X2,X4) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | (? [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) | ? [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7)))))) & (s(X2,X4) != s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) | (! [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) | s(X2,X4) != s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) & ! [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) | s(X2,X4) != s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7)))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) <~> (? [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) | ? [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7))))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) <=> (? [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) | ? [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7))))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s(X2,X4) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.error_result'(X1),X2),X5))) <=> (? [X8] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X8))) & s(X2,X4) = s(X2,happ(s(fun(X0,X2),X6),s(X0,X8)))) | ? [X7] : (s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X7))) & s(X2,X4) = s(X2,happ(s(fun('type.semanticPrimitives.error_result'(X1),X2),X5),s('type.semanticPrimitives.error_result'(X1),X7))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12] : (s(X8,X10) = s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X9),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.error_result'(X7),X8),X11))) <=> (? [X13] : (s('type.semanticPrimitives.result'(X6,X7),X9) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X13))) & s(X8,X10) = s(X8,happ(s(fun(X6,X8),X12),s(X6,X13)))) | ? [X14] : (s('type.semanticPrimitives.result'(X6,X7),X9) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X14))) & s(X8,X10) = s(X8,happ(s(fun('type.semanticPrimitives.error_result'(X7),X8),X11),s('type.semanticPrimitives.error_result'(X7),X14))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12] : (s(X8,X10) = s(X8,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X7),X9),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.error_result'(X7),X8),X11))) <=> (? [X13] : (s('type.semanticPrimitives.result'(X6,X7),X9) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rval_1'(s(X6,X13))) & s(X8,X10) = s(X8,happ(s(fun(X6,X8),X12),s(X6,X13)))) | ? [X14] : (s('type.semanticPrimitives.result'(X6,X7),X9) = s('type.semanticPrimitives.result'(X6,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X14))) & s(X8,X10) = s(X8,happ(s(fun('type.semanticPrimitives.error_result'(X7),X8),X11),s('type.semanticPrimitives.error_result'(X7),X14))))))),
  file('thm.closProps.eqs|split|5','thm.closProps.eqs|split|5')).
fof(f97,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | $spl6),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f71,plain,(
  s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),sK8))) | s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))),
  inference(cnf_transformation,[],[f67])).
fof(f99,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | $spl8),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f72,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) = s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,sK7))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))),
  inference(cnf_transformation,[],[f67])).
fof(f73,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),sK8))) | s(sK2,sK4) = s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,sK7))) | s(sK2,sK4) = s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))),
  inference(cnf_transformation,[],[f67])).
fof(f93,plain,(
  s(sK2,sK4) != s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) | $spl1),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(sK2,sK4) != s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f95,plain,(
  ( ! [X8] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X8)))) ) | $spl4),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  ( ! [X8] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X8)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f74,plain,(
  ( ! [X8] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X8))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rval_1'(s(sK0,X8))) | s(sK2,sK4) != s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f94,plain,(
  ( ! [X7] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),X7))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  ( ! [X7] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),X7))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f75,plain,(
  ( ! [X7] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5),s('type.semanticPrimitives.error_result'(sK1),X7))) | s('type.semanticPrimitives.result'(sK0,sK1),sK3) != s('type.semanticPrimitives.result'(sK0,sK1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK1),X7))) | s(sK2,sK4) != s(sK2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK1),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.error_result'(sK1),sK2),sK5)))) )),
  inference(cnf_transformation,[],[f67])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 181
% Passive clauses: 12262
% Generated clauses: 33307
% Final active clauses: 80
% Final passive clauses: 1714
% Input formulas: 36
% Initial clauses: 22
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 107
% Trivial inequalities: 37
% Fw subsumption resolutions: 444
% Fw demodulations: 6997
% 
% Simple tautologies: 99
% Equational tautologies: 3
% Forward subsumptions: 13367
% 
% Binary resolution: 214
% Factoring: 2
% Forward superposition: 13639
% Backward superposition: 11249
% Self superposition: 322
% Equality factoring: 12
% Equality resolution: 36
% 
% Split clauses: 111
% Split components: 226
% SAT solver clauses: 121
% SAT solver binary clauses: 3
% 
% Sat splits: 111
% Sat splitting refutations: 9
% 
% Memory used [KB]: 66011
% Time elapsed: 2.183 s
% ------------------------------
% ------------------------------
% Success in time 13.95 s
