Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.mod_to_conProof
% (35967)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 879
% Passive clauses: 5520
% Generated clauses: 23366
% Final active clauses: 870
% Final passive clauses: 4374
% Input formulas: 260
% Initial clauses: 879
% 
% Duplicate literals: 228
% Trivial inequalities: 41
% Fw subsumption resolutions: 1203
% Fw demodulations: 2471
% 
% Simple tautologies: 338
% Equational tautologies: 57
% Forward subsumptions: 13094
% 
% Binary resolution: 5820
% Factoring: 1073
% Forward superposition: 5656
% Backward superposition: 6163
% Self superposition: 27
% Equality factoring: 2
% Equality resolution: 32
% 
% Split clauses: 512
% Split components: 1026
% SAT solver clauses: 682
% SAT solver unit clauses: 11
% SAT solver binary clauses: 51
% 
% Sat splits: 512
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 682
% TWLsolver calls for satisfiability: 358
% 
% Memory used [KB]: 13432
% Time elapsed: 0.622 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.mod_to_conProof
% (35980)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 566
% Passive clauses: 3153
% Generated clauses: 15139
% Final active clauses: 552
% Final passive clauses: 2548
% Input formulas: 260
% Initial clauses: 879
% 
% Duplicate literals: 106
% Trivial inequalities: 3
% Fw subsumption resolutions: 573
% Fw demodulations: 1663
% Bw demodulations: 53
% Condensations: 104
% Global subsumptions: 26
% 
% Simple tautologies: 289
% Equational tautologies: 22
% Forward subsumptions: 5558
% 
% Binary resolution: 2722
% Factoring: 306
% Forward superposition: 3431
% Backward superposition: 1943
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 21
% Split components: 42
% SAT solver clauses: 4538
% SAT solver unit clauses: 963
% SAT solver binary clauses: 275
% SAT solver learnt clauses: 52
% SAT solver learnt literals: 72
% 
% Sat splits: 21
% 
% TWLsolver clauses: 4434
% TWLsolver calls for satisfiability: 26412
% 
% Memory used [KB]: 21364
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.mod_to_conProof
% (35996)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 956
% Passive clauses: 10259
% Generated clauses: 83617
% Final active clauses: 934
% Final passive clauses: 7581
% Input formulas: 260
% Initial clauses: 879
% Discarded non-redundant clauses: 3064
% 
% Duplicate literals: 10
% Trivial inequalities: 11
% Fw subsumption resolutions: 918
% Fw demodulations: 29373
% Bw demodulations: 426
% Condensations: 45
% 
% Simple tautologies: 161
% Equational tautologies: 70
% Forward subsumptions: 24884
% Fw demodulations to eq. taut.: 14747
% 
% Unit resulting resolution: 33525
% Factoring: 1
% Forward superposition: 8727
% Backward superposition: 9557
% Self superposition: 82
% Equality factoring: 2
% Equality resolution: 32
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 11522
% SAT solver unit clauses: 4966
% SAT solver binary clauses: 4658
% 
% Sat splits: 27
% Sat splitting refutations: 1207
% 
% Memory used [KB]: 216840
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 31 next slice time: 31
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.mod_to_conProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.mod_to_conProof
% SZS output start Proof for thm.mod_to_conProof
fof(f24422,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1062,f1064,f1065,f1066,f1069,f1165,f1218,f1225,f1227,f1229,f1231,f1232,f1327,f1377,f1378,f1379,f1380,f1381,f1382,f1383,f1519,f1683,f1685,f1686,f1687,f1690,f1786,f2020,f2329,f2325,f2323,f3700,f3702,f3703,f3704,f3707,f3803,f5564,f5566,f5567,f5568,f5571,f5667,f5762,f5763,f5764,f5765,f5766,f1388,f5767,f1392,f6137,f6741,f6743,f6744,f6745,f6748,f6844,f8294,f8296,f8298,f8300,f8301,f8396,f8716,f8718,f8719,f8720,f8723,f8819,f9110,f9111,f9112,f9130,f1831,f9131,f5728,f9132,f3848,f9204,f9216,f9418,f9425,f8941,f9452,f9460,f9454,f9461,f9455,f9462,f9456,f9463,f9642,f9588,f9578,f9574,f9572,f9772,f8507,f12108,f12110,f12112,f12114,f12115,f12210,f12802,f12831,f12884,f12898,f12907,f12954,f12965,f12968,f12976,f12974,f12986,f12996,f13050,f12322,f13058,f13231,f13233,f13235,f13237,f13238,f13333,f13450,f13507,f13508,f13525,f13526,f13534,f13535,f13542,f13594,f13597,f13605,f13603,f13615,f13625,f13628,f13445,f13636,f13644,f13649,f13675,f13680,f13968,f13973,f13999,f14004,f14035,f14040,f14067,f14072,f13516,f14242,f14241,f14240,f14264,f12826,f14263,f14247,f14262,f13514,f14268,f14270,f14272,f14274,f14275,f14370,f14482,f14483,f14485,f14487,f14489,f14490,f14585,f14697,f14706,f14708,f14710,f14712,f14713,f14808,f14892,f15063,f15068,f1396,f15069,f1394,f15070,f1390,f1386,f15283,f15412,f15725,f15727,f15728,f15729,f15732,f15828,f15860,f15903,f15410,f15917,f15919,f15920,f15921,f15924,f16020,f16307,f16351,f16479,f16521,f16094,f16614,f16616,f16618,f16620,f16621,f16716,f16782,f17002,f16932,f16836,f16833,f16832,f16831,f16829,f17068,f17070,f17072,f17074,f17075,f17170,f17303,f17306,f17308,f17310,f17312,f17313,f17408,f17646,f17648,f17650,f17652,f17653,f17748,f17881,f17972,f18015,f18140,f18142,f18144,f18146,f18147,f18242,f18375,f18376,f17470,f18627,f18629,f18631,f18633,f18634,f18729,f18862,f18863,f18907,f18998,f19040,f17013,f19084,f19086,f19087,f19088,f19091,f19187,f19320,f19414,f19416,f19417,f19418,f19421,f19517,f19590,f19591,f19616,f19592,f19593,f19620,f19594,f19614,f19595,f19622,f19618,f19612,f19857,f19859,f19860,f19861,f19864,f19960,f20093,f20394,f20396,f20397,f20398,f20401,f20497,f20570,f20571,f20596,f20572,f20573,f20600,f20574,f20594,f20575,f20602,f20598,f20592,f21311,f21359,f21458,f21504,f19589,f21628,f21630,f21632,f21634,f21635,f21730,f21865,f22037,f22039,f22041,f22043,f22044,f22139,f22274,f22369,f22417,f22547,f22549,f22551,f22553,f22554,f22649,f22784,f22786,f22832,f20569,f23000,f23002,f23004,f23006,f23007,f23102,f23237,f23270,f23274,f23272,f23301,f23303,f23305,f23307,f23308,f23403,f23521,f23722,f23765,f23853,f23855,f23857,f1219,f23859,f1221,f23860,f1220,f23955,f24073,f24075,f24116,f1222,f15901,f24417])).
fof(f24417,plain,(
  $false | ($spl16 | $spl20 | $spl22)),
  inference(subsumption_resolution,[],[f24416,f679])).
fof(f679,plain,(
  ~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))),
  inference(cnf_transformation,[],[f508])).
fof(f508,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) & ~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20,sK21,sK22])],[f433])).
fof(f433,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) & ~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))))),
  inference(flattening,[],[f432])).
fof(f432,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3))))) & ~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))))),
  inference(ennf_transformation,[],[f262])).
fof(f262,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3))))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.semanticPrimitives.store_v'('type.modSem.v'),X7),s('type.semanticPrimitives.store_v'('type.conSem.v'),X8)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9))))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9),s('type.semanticPrimitives.store_v'('type.modSem.v'),X7),s('type.semanticPrimitives.store_v'('type.conSem.v'),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.semanticPrimitives.store_v'('type.modSem.v'),X7),s('type.semanticPrimitives.store_v'('type.conSem.v'),X8)))) & p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9))))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9),s('type.semanticPrimitives.store_v'('type.modSem.v'),X7),s('type.semanticPrimitives.store_v'('type.conSem.v'),X8)))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.mod_to_conProof_58_prove_ltb')).
fof(f24416,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | ($spl16 | $spl20 | $spl22)),
  inference(forward_demodulation,[],[f24415,f1390])).
fof(f24415,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | ($spl20 | $spl22)),
  inference(forward_demodulation,[],[f24239,f1394])).
fof(f24239,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))))) | $spl22),
  inference(resolution,[],[f16605,f953])).
fof(f953,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X2)))))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2))))) )),
  inference(cnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2)))) | p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X2)))))))),
  inference(ennf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2)))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X2)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X25,X19,X20] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X25),s('type.list.list'('type.modSem.v'),X19),s('type.list.list'('type.conSem.v'),X20)))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X25),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X19))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X20)))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.sv_rel_rules|split|2')).
fof(f16605,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | $spl22),
  inference(resolution,[],[f1676,f1396])).
fof(f1676,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X2),s('type.list.list'('type.conSem.v'),X3)))) | p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X2),s('type.list.list'('type.conSem.v'),X3))))) )),
  inference(resolution,[],[f678,f683])).
fof(f683,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2)))) | ~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2))))) )),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2)))) | ! [X3] : (~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) | p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2))))))),
  inference(ennf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2)))) => ! [X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) => p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.list.list'('type.modSem.v'),X1),s('type.list.list'('type.conSem.v'),X2))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X19,X29] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.list.list'('type.modSem.v'),X19),s('type.list.list'('type.conSem.v'),X29)))) => ! [X9] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9)))) => p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9),s('type.list.list'('type.modSem.v'),X19),s('type.list.list'('type.conSem.v'),X29))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.mod_to_conProof_37_prove_ltb|split|1')).
fof(f678,plain,(
  p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))))),
  inference(cnf_transformation,[],[f508])).
fof(f15901,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl70),
  inference(cnf_transformation,[],[f15901_D])).
fof(f15901_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1222,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl8),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f24116,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | p(s(bool,X1))) ) | $spl22),
  inference(resolution,[],[f15789,f9058])).
fof(f9058,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | s(bool,X7) != s(bool,X8) | p(s(bool,X8))) )),
  inference(resolution,[],[f9054,f794])).
fof(f794,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f551])).
fof(f551,plain,(
  ! [X1,X0,X2] : ((~sP9(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP9(X1,X0,X2)))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X1,X0,X2] : ((~sP9(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP9(X1,X0,X2)))),
  inference(nnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X1,X0,X2] : (sP9(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f9054,plain,(
  ( ! [X4,X5,X3] : (~sP9(X4,X5,X3) | p(s(bool,X3)) | ~sP8(X5,X4,X3) | s(bool,X4) != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f9047])).
fof(f9047,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP9(X4,X5,X3) | ~sP8(X5,X4,X3) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5)) )),
  inference(resolution,[],[f9045,f802])).
fof(f802,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP11(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f555])).
fof(f555,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP11(X2,X1,X0)) & (~sP11(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f495])).
fof(f495,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP11(X2,X1,X0))),
  inference(definition_folding,[],[f283,f494,f493,f492,f491])).
fof(f491,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f493,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f494,plain,(
  ! [X2,X1,X0] : (sP11(X2,X1,X0) <=> (sP8(X0,X1,X2) & sP9(X1,X0,X2) & sP10(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X55,X49,X50] : ((p(s(bool,X50)) <=> s(bool,X49) = s(bool,X55)) <=> ((p(s(bool,X50)) | p(s(bool,X49)) | p(s(bool,X55))) & (p(s(bool,X50)) | ~p(s(bool,X55)) | ~p(s(bool,X49))) & (p(s(bool,X49)) | ~p(s(bool,X55)) | ~p(s(bool,X50))) & (p(s(bool,X55)) | ~p(s(bool,X49)) | ~p(s(bool,X50)))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f9045,plain,(
  ( ! [X10,X11,X9] : (sP11(X9,X10,X11) | p(s(bool,X9)) | ~sP9(X10,X11,X9) | ~sP8(X11,X10,X9)) )),
  inference(subsumption_resolution,[],[f9042,f1155])).
fof(f1155,plain,(
  ( ! [X190] : (s(bool,X190) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X190))) )),
  inference(resolution,[],[f677,f862])).
fof(f862,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f241])).
fof(f241,axiom,(
  ! [X21,X22] : (s(bool,X21) = s(bool,X22) <=> ((p(s(bool,X21)) & p(s(bool,X22))) | (~p(s(bool,X21)) & ~p(s(bool,X22)))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.EQ_EXPAND')).
fof(f677,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))),
  inference(cnf_transformation,[],[f508])).
fof(f9042,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | sP11(X9,X10,X11) | p(s(bool,X9)) | ~sP9(X10,X11,X9) | ~sP8(X11,X10,X9)) )),
  inference(resolution,[],[f8352,f786])).
fof(f786,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0)) | ~sP10(X0,X2,X1) | ~sP9(X1,X2,X0) | ~sP8(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | (sP8(X2,X1,X0) & sP9(X1,X2,X0) & sP10(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X2,X1,X0) | ~sP9(X1,X2,X0) | ~sP10(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f545])).
fof(f545,plain,(
  ! [X2,X1,X0] : ((~sP11(X2,X1,X0) | (sP8(X0,X1,X2) & sP9(X1,X0,X2) & sP10(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | ~sP9(X1,X0,X2) | ~sP10(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP11(X2,X1,X0)))),
  inference(flattening,[],[f544])).
fof(f544,plain,(
  ! [X2,X1,X0] : ((~sP11(X2,X1,X0) | (sP8(X0,X1,X2) & sP9(X1,X0,X2) & sP10(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | ~sP9(X1,X0,X2) | ~sP10(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP11(X2,X1,X0)))),
  inference(nnf_transformation,[],[f494])).
fof(f8352,plain,(
  ( ! [X182,X184,X183] : (sP10(X182,X183,X184) | s(bool,X182) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f790])).
fof(f790,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f548])).
fof(f548,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f547])).
fof(f547,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f493])).
fof(f1157,plain,(
  ( ! [X192] : (~p(s(bool,X192)) | s(bool,X192) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f677,f865])).
fof(f865,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f15789,plain,(
  ( ! [X149,X148] : (sP8('const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))),X148,X149)) ) | $spl22),
  inference(resolution,[],[f1396,f798])).
fof(f798,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f491])).
fof(f24075,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | p(s(bool,X3))) ) | $spl22),
  inference(resolution,[],[f15788,f9057])).
fof(f9057,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X5,X3) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) )),
  inference(resolution,[],[f9054,f793])).
fof(f793,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f552])).
fof(f15788,plain,(
  ( ! [X146,X147] : (sP8(X146,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))),X147)) ) | $spl22),
  inference(resolution,[],[f1396,f797])).
fof(f797,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f554])).
fof(f24073,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | s(bool,X2) != s(bool,X5) | ~p(s(bool,X5))) ) | $spl148),
  inference(cnf_transformation,[],[f24073_D])).
fof(f24073_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | s(bool,X2) != s(bool,X5) | ~p(s(bool,X5))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f23955,plain,(
  ( ! [X613,X615,X611,X612,X614,X610] : (s('type.pair.prod'(X610,bool),'const.pair.,_2'(s(X610,X611),s(bool,X612))) != s('type.pair.prod'(X610,bool),'const.pair.,_2'(s(X610,X613),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X614)) | ~p(s(bool,X615)) | s(bool,X612) != s(bool,X615)) )),
  inference(resolution,[],[f1367,f1048])).
fof(f1048,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f725])).
fof(f725,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X51,X52,X53,X54] : ((s(bool,X51) = s(bool,X52) & (p(s(bool,X52)) => s(bool,X53) = s(bool,X54))) => ((p(s(bool,X51)) => p(s(bool,X53))) <=> (p(s(bool,X52)) => p(s(bool,X54)))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.IMP_CONG')).
fof(f1367,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X26),s(bool,X24))) != s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X27),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))))) )),
  inference(superposition,[],[f679,f1042])).
fof(f1042,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X3) = s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f673])).
fof(f673,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f672])).
fof(f672,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f426])).
fof(f426,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X33,X34] : ! [X51,X53,X68,X124] : (s('type.pair.prod'(X33,X34),'const.pair.,_2'(s(X33,X51),s(X34,X53))) = s('type.pair.prod'(X33,X34),'const.pair.,_2'(s(X33,X68),s(X34,X124))) <=> (s(X33,X51) = s(X33,X68) & s(X34,X53) = s(X34,X124)))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f1220,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) | $spl4),
  inference(cnf_transformation,[],[f1220_D])).
fof(f1220_D,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f23860,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X50),s(bool,X51))) != s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X52),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X51) != s(bool,X54) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f1367,f728])).
fof(f728,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f1221,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl6),
  inference(cnf_transformation,[],[f1221_D])).
fof(f1221_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f23859,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X44))) != s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X45),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X44) != s(bool,X47)) )),
  inference(resolution,[],[f1367,f727])).
fof(f727,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f1219,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) | $spl2),
  inference(cnf_transformation,[],[f1219_D])).
fof(f1219_D,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f23857,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X29),s(bool,X30))) != s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X31),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X30) != s(bool,X33)) )),
  inference(resolution,[],[f1367,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f23855,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16))) != s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X17),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X16) != s(bool,X19) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f1367,f724])).
fof(f724,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f23853,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X2) != s(bool,X5)) )),
  inference(resolution,[],[f1367,f722])).
fof(f722,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f23765,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl146),
  inference(cnf_transformation,[],[f23765_D])).
fof(f23765_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f23722,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | ~p(s(bool,X1))) ) | $spl22),
  inference(resolution,[],[f15783,f9442])).
fof(f9442,plain,(
  ( ! [X2,X0,X1] : (~sP10(X1,X0,X2) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f9434,f792])).
fof(f792,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f552])).
fof(f9434,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X1,X0,X2) | ~sP9(X2,X0,X1) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f9431,f801])).
fof(f801,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP11(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f555])).
fof(f9431,plain,(
  ( ! [X10,X11,X9] : (sP11(X10,X11,X9) | ~p(s(bool,X9)) | ~sP10(X10,X9,X11) | ~sP9(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f9428,f1316])).
fof(f1316,plain,(
  ( ! [X182] : (s(bool,X182) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X182))) )),
  inference(resolution,[],[f679,f862])).
fof(f9428,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | sP11(X10,X11,X9) | ~p(s(bool,X9)) | ~sP10(X10,X9,X11) | ~sP9(X11,X9,X10)) )),
  inference(resolution,[],[f8780,f784])).
fof(f784,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X2)) | ~sP10(X0,X2,X1) | ~sP9(X1,X2,X0) | ~sP8(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f546])).
fof(f8780,plain,(
  ( ! [X202,X200,X201] : (sP8(X200,X201,X202) | s(bool,X200) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1318,f798])).
fof(f1318,plain,(
  ( ! [X184] : (p(s(bool,X184)) | s(bool,X184) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f679,f866])).
fof(f866,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f15783,plain,(
  ( ! [X136,X137] : (sP10(X136,X137,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl22),
  inference(resolution,[],[f1396,f788])).
fof(f788,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f549])).
fof(f23521,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5))) ) | $spl144),
  inference(cnf_transformation,[],[f23521_D])).
fof(f23521_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f23403,plain,(
  ( ! [X613,X615,X611,X612,X614,X610] : (s('type.pair.prod'(bool,X610),'const.pair.,_2'(s(bool,X611),s(X610,X612))) != s('type.pair.prod'(bool,X610),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X610,X613))) | p(s(bool,X614)) | ~p(s(bool,X615)) | s(bool,X611) != s(bool,X615)) )),
  inference(resolution,[],[f1365,f1048])).
fof(f1365,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X19))) != s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X17,X18)))) )),
  inference(superposition,[],[f679,f1041])).
fof(f1041,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) = s(X0,X4) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f673])).
fof(f23308,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,X50),s(X49,X51))) != s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X49,X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f1365,f728])).
fof(f23307,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,X43),s(X42,X44))) != s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X42,X45))) | p(s(bool,X46)) | ~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X43) != s(bool,X47)) )),
  inference(resolution,[],[f1365,f727])).
fof(f23305,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30))) != s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X28,X31))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X29) != s(bool,X33)) )),
  inference(resolution,[],[f1365,f726])).
fof(f23303,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16))) != s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X14,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f1365,f724])).
fof(f23301,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f1365,f722])).
fof(f23272,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X19),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X18))))) ) | $spl140),
  inference(cnf_transformation,[],[f23272_D])).
fof(f23272_D,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X19),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X18))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f23274,plain,(
  s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | $spl143),
  inference(cnf_transformation,[],[f23274_D])).
fof(f23274_D,plain,(
  s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f23270,plain,(
  ( ! [X19,X18] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X19),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X18))))) ) | $spl66),
  inference(trivial_inequality_removal,[],[f23258])).
fof(f23258,plain,(
  ( ! [X19,X18] : (s('type.list.list'('type.conSem.v'),X18) != s('type.list.list'('type.conSem.v'),X18) | s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X19),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X18))))) ) | $spl66),
  inference(superposition,[],[f17471,f997])).
fof(f997,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X1))))) )),
  inference(cnf_transformation,[],[f658])).
fof(f658,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X1)))) | s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1)) & (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X1) | p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X1))))))),
  inference(nnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X1)))) <=> s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X6,X19] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X19)))) <=> s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X19))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.v_rel_eqns|split|4')).
fof(f17471,plain,(
  ( ! [X0] : (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X0) | s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(superposition,[],[f17304,f999])).
fof(f999,plain,(
  ( ! [X0,X1] : (s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1) | s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f659])).
fof(f659,plain,(
  ! [X0,X1] : ((s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))) | s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1)) & (s('type.list.list'('type.conSem.v'),X0) != s('type.list.list'('type.conSem.v'),X1) | s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1)))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))) <=> s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1))),
  inference(rectify,[],[f135])).
fof(f135,axiom,(
  ! [X68,X131] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X68))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X131))) <=> s('type.list.list'('type.conSem.v'),X68) = s('type.list.list'('type.conSem.v'),X131))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.conSem.v_|49||49||split|5')).
fof(f17304,plain,(
  s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)) | $spl66),
  inference(resolution,[],[f15410,f998])).
fof(f998,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),X1)))) | s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X1)) )),
  inference(cnf_transformation,[],[f658])).
fof(f23237,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK40(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl138),
  inference(cnf_transformation,[],[f23237_D])).
fof(f23237_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK40(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f23102,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK40(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1325,f1048])).
fof(f1325,plain,(
  ( ! [X198,X199,X197] : (~p(s(bool,happ(s(fun(X197,bool),X198),s(X197,sK40(X198,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X197))))) | p(s(bool,happ(s(fun(X197,bool),X198),s(X197,X199))))) )),
  inference(resolution,[],[f679,f888])).
fof(f888,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK39(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK39,sK40])],[f596])).
fof(f596,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f595])).
fof(f595,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f594])).
fof(f594,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X33] : ! [X57,X58] : (! [X51] : (p(s(bool,X57)) | p(s(bool,happ(s(fun(X33,bool),X58),s(X33,X51))))) <=> (p(s(bool,X57)) | ! [X51] : p(s(bool,happ(s(fun(X33,bool),X58),s(X33,X51))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f23007,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK40(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1325,f728])).
fof(f23006,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK40(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1325,f727])).
fof(f23004,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK40(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1325,f726])).
fof(f23002,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK40(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1325,f724])).
fof(f23000,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK40(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1325,f722])).
fof(f20569,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | $spl114),
  inference(cnf_transformation,[],[f20569_D])).
fof(f20569_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f22832,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | p(s(bool,X1))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f20458,f9058])).
fof(f20458,plain,(
  ( ! [X140,X141] : (sP8('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))),X140,X141)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f798])).
fof(f17009,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) | ($spl16 | $spl72)),
  inference(forward_demodulation,[],[f16825,f1390])).
fof(f16825,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) | $spl72),
  inference(resolution,[],[f16606,f953])).
fof(f16606,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) | $spl72),
  inference(resolution,[],[f1676,f15903])).
fof(f22786,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | p(s(bool,X3))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f20457,f9057])).
fof(f20457,plain,(
  ( ! [X138,X139] : (sP8(X138,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))),X139)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f797])).
fof(f22784,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK39(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl136),
  inference(cnf_transformation,[],[f22784_D])).
fof(f22784_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK39(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f22649,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK39(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1324,f1048])).
fof(f1324,plain,(
  ( ! [X196,X194,X195] : (~p(s(bool,happ(s(fun(X194,bool),X195),s(X194,sK39(X195,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X194))))) | p(s(bool,happ(s(fun(X194,bool),X195),s(X194,X196))))) )),
  inference(resolution,[],[f679,f886])).
fof(f886,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f597])).
fof(f22554,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK39(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1324,f728])).
fof(f22553,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK39(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1324,f727])).
fof(f22551,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK39(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1324,f726])).
fof(f22549,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK39(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1324,f724])).
fof(f22547,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK39(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1324,f722])).
fof(f22417,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | $spl134),
  inference(cnf_transformation,[],[f22417_D])).
fof(f22417_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f22369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | ~p(s(bool,X1))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f20452,f9442])).
fof(f20452,plain,(
  ( ! [X128,X129] : (sP10(X128,X129,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f788])).
fof(f22274,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK38(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl132),
  inference(cnf_transformation,[],[f22274_D])).
fof(f22274_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK38(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f22139,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK38(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1323,f1048])).
fof(f1323,plain,(
  ( ! [X191,X192,X193] : (~p(s(bool,happ(s(fun(X191,bool),X192),s(X191,sK38(X192,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X191))))) | p(s(bool,happ(s(fun(X191,bool),X192),s(X191,X193))))) )),
  inference(resolution,[],[f679,f883])).
fof(f883,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK38(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f593])).
fof(f593,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK38(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37,sK38])],[f592])).
fof(f592,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X33] : ! [X58,X57] : (! [X51] : (p(s(bool,happ(s(fun(X33,bool),X57),s(X33,X51)))) | p(s(bool,X58))) <=> (! [X51] : p(s(bool,happ(s(fun(X33,bool),X57),s(X33,X51)))) | p(s(bool,X58))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.LEFT_FORALL_OR_THM')).
fof(f22044,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK38(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1323,f728])).
fof(f22043,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK38(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1323,f727])).
fof(f22041,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK38(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1323,f726])).
fof(f22039,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK38(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1323,f724])).
fof(f22037,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK38(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1323,f722])).
fof(f21865,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK37(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl130),
  inference(cnf_transformation,[],[f21865_D])).
fof(f21865_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK37(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f21730,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK37(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1322,f1048])).
fof(f1322,plain,(
  ( ! [X189,X190,X188] : (~p(s(bool,happ(s(fun(X188,bool),X189),s(X188,sK37(X189,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X188))))) | p(s(bool,happ(s(fun(X188,bool),X189),s(X188,X190))))) )),
  inference(resolution,[],[f679,f881])).
fof(f881,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f593])).
fof(f21635,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK37(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1322,f728])).
fof(f21634,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK37(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1322,f727])).
fof(f21632,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK37(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1322,f726])).
fof(f21630,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK37(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1322,f724])).
fof(f21628,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK37(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1322,f722])).
fof(f19589,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | $spl98),
  inference(cnf_transformation,[],[f19589_D])).
fof(f19589_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f21504,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | p(s(bool,X1))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f19478,f9058])).
fof(f19478,plain,(
  ( ! [X140,X141] : (sP8('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))),X140,X141)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f798])).
fof(f17012,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) | ($spl16 | $spl72)),
  inference(forward_demodulation,[],[f15913,f1390])).
fof(f15913,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) | $spl72),
  inference(resolution,[],[f15903,f953])).
fof(f21458,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | p(s(bool,X3))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f19477,f9057])).
fof(f19477,plain,(
  ( ! [X138,X139] : (sP8(X138,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))),X139)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f797])).
fof(f21359,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | $spl128),
  inference(cnf_transformation,[],[f21359_D])).
fof(f21359_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f21311,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | ~p(s(bool,X1))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f19472,f9442])).
fof(f19472,plain,(
  ( ! [X128,X129] : (sP10(X128,X129,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f788])).
fof(f20592,plain,(
  sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | $spl116),
  inference(cnf_transformation,[],[f20592_D])).
fof(f20592_D,plain,(
  sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f20598,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | $spl122),
  inference(cnf_transformation,[],[f20598_D])).
fof(f20598_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f20602,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22))))) | $spl126),
  inference(cnf_transformation,[],[f20602_D])).
fof(f20602_D,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f20575,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f943])).
fof(f943,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | sP17(X0,X1,X2) | (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0))))) | (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0)))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0))))))) & ((~sP17(X0,X1,X2) & ! [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) & ! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))) | sP18(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK63,sK64,sK65])],[f636])).
fof(f636,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | sP17(X0,X1,X2) | ? [X8] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X8))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X8)))) | ? [X6,X7] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X6))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X7))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X6),s('type.list.list'('type.conSem.v'),X7)))))) & ((~sP17(X0,X1,X2) & ! [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) & ! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))) | sP18(X0,X1,X2)))),
  inference(rectify,[],[f635])).
fof(f635,plain,(
  ! [X0,X2,X1] : ((~sP18(X0,X2,X1) | sP17(X0,X2,X1) | ? [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) | ? [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))) & ((~sP17(X0,X2,X1) & ! [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) & ! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))) | sP18(X0,X2,X1)))),
  inference(flattening,[],[f634])).
fof(f634,plain,(
  ! [X0,X2,X1] : ((~sP18(X0,X2,X1) | (sP17(X0,X2,X1) | ? [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) | ? [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4))))))) & ((~sP17(X0,X2,X1) & ! [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) & ! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))) | sP18(X0,X2,X1)))),
  inference(nnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0,X2,X1] : (sP18(X0,X2,X1) <=> (sP17(X0,X2,X1) | ? [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) | ? [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f20392,plain,(
  sP18(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f951])).
fof(f951,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X2,X1) | ~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2))))) )),
  inference(cnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) | sP18(X0,X2,X1)) & (~sP18(X0,X2,X1) | p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2))))))),
  inference(nnf_transformation,[],[f507])).
fof(f507,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) <=> sP18(X0,X2,X1))),
  inference(definition_folding,[],[f338,f506,f505])).
fof(f505,plain,(
  ! [X0,X2,X1] : (sP17(X0,X2,X1) <=> ? [X6,X7] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X6))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X7))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X6),s('type.conSem.v',X7))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f338,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),X1),s('type.semanticPrimitives.store_v'('type.conSem.v'),X2)))) <=> (? [X6,X7] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X6))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X7))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X6),s('type.conSem.v',X7))))) | ? [X5] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X5)))) | ? [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X3))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X4))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),X3),s('type.list.list'('type.conSem.v'),X4)))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X26,X27,X28] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X26),s('type.semanticPrimitives.store_v'('type.modSem.v'),X27),s('type.semanticPrimitives.store_v'('type.conSem.v'),X28)))) <=> (? [X16,X17] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X27) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X16))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X28) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X17))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X26),s('type.modSem.v',X16),s('type.conSem.v',X17))))) | ? [X18] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X27) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X18))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X28) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X18)))) | ? [X19,X20] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X27) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),X19))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X28) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),X20))) & p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X26),s('type.list.list'('type.modSem.v'),X19),s('type.list.list'('type.conSem.v'),X20)))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.sv_rel_cases')).
fof(f20594,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | $spl118),
  inference(cnf_transformation,[],[f20594_D])).
fof(f20594_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f20574,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f942])).
fof(f942,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f20600,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | $spl124),
  inference(cnf_transformation,[],[f20600_D])).
fof(f20600_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f20573,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f941])).
fof(f941,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f20572,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f940])).
fof(f940,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(X2,X1,X0)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f20596,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | $spl120),
  inference(cnf_transformation,[],[f20596_D])).
fof(f20596_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f20571,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f939])).
fof(f939,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f20570,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK22)))) | sP17(sK22,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f20392,f938])).
fof(f938,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(X2,X1,X0)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(X2,X1,X0)))) | sP17(X0,X1,X2) | ~sP18(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f637])).
fof(f20497,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f1048])).
fof(f20401,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | s(bool,X26) != s(bool,X27)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f728])).
fof(f20398,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f727])).
fof(f20397,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f726])).
fof(f20396,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | s(bool,X11) != s(bool,X12)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f724])).
fof(f20394,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17009,f722])).
fof(f20093,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | s(bool,X2) != s(bool,X5) | p(s(bool,X5))) ) | $spl112),
  inference(cnf_transformation,[],[f20093_D])).
fof(f20093_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | s(bool,X2) != s(bool,X5) | p(s(bool,X5))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f19960,plain,(
  ( ! [X621,X617,X619,X620,X616,X618] : (s('type.pair.prod'(X616,bool),'const.pair.,_2'(s(X616,X617),s(bool,X618))) != s('type.pair.prod'(X616,bool),'const.pair.,_2'(s(X616,X619),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X620)) | p(s(bool,X621)) | s(bool,X618) != s(bool,X621)) )),
  inference(resolution,[],[f1208,f1048])).
fof(f1208,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X26),s(bool,X24))) != s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X27),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))))) )),
  inference(superposition,[],[f677,f1042])).
fof(f19864,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X50),s(bool,X51))) != s('type.pair.prod'(X49,bool),'const.pair.,_2'(s(X49,X52),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X51) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1208,f728])).
fof(f19861,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X29),s(bool,X30))) != s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X31),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X30) != s(bool,X34)) )),
  inference(resolution,[],[f1208,f727])).
fof(f19860,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.pair.prod'(X21,bool),'const.pair.,_2'(s(X21,X22),s(bool,X23))) != s('type.pair.prod'(X21,bool),'const.pair.,_2'(s(X21,X24),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X23) != s(bool,X26)) )),
  inference(resolution,[],[f1208,f726])).
fof(f19859,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16))) != s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X17),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f1208,f724])).
fof(f19857,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X3),s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X2) != s(bool,X5)) )),
  inference(resolution,[],[f1208,f722])).
fof(f19612,plain,(
  sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | $spl100),
  inference(cnf_transformation,[],[f19612_D])).
fof(f19612_D,plain,(
  sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f19618,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | $spl106),
  inference(cnf_transformation,[],[f19618_D])).
fof(f19618_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f19622,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19))))) | $spl110),
  inference(cnf_transformation,[],[f19622_D])).
fof(f19622_D,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f19595,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f943])).
fof(f19412,plain,(
  sP18(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f951])).
fof(f19614,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | $spl102),
  inference(cnf_transformation,[],[f19614_D])).
fof(f19614_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f19594,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f942])).
fof(f19620,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | $spl108),
  inference(cnf_transformation,[],[f19620_D])).
fof(f19620_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f19593,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f941])).
fof(f19592,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f940])).
fof(f19616,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | $spl104),
  inference(cnf_transformation,[],[f19616_D])).
fof(f19616_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f19591,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f939])).
fof(f19590,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK19)))) | sP17(sK19,'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),sK20) | ($spl16 | $spl72)),
  inference(resolution,[],[f19412,f938])).
fof(f19517,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f1048])).
fof(f19421,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | s(bool,X26) != s(bool,X27)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f728])).
fof(f19418,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f727])).
fof(f19417,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f726])).
fof(f19416,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))) | s(bool,X11) != s(bool,X12)) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f724])).
fof(f19414,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))))) ) | ($spl16 | $spl72)),
  inference(resolution,[],[f17012,f722])).
fof(f19320,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | s(bool,X1) != s(bool,X5) | p(s(bool,X5))) ) | $spl96),
  inference(cnf_transformation,[],[f19320_D])).
fof(f19320_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | s(bool,X1) != s(bool,X5) | p(s(bool,X5))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f19187,plain,(
  ( ! [X621,X617,X619,X620,X616,X618] : (s('type.pair.prod'(bool,X616),'const.pair.,_2'(s(bool,X617),s(X616,X618))) != s('type.pair.prod'(bool,X616),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X616,X619))) | p(s(bool,X620)) | p(s(bool,X621)) | s(bool,X617) != s(bool,X621)) )),
  inference(resolution,[],[f1206,f1048])).
fof(f1206,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X19))) != s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X17,X18)))) )),
  inference(superposition,[],[f677,f1041])).
fof(f19091,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,X50),s(X49,X51))) != s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X49,X52))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1206,f728])).
fof(f19088,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30))) != s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X28,X31))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f1206,f727])).
fof(f19087,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X22),s(X21,X23))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X21,X24))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f1206,f726])).
fof(f19086,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16))) != s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X14,X17))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f1206,f724])).
fof(f19084,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))),s(X0,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f1206,f722])).
fof(f17013,plain,(
  ( ! [X217] : (p(s(bool,X217)) | s(bool,X217) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl80),
  inference(cnf_transformation,[],[f17013_D])).
fof(f17013_D,plain,(
  ( ! [X217] : (p(s(bool,X217)) | s(bool,X217) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f19040,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | p(s(bool,X1))) ) | $spl72),
  inference(resolution,[],[f16893,f9058])).
fof(f16893,plain,(
  ( ! [X149,X148] : (sP8('const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),X148,X149)) ) | $spl72),
  inference(resolution,[],[f16606,f798])).
fof(f18998,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | p(s(bool,X3))) ) | $spl72),
  inference(resolution,[],[f16892,f9057])).
fof(f16892,plain,(
  ( ! [X146,X147] : (sP8(X146,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),X147)) ) | $spl72),
  inference(resolution,[],[f16606,f797])).
fof(f18907,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f18907_D])).
fof(f18907_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f18863,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | ~p(s(bool,X1))) ) | $spl72),
  inference(resolution,[],[f16887,f9442])).
fof(f16887,plain,(
  ( ! [X136,X137] : (sP10(X136,X137,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f16606,f788])).
fof(f18862,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK47(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl92),
  inference(cnf_transformation,[],[f18862_D])).
fof(f18862_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK47(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f18729,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK47(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1164,f1048])).
fof(f1164,plain,(
  ( ! [X206,X207,X205] : (~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,sK47(X206,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X205))))) | p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X207))))) )),
  inference(resolution,[],[f677,f900])).
fof(f900,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK47(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f609])).
fof(f609,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK46(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK47(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46,sK47])],[f608])).
fof(f608,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X33] : ! [X57,X58] : ((p(s(bool,X57)) & ! [X51] : p(s(bool,happ(s(fun(X33,bool),X58),s(X33,X51))))) <=> ! [X51] : (p(s(bool,X57)) & p(s(bool,happ(s(fun(X33,bool),X58),s(X33,X51))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f18634,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK47(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1164,f728])).
fof(f18633,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK47(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1164,f727])).
fof(f18631,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK47(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1164,f726])).
fof(f18629,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK47(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1164,f724])).
fof(f18627,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK47(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1164,f722])).
fof(f17470,plain,(
  ( ! [X5,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X5),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl84),
  inference(cnf_transformation,[],[f17470_D])).
fof(f17470_D,plain,(
  ( ! [X5,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X5),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f18376,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X2),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(resolution,[],[f17367,f9443])).
fof(f9443,plain,(
  ( ! [X4,X5,X3] : (~sP9(X5,X3,X4) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f9435,f790])).
fof(f9435,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP10(X4,X3,X5) | ~sP9(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f9431,f802])).
fof(f17367,plain,(
  ( ! [X194,X192,X193] : (sP9('const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X192),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))),X193,X194)) ) | $spl66),
  inference(resolution,[],[f15410,f794])).
fof(f18375,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK46(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) | $spl90),
  inference(cnf_transformation,[],[f18375_D])).
fof(f18375_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK46(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133)))) | ~p(s(bool,X137))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f18242,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK46(X642,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X641))))) )),
  inference(resolution,[],[f1163,f1048])).
fof(f1163,plain,(
  ( ! [X204,X202,X203] : (~p(s(bool,happ(s(fun(X202,bool),X203),s(X202,sK46(X203,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X202))))) | p(s(bool,happ(s(fun(X202,bool),X203),s(X202,X204))))) )),
  inference(resolution,[],[f677,f898])).
fof(f898,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK46(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f609])).
fof(f18147,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK46(X176,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1163,f728])).
fof(f18146,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK46(X170,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X169))))) )),
  inference(resolution,[],[f1163,f727])).
fof(f18144,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK46(X158,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X157))))) )),
  inference(resolution,[],[f1163,f726])).
fof(f18142,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK46(X146,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1163,f724])).
fof(f18140,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK46(X134,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133))))) )),
  inference(resolution,[],[f1163,f722])).
fof(f18015,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X4),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl88),
  inference(cnf_transformation,[],[f18015_D])).
fof(f18015_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X4),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f17972,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X4),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | ~p(s(bool,X5))) ) | $spl66),
  inference(resolution,[],[f17363,f9156])).
fof(f9156,plain,(
  ( ! [X2,X0,X1] : (~sP10(X1,X2,X0) | p(s(bool,X0)) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f9148,f796])).
fof(f796,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f554])).
fof(f9148,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP10(X1,X2,X0) | ~sP8(X2,X0,X1) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f9146,f801])).
fof(f9146,plain,(
  ( ! [X10,X11,X9] : (sP11(X10,X9,X11) | p(s(bool,X9)) | ~sP10(X10,X11,X9) | ~sP8(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f9143,f1155])).
fof(f9143,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | sP11(X10,X9,X11) | p(s(bool,X9)) | ~sP10(X10,X11,X9) | ~sP8(X11,X9,X10)) )),
  inference(resolution,[],[f8355,f785])).
fof(f785,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X1)) | ~sP10(X0,X2,X1) | ~sP9(X1,X2,X0) | ~sP8(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f546])).
fof(f8355,plain,(
  ( ! [X191,X192,X193] : (sP9(X191,X192,X193) | s(bool,X191) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f794])).
fof(f17363,plain,(
  ( ! [X182,X180,X181] : (sP10(X180,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X181),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))),X182)) ) | $spl66),
  inference(resolution,[],[f15410,f789])).
fof(f789,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f549])).
fof(f17881,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133)))) | ~p(s(bool,X137))) ) | $spl86),
  inference(cnf_transformation,[],[f17881_D])).
fof(f17881_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133)))) | ~p(s(bool,X137))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f17748,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X642,X641))))) )),
  inference(resolution,[],[f1162,f1048])).
fof(f1162,plain,(
  ( ! [X200,X199,X201] : (~p(s(bool,happ(s(fun(X199,bool),X200),s(X199,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X200,X199))))) | p(s(bool,happ(s(fun(X199,bool),X200),s(X199,X201))))) )),
  inference(resolution,[],[f677,f895])).
fof(f895,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK44,sK45])],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f602])).
fof(f602,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X33] : ! [X57,X58] : ((! [X51] : p(s(bool,happ(s(fun(X33,bool),X57),s(X33,X51)))) & p(s(bool,X58))) <=> ! [X51] : (p(s(bool,happ(s(fun(X33,bool),X57),s(X33,X51)))) & p(s(bool,X58))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.bool.LEFT_AND_FORALL_THM')).
fof(f17653,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X176,X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1162,f728])).
fof(f17652,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X170,X169))))) )),
  inference(resolution,[],[f1162,f727])).
fof(f17650,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X158,X157))))) )),
  inference(resolution,[],[f1162,f726])).
fof(f17648,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X146,X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1162,f724])).
fof(f17646,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK45('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133))))) )),
  inference(resolution,[],[f1162,f722])).
fof(f17408,plain,(
  ( ! [X304,X306,X305] : (p(s(bool,X304)) | ~p(s(bool,X305)) | s(bool,X305) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X306),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(resolution,[],[f15410,f1048])).
fof(f17313,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X31),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s(bool,X29) != s(bool,X32)) ) | $spl66),
  inference(resolution,[],[f15410,f728])).
fof(f17312,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X28),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(resolution,[],[f15410,f727])).
fof(f17310,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X21),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(resolution,[],[f15410,f726])).
fof(f17308,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X12),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s(bool,X10) != s(bool,X13)) ) | $spl66),
  inference(resolution,[],[f15410,f724])).
fof(f17306,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X5),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl66),
  inference(resolution,[],[f15410,f722])).
fof(f17303,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133)))) | ~p(s(bool,X137))) ) | $spl82),
  inference(cnf_transformation,[],[f17303_D])).
fof(f17303_D,plain,(
  ( ! [X134,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133)))) | ~p(s(bool,X137))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f17170,plain,(
  ( ! [X644,X642,X645,X641,X643] : (p(s(bool,happ(s(fun(X641,bool),X642),s(X641,X643)))) | p(s(bool,X644)) | ~p(s(bool,X645)) | s(bool,X645) != s(bool,happ(s(fun(X641,bool),X642),s(X641,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X642,X641))))) )),
  inference(resolution,[],[f1161,f1048])).
fof(f1161,plain,(
  ( ! [X198,X196,X197] : (~p(s(bool,happ(s(fun(X196,bool),X197),s(X196,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X197,X196))))) | p(s(bool,happ(s(fun(X196,bool),X197),s(X196,X198))))) )),
  inference(resolution,[],[f677,f893])).
fof(f893,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f605])).
fof(f17075,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X175,bool),X176),s(X175,X177)))) | ~p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,happ(s(fun(X175,bool),X176),s(X175,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X176,X175)))) | s(bool,X178) != s(bool,X180)) )),
  inference(resolution,[],[f1161,f728])).
fof(f17074,plain,(
  ( ! [X173,X171,X169,X174,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | p(s(bool,X172)) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,happ(s(fun(X169,bool),X170),s(X169,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X170,X169))))) )),
  inference(resolution,[],[f1161,f727])).
fof(f17072,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X160) != s(bool,X162) | s(bool,X161) != s(bool,happ(s(fun(X157,bool),X158),s(X157,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X158,X157))))) )),
  inference(resolution,[],[f1161,f726])).
fof(f17070,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(X145,bool),X146),s(X145,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X146,X145)))) | s(bool,X148) != s(bool,X150)) )),
  inference(resolution,[],[f1161,f724])).
fof(f17068,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,happ(s(fun(X133,bool),X134),s(X133,X135)))) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,X138) | s(bool,X137) != s(bool,happ(s(fun(X133,bool),X134),s(X133,sK44('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X134,X133))))) )),
  inference(resolution,[],[f1161,f722])).
fof(f16829,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f16606,f722])).
fof(f16831,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | s(bool,X19) != s(bool,X20)) ) | $spl72),
  inference(resolution,[],[f16606,f724])).
fof(f16832,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f16606,f726])).
fof(f16833,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f16606,f727])).
fof(f16836,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | s(bool,X34) != s(bool,X35)) ) | $spl72),
  inference(resolution,[],[f16606,f728])).
fof(f16932,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f16606,f1048])).
fof(f17002,plain,(
  $false | ($spl16 | $spl68 | $spl72)),
  inference(subsumption_resolution,[],[f17001,f679])).
fof(f17001,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | ($spl16 | $spl68 | $spl72)),
  inference(forward_demodulation,[],[f17000,f1390])).
fof(f17000,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | ($spl68 | $spl72)),
  inference(forward_demodulation,[],[f16825,f15412])).
fof(f16782,plain,(
  ( ! [X21,X19,X18] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19)))) | s(bool,X21) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19))) | ~p(s(bool,X21))) ) | $spl78),
  inference(cnf_transformation,[],[f16782_D])).
fof(f16782_D,plain,(
  ( ! [X21,X19,X18] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19)))) | s(bool,X21) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19))) | ~p(s(bool,X21))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f16716,plain,(
  ( ! [X424,X426,X425,X427] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X424),s('type.list.list'('type.conSem.v'),X425)))) | p(s(bool,X426)) | ~p(s(bool,X427)) | s(bool,X427) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X424),s('type.list.list'('type.conSem.v'),X425)))) )),
  inference(resolution,[],[f1676,f1048])).
fof(f16621,plain,(
  ( ! [X57,X54,X56,X55,X53] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X53),s('type.list.list'('type.conSem.v'),X54)))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X53),s('type.list.list'('type.conSem.v'),X54))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f1676,f728])).
fof(f16620,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X48),s('type.list.list'('type.conSem.v'),X49)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X48),s('type.list.list'('type.conSem.v'),X49)))) )),
  inference(resolution,[],[f1676,f727])).
fof(f16618,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X38),s('type.list.list'('type.conSem.v'),X39)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X38),s('type.list.list'('type.conSem.v'),X39)))) )),
  inference(resolution,[],[f1676,f726])).
fof(f16616,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X28),s('type.list.list'('type.conSem.v'),X29)))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X28),s('type.list.list'('type.conSem.v'),X29))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f1676,f724])).
fof(f16614,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19)))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),X18),s('type.list.list'('type.conSem.v'),X19)))) )),
  inference(resolution,[],[f1676,f722])).
fof(f16094,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f16094_D])).
fof(f16094_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f16521,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | p(s(bool,X1))) ) | $spl72),
  inference(resolution,[],[f15981,f9058])).
fof(f15981,plain,(
  ( ! [X149,X148] : (sP8('const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),X148,X149)) ) | $spl72),
  inference(resolution,[],[f15903,f798])).
fof(f16479,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | p(s(bool,X3))) ) | $spl72),
  inference(resolution,[],[f15980,f9057])).
fof(f15980,plain,(
  ( ! [X146,X147] : (sP8(X146,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')),X147)) ) | $spl72),
  inference(resolution,[],[f15903,f797])).
fof(f16351,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f16351_D])).
fof(f16351_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f16307,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | ~p(s(bool,X1))) ) | $spl72),
  inference(resolution,[],[f15975,f9442])).
fof(f15975,plain,(
  ( ! [X136,X137] : (sP10(X136,X137,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f15903,f788])).
fof(f16020,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f15903,f1048])).
fof(f15924,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | s(bool,X34) != s(bool,X35)) ) | $spl72),
  inference(resolution,[],[f15903,f728])).
fof(f15921,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f15903,f727])).
fof(f15920,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f15903,f726])).
fof(f15919,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | s(bool,X19) != s(bool,X20)) ) | $spl72),
  inference(resolution,[],[f15903,f724])).
fof(f15917,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) ) | $spl72),
  inference(resolution,[],[f15903,f722])).
fof(f15410,plain,(
  ( ! [X2] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X2),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))))) ) | $spl66),
  inference(cnf_transformation,[],[f15410_D])).
fof(f15410_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X2),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f15903,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) | $spl72),
  inference(cnf_transformation,[],[f15903_D])).
fof(f15903_D,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f15860,plain,(
  ( ! [X2] : (p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X2),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))))) ) | $spl22),
  inference(superposition,[],[f1396,f997])).
fof(f15828,plain,(
  ( ! [X216,X217] : (p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl22),
  inference(resolution,[],[f1396,f1048])).
fof(f15732,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s(bool,X34) != s(bool,X35)) ) | $spl22),
  inference(resolution,[],[f1396,f728])).
fof(f15729,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl22),
  inference(resolution,[],[f1396,f727])).
fof(f15728,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl22),
  inference(resolution,[],[f1396,f726])).
fof(f15727,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s(bool,X19) != s(bool,X20)) ) | $spl22),
  inference(resolution,[],[f1396,f724])).
fof(f15725,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) ) | $spl22),
  inference(resolution,[],[f1396,f722])).
fof(f15412,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | $spl68),
  inference(cnf_transformation,[],[f15412_D])).
fof(f15412_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f15283,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))) | ~p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X2),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))))) ) | $spl20),
  inference(superposition,[],[f1394,f997])).
fof(f1386,plain,(
  sP17(sK19,sK21,sK20) | $spl12),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  sP17(sK19,sK21,sK20) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1390,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | $spl16),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f15070,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | $spl15),
  inference(subsumption_resolution,[],[f6735,f1060])).
fof(f1060,plain,(
  sP18(sK19,sK21,sK20)),
  inference(resolution,[],[f677,f951])).
fof(f6735,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(trivial_inequality_removal,[],[f6722])).
fof(f6722,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) != s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(superposition,[],[f1387,f938])).
fof(f1387,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl15),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1394,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | $spl20),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f15069,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | $spl15),
  inference(subsumption_resolution,[],[f6734,f1060])).
fof(f6734,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(trivial_inequality_removal,[],[f6723])).
fof(f6723,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) != s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(superposition,[],[f1387,f940])).
fof(f1396,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | $spl22),
  inference(cnf_transformation,[],[f1396_D])).
fof(f1396_D,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f15068,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | sP17(sK19,sK21,sK20) | $spl15),
  inference(subsumption_resolution,[],[f6733,f1060])).
fof(f6733,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(trivial_inequality_removal,[],[f6724])).
fof(f6724,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) != s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) | p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | sP17(sK19,sK21,sK20) | ~sP18(sK19,sK21,sK20) | $spl15),
  inference(superposition,[],[f1387,f942])).
fof(f15063,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f15062,f679])).
fof(f15062,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | $spl12),
  inference(forward_demodulation,[],[f15061,f6296])).
fof(f6296,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',sK66(sK20,sK21,sK19)))) | $spl12),
  inference(resolution,[],[f1386,f947])).
fof(f947,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',sK66(X2,X1,X0)))) | ~sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',sK66(X2,X1,X0)))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',sK67(X2,X1,X0)))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',sK66(X2,X1,X0)),s('type.conSem.v',sK67(X2,X1,X0))))))) & (! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X4))) | ~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X3),s('type.conSem.v',X4))))) | sP17(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK66,sK67])],[f639])).
fof(f639,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | ? [X5,X6] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X5))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X6))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X5),s('type.conSem.v',X6)))))) & (! [X3,X4] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X3))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X4))) | ~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X3),s('type.conSem.v',X4))))) | sP17(X0,X1,X2)))),
  inference(rectify,[],[f638])).
fof(f638,plain,(
  ! [X0,X2,X1] : ((~sP17(X0,X2,X1) | ? [X6,X7] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X6))) & s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X7))) & p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X6),s('type.conSem.v',X7)))))) & (! [X6,X7] : (s('type.semanticPrimitives.store_v'('type.modSem.v'),X1) != s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X6))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),X2) != s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X7))) | ~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X6),s('type.conSem.v',X7))))) | sP17(X0,X2,X1)))),
  inference(nnf_transformation,[],[f505])).
fof(f15061,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',sK66(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) | $spl12),
  inference(forward_demodulation,[],[f14894,f6297])).
fof(f6297,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',sK67(sK20,sK21,sK19)))) | $spl12),
  inference(resolution,[],[f1386,f948])).
fof(f948,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.store_v'('type.conSem.v'),X1) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',sK67(X2,X1,X0)))) | ~sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f640])).
fof(f14894,plain,(
  p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',sK66(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',sK67(sK20,sK21,sK19))))))) | $spl12),
  inference(resolution,[],[f14698,f954])).
fof(f954,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X2)))))) | ~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2))))) )),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2)))) | p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X2)))))))),
  inference(ennf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2)))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X2)))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X25,X16,X17] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X25),s('type.modSem.v',X16),s('type.conSem.v',X17)))) => p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X25),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Refv_1'(s('type.modSem.v',X16))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Refv_1'(s('type.conSem.v',X17)))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.sv_rel_rules|split|0')).
fof(f14698,plain,(
  p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) | $spl12),
  inference(resolution,[],[f1675,f6298])).
fof(f6298,plain,(
  p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) | $spl12),
  inference(resolution,[],[f1386,f949])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',sK66(X2,X1,X0)),s('type.conSem.v',sK67(X2,X1,X0))))) | ~sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f640])).
fof(f1675,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X0),s('type.conSem.v',X1)))) | p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X0),s('type.conSem.v',X1))))) )),
  inference(resolution,[],[f678,f684])).
fof(f684,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.modSem.v',X1),s('type.conSem.v',X2)))) | ~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) | ~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2)))) | ! [X3] : (~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) | p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.modSem.v',X1),s('type.conSem.v',X2))))))),
  inference(ennf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.modSem.v',X1),s('type.conSem.v',X2)))) => ! [X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3)))) => p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X3),s('type.modSem.v',X1),s('type.conSem.v',X2))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X16,X32] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.modSem.v',X16),s('type.conSem.v',X32)))) => ! [X9] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X6),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9)))) => p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X9),s('type.modSem.v',X16),s('type.conSem.v',X32))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.mod_to_conProof_37_prove_ltb|split|0')).
fof(f14892,plain,(
  ( ! [X17,X20,X18] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X17),s('type.conSem.v',X18)))) | s(bool,X20) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X17),s('type.conSem.v',X18))) | ~p(s(bool,X20))) ) | $spl64),
  inference(cnf_transformation,[],[f14892_D])).
fof(f14892_D,plain,(
  ( ! [X17,X20,X18] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X17),s('type.conSem.v',X18)))) | s(bool,X20) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X17),s('type.conSem.v',X18))) | ~p(s(bool,X20))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14808,plain,(
  ( ! [X424,X426,X423,X425] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X423),s('type.conSem.v',X424)))) | p(s(bool,X425)) | ~p(s(bool,X426)) | s(bool,X426) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X423),s('type.conSem.v',X424)))) )),
  inference(resolution,[],[f1675,f1048])).
fof(f14713,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X52),s('type.conSem.v',X53)))) | ~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X52),s('type.conSem.v',X53))) | s(bool,X54) != s(bool,X56)) )),
  inference(resolution,[],[f1675,f728])).
fof(f14712,plain,(
  ( ! [X47,X50,X48,X51,X49] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X47),s('type.conSem.v',X48)))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X47),s('type.conSem.v',X48)))) )),
  inference(resolution,[],[f1675,f727])).
fof(f14710,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X37),s('type.conSem.v',X38)))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X37),s('type.conSem.v',X38)))) )),
  inference(resolution,[],[f1675,f726])).
fof(f14708,plain,(
  ( ! [X30,X28,X31,X29,X27] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X27),s('type.conSem.v',X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X27),s('type.conSem.v',X28))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1675,f724])).
fof(f14706,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.modSem.v',X17),s('type.conSem.v',X18)))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',X17),s('type.conSem.v',X18)))) )),
  inference(resolution,[],[f1675,f722])).
fof(f14697,plain,(
  ( ! [X0,X3,X1] : (sP14(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) | $spl62),
  inference(cnf_transformation,[],[f14697_D])).
fof(f14697_D,plain,(
  ( ! [X0,X3,X1] : (sP14(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f14585,plain,(
  ( ! [X408,X407,X409,X406] : (sP14(X406,X407) | s(bool,X406) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X408)) | ~p(s(bool,X409)) | s(bool,X407) != s(bool,X409)) )),
  inference(resolution,[],[f8380,f1048])).
fof(f8380,plain,(
  ( ! [X263,X264] : (~p(s(bool,X264)) | sP14(X263,X264) | s(bool,X263) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f848])).
fof(f848,plain,(
  ( ! [X0,X1] : (sP14(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1)))),
  inference(rectify,[],[f575])).
fof(f575,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP14(X1,X0)))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP14(X1,X0)))),
  inference(nnf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X1,X0] : (sP14(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f14490,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP14(X35,X36) | s(bool,X35) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f8380,f728])).
fof(f14489,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP14(X30,X31) | s(bool,X30) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f8380,f727])).
fof(f14487,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP14(X20,X21) | s(bool,X20) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f8380,f726])).
fof(f14485,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP14(X10,X11) | s(bool,X10) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f8380,f724])).
fof(f14483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP14(X0,X1) | s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f8380,f722])).
fof(f14482,plain,(
  ( ! [X0,X3,X1] : (sP14(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,X1) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) | $spl60),
  inference(cnf_transformation,[],[f14482_D])).
fof(f14482_D,plain,(
  ( ! [X0,X3,X1] : (sP14(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | s(bool,X1) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f14370,plain,(
  ( ! [X408,X407,X409,X406] : (sP14(X406,X407) | s(bool,X407) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X408)) | ~p(s(bool,X409)) | s(bool,X406) != s(bool,X409)) )),
  inference(resolution,[],[f8379,f1048])).
fof(f8379,plain,(
  ( ! [X261,X262] : (~p(s(bool,X262)) | sP14(X262,X261) | s(bool,X261) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f845])).
fof(f845,plain,(
  ( ! [X0,X1] : (sP14(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f576])).
fof(f14275,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP14(X35,X36) | s(bool,X36) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f8379,f728])).
fof(f14274,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP14(X30,X31) | s(bool,X31) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X30) != s(bool,X33)) )),
  inference(resolution,[],[f8379,f727])).
fof(f14272,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP14(X20,X21) | s(bool,X21) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) )),
  inference(resolution,[],[f8379,f726])).
fof(f14270,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP14(X10,X11) | s(bool,X11) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f8379,f724])).
fof(f14268,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP14(X0,X1) | s(bool,X1) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f8379,f722])).
fof(f13514,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | $spl52),
  inference(cnf_transformation,[],[f13514_D])).
fof(f13514_D,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f14262,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(duplicate_literal_removal,[],[f14256])).
fof(f14256,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X12))) )),
  inference(resolution,[],[f8378,f3788])).
fof(f3788,plain,(
  ( ! [X181] : (sP14(X181,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X181))) )),
  inference(resolution,[],[f1844,f848])).
fof(f1844,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))))),
  inference(resolution,[],[f1680,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2,X3] : ((~sP1(X0,X1,X2,X3) | (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))) & ! [X11] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | ! [X12] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X12),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X12),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))))) & sP0(X0,X1,X2) & ! [X9,X10] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | s('type.list.list'('type.string.char'),X10) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9)))))) & (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))) | (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),sK25(X3,X2,X1,X0)),s('type.semanticPrimitives.tid_or_exn',sK24(X3,X2,X1,X0)))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),sK26(X3,X2,X1,X0)),s('type.semanticPrimitives.tid_or_exn',sK24(X3,X2,X1,X0)))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),sK26(X3,X2,X1,X0)),s('type.semanticPrimitives.tid_or_exn',sK24(X3,X2,X1,X0)))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))))) | ~sP0(X0,X1,X2) | (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),sK28(X3,X2,X1,X0)),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK27(X3,X2,X1,X0)))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & s('type.list.list'('type.string.char'),sK28(X3,X2,X1,X0)) != s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK27(X3,X2,X1,X0))))) | sP1(X0,X1,X2,X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25,sK26,sK27,sK28])],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2,X3] : ((~sP1(X0,X1,X2,X3) | (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))) & ! [X11] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | ! [X12] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X12),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X12),s('type.semanticPrimitives.tid_or_exn',X11))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))))) & sP0(X0,X1,X2) & ! [X9,X10] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | s('type.list.list'('type.string.char'),X10) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9)))))) & (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))) | ? [X6] : (? [X8] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X8),s('type.semanticPrimitives.tid_or_exn',X6))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & ? [X7] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X7),s('type.semanticPrimitives.tid_or_exn',X6))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X7),s('type.semanticPrimitives.tid_or_exn',X6))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))))) | ~sP0(X0,X1,X2) | ? [X4,X5] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & s('type.list.list'('type.string.char'),X5) != s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4)))) | sP1(X0,X1,X2,X3)))),
  inference(rectify,[],[f512])).
fof(f512,plain,(
  ! [X0,X1,X3,X2] : ((~sP1(X0,X1,X3,X2) | (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | ! [X14] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & sP0(X0,X1,X3) & ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4)))))) & (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) | ? [X12] : (? [X13] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & ? [X14] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) | ~sP0(X0,X1,X3) | ? [X4,X5] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & s('type.list.list'('type.string.char'),X5) != s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4)))) | sP1(X0,X1,X3,X2)))),
  inference(flattening,[],[f511])).
fof(f511,plain,(
  ! [X0,X1,X3,X2] : ((~sP1(X0,X1,X3,X2) | (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | ! [X14] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & sP0(X0,X1,X3) & ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4)))))) & ((~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) | ? [X12] : (? [X13] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) & ? [X14] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) | ~sP0(X0,X1,X3) | ? [X4,X5] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) & s('type.list.list'('type.string.char'),X5) != s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))) | sP1(X0,X1,X3,X2)))),
  inference(nnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1,X3,X2] : (sP1(X0,X1,X3,X2) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | ! [X14] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & sP0(X0,X1,X3) & ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1680,plain,(
  sP1('type.num.num','type.num.num',sK22,sK19)),
  inference(resolution,[],[f678,f703])).
fof(f703,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X3,X2) | ~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3))))) )),
  inference(cnf_transformation,[],[f518])).
fof(f518,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) | sP1(X0,X1,X3,X2)) & (~sP1(X0,X1,X3,X2) | p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3))))))),
  inference(nnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) <=> sP1(X0,X1,X3,X2))),
  inference(definition_folding,[],[f445,f480,f479])).
fof(f479,plain,(
  ! [X0,X1,X3] : (sP0(X0,X1,X3) <=> ! [X6,X7,X8,X9,X10,X11] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X6),s('type.semanticPrimitives.tid_or_exn',X7)))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.semanticPrimitives.tid_or_exn',X6))))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn',X7))))) | (s('type.list.list'('type.string.char'),X9) = s('type.list.list'('type.string.char'),X10) & s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn',X7))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f445,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | ! [X14] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & ! [X6,X7,X8,X9,X10,X11] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X6),s('type.semanticPrimitives.tid_or_exn',X7)))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.semanticPrimitives.tid_or_exn',X6))))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn',X7))))) | (s('type.list.list'('type.string.char'),X9) = s('type.list.list'('type.string.char'),X10) & s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn',X7))) & ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))))),
  inference(flattening,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (! [X13] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) | ! [X14] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & ! [X6,X7,X8,X9,X10,X11] : ((~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X6),s('type.semanticPrimitives.tid_or_exn',X7)))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.semanticPrimitives.tid_or_exn',X6))))) | s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) != s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn',X7)))))) | (s('type.list.list'('type.string.char'),X9) = s('type.list.list'('type.string.char'),X10) & s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn',X7))) & ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) | s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))))),
  inference(ennf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (? [X13] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) => ! [X14] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & ! [X6,X7,X8,X9,X10,X11] : ((p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X6),s('type.semanticPrimitives.tid_or_exn',X7)))) & s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) = s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.semanticPrimitives.tid_or_exn',X6))))) & s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) = s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn',X7)))))) => (s('type.list.list'('type.string.char'),X9) = s('type.list.list'('type.string.char'),X10) & s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn',X7))) & ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) => s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))) & ! [X12] : (? [X13] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X13),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))) => ! [X14] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X14),s('type.semanticPrimitives.tid_or_exn',X12))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X2)))))))) & ! [X6,X7,X8,X9,X10,X11] : ((p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X6),s('type.semanticPrimitives.tid_or_exn',X7)))) & s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) = s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.semanticPrimitives.tid_or_exn',X6))))) & s('type.option.option'('type.pair.prod'(X0,X1)),'const.option.SOME_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X11))))) = s('type.option.option'('type.pair.prod'(X0,X1)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X10),s('type.semanticPrimitives.tid_or_exn',X7)))))) => (s('type.list.list'('type.string.char'),X9) = s('type.list.list'('type.string.char'),X10) & s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn',X7))) & ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X5),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X0,X1)),X3)))))) => s('type.list.list'('type.string.char'),X5) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X4))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X33,X34] : ! [X40,X41] : (p(s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X40),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41)))) <=> (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X40),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41)))) & ! [X42] : (? [X43] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X43),s('type.semanticPrimitives.tid_or_exn',X42))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X40)))))) => ! [X43] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X43),s('type.semanticPrimitives.tid_or_exn',X42))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41)))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X43),s('type.semanticPrimitives.tid_or_exn',X42))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X40)))))))) & ! [X21,X22,X44,X43,X45,X46] : ((p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',X21),s('type.semanticPrimitives.tid_or_exn',X22)))) & s('type.option.option'('type.pair.prod'(X33,X34)),'const.option.SOME_1'(s('type.pair.prod'(X33,X34),'const.pair.,_2'(s(X33,X44),s(X34,X46))))) = s('type.option.option'('type.pair.prod'(X33,X34)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X43),s('type.semanticPrimitives.tid_or_exn',X21))))) & s('type.option.option'('type.pair.prod'(X33,X34)),'const.option.SOME_1'(s('type.pair.prod'(X33,X34),'const.pair.,_2'(s(X33,X44),s(X34,X46))))) = s('type.option.option'('type.pair.prod'(X33,X34)),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41),s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X45),s('type.semanticPrimitives.tid_or_exn',X22)))))) => (s('type.list.list'('type.string.char'),X43) = s('type.list.list'('type.string.char'),X45) & s('type.semanticPrimitives.tid_or_exn',X21) = s('type.semanticPrimitives.tid_or_exn',X22))) & ! [X42,X43] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X43),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X42))))),s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'(X33,X34)),X41)))))) => s('type.list.list'('type.string.char'),X43) = s('type.list.list'('type.string.char'),'const.ast.id_to_n_1'(s('type.ast.id'('type.list.list'('type.string.char')),X42))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.gtagenv_weak_def')).
fof(f8378,plain,(
  ( ! [X259,X260] : (~sP14(X260,X259) | p(s(bool,X260)) | s(bool,X259) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f843])).
fof(f843,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP14(X0,X1)) )),
  inference(cnf_transformation,[],[f576])).
fof(f14247,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | $spl58),
  inference(cnf_transformation,[],[f14247_D])).
fof(f14247_D,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f14263,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f14255])).
fof(f14255,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | p(s(bool,X11))) ) | $spl12),
  inference(resolution,[],[f8378,f6829])).
fof(f6829,plain,(
  ( ! [X189] : (sP14(X189,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | p(s(bool,X189))) ) | $spl12),
  inference(resolution,[],[f6298,f848])).
fof(f12826,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | $spl46),
  inference(cnf_transformation,[],[f12826_D])).
fof(f12826_D,plain,(
  s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f14264,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(duplicate_literal_removal,[],[f14254])).
fof(f14254,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X10))) )),
  inference(resolution,[],[f8378,f1771])).
fof(f1771,plain,(
  ( ! [X195] : (sP14(X195,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X195))) )),
  inference(resolution,[],[f678,f848])).
fof(f14240,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(duplicate_literal_removal,[],[f14239])).
fof(f14239,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17))) )),
  inference(resolution,[],[f8377,f3787])).
fof(f3787,plain,(
  ( ! [X180] : (sP14('const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X180) | p(s(bool,X180))) )),
  inference(resolution,[],[f1844,f845])).
fof(f8377,plain,(
  ( ! [X257,X258] : (~sP14(X257,X258) | p(s(bool,X258)) | s(bool,X257) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f843])).
fof(f14241,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f14238])).
fof(f14238,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | p(s(bool,X16))) ) | $spl12),
  inference(resolution,[],[f8377,f6828])).
fof(f6828,plain,(
  ( ! [X188] : (sP14('const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))),X188) | p(s(bool,X188))) ) | $spl12),
  inference(resolution,[],[f6298,f845])).
fof(f14242,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(duplicate_literal_removal,[],[f14237])).
fof(f14237,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X15))) )),
  inference(resolution,[],[f8377,f1770])).
fof(f1770,plain,(
  ( ! [X194] : (sP14('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X194) | p(s(bool,X194))) )),
  inference(resolution,[],[f678,f845])).
fof(f13516,plain,(
  s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | $spl54),
  inference(cnf_transformation,[],[f13516_D])).
fof(f13516_D,plain,(
  s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f14072,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1743,f13291])).
fof(f13291,plain,(
  ( ! [X189,X190,X188] : (sP9(X189,X188,X190) | s(bool,X188) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f793])).
fof(f3795,plain,(
  ( ! [X188] : (~p(s(bool,X188)) | s(bool,X188) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f865])).
fof(f1743,plain,(
  ( ! [X142,X143] : (~sP9(X142,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X143) | p(s(bool,X143)) | ~p(s(bool,X142))) )),
  inference(resolution,[],[f678,f791])).
fof(f791,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f552])).
fof(f14067,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1743,f8354])).
fof(f8354,plain,(
  ( ! [X189,X190,X188] : (sP9(X189,X188,X190) | s(bool,X188) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f793])).
fof(f14040,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1742,f13292])).
fof(f13292,plain,(
  ( ! [X191,X192,X193] : (sP9(X191,X192,X193) | s(bool,X191) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f794])).
fof(f1742,plain,(
  ( ! [X140,X141] : (~sP9('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X140,X141) | p(s(bool,X141)) | ~p(s(bool,X140))) )),
  inference(resolution,[],[f678,f791])).
fof(f14035,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1742,f8355])).
fof(f14004,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1740,f13288])).
fof(f13288,plain,(
  ( ! [X180,X181,X179] : (sP10(X180,X179,X181) | s(bool,X179) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f789])).
fof(f1740,plain,(
  ( ! [X136,X137] : (~sP10(X136,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X137) | p(s(bool,X137)) | ~p(s(bool,X136))) )),
  inference(resolution,[],[f678,f787])).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f549])).
fof(f13999,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1740,f8351])).
fof(f8351,plain,(
  ( ! [X180,X181,X179] : (sP10(X180,X179,X181) | s(bool,X179) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f789])).
fof(f13973,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1739,f13289])).
fof(f13289,plain,(
  ( ! [X182,X184,X183] : (sP10(X182,X183,X184) | s(bool,X182) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f790])).
fof(f1739,plain,(
  ( ! [X134,X135] : (~sP10('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X134,X135) | p(s(bool,X135)) | ~p(s(bool,X134))) )),
  inference(resolution,[],[f678,f787])).
fof(f13968,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1739,f8352])).
fof(f13680,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1727,f13275])).
fof(f13275,plain,(
  ( ! [X142,X140,X141] : (sP6(X141,X140,X142) | s(bool,X140) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f772])).
fof(f772,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f540])).
fof(f540,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f488])).
fof(f488,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f1727,plain,(
  ( ! [X111,X110] : (~sP6(X110,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X111) | p(s(bool,X111)) | ~p(s(bool,X110))) )),
  inference(resolution,[],[f678,f770])).
fof(f770,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f541])).
fof(f13675,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1727,f8338])).
fof(f8338,plain,(
  ( ! [X142,X140,X141] : (sP6(X141,X140,X142) | s(bool,X140) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f772])).
fof(f13649,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1726,f13276])).
fof(f13276,plain,(
  ( ! [X144,X145,X143] : (sP6(X143,X144,X145) | s(bool,X143) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f773])).
fof(f773,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f541])).
fof(f1726,plain,(
  ( ! [X109,X108] : (~sP6('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X108,X109) | p(s(bool,X109)) | ~p(s(bool,X108))) )),
  inference(resolution,[],[f678,f770])).
fof(f13644,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1726,f8339])).
fof(f8339,plain,(
  ( ! [X144,X145,X143] : (sP6(X143,X144,X145) | s(bool,X143) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f773])).
fof(f13636,plain,(
  ( ! [X23,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f13292,f1121])).
fof(f1121,plain,(
  ( ! [X130,X131] : (~sP9('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X130,X131) | p(s(bool,X131)) | ~p(s(bool,X130))) )),
  inference(resolution,[],[f677,f791])).
fof(f13445,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl50),
  inference(cnf_transformation,[],[f13445_D])).
fof(f13445_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13628,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(resolution,[],[f13292,f9443])).
fof(f13625,plain,(
  ( ! [X21,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f13291,f1122])).
fof(f1122,plain,(
  ( ! [X132,X133] : (~sP9(X132,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X133) | p(s(bool,X133)) | ~p(s(bool,X132))) )),
  inference(resolution,[],[f677,f791])).
fof(f13615,plain,(
  ( ! [X23,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f13289,f1118])).
fof(f1118,plain,(
  ( ! [X125,X124] : (~sP10('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X124,X125) | p(s(bool,X125)) | ~p(s(bool,X124))) )),
  inference(resolution,[],[f677,f787])).
fof(f13603,plain,(
  ( ! [X21,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f13288,f1119])).
fof(f1119,plain,(
  ( ! [X127,X126] : (~sP10(X126,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X127) | p(s(bool,X127)) | ~p(s(bool,X126))) )),
  inference(resolution,[],[f677,f787])).
fof(f13605,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X3) = s(bool,X4) | p(s(bool,X4))) ) | $spl56),
  inference(cnf_transformation,[],[f13605_D])).
fof(f13605_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X3) = s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13597,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f13288,f9156])).
fof(f13594,plain,(
  ( ! [X19,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X19)) | ~p(s(bool,X20))) )),
  inference(resolution,[],[f13276,f1105])).
fof(f1105,plain,(
  ( ! [X99,X98] : (~sP6('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X98,X99) | p(s(bool,X99)) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f677,f770])).
fof(f13542,plain,(
  ( ! [X17,X18] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f13275,f1106])).
fof(f1106,plain,(
  ( ! [X101,X100] : (~sP6(X100,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X101) | p(s(bool,X101)) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f677,f770])).
fof(f13535,plain,(
  ( ! [X21,X22] : (s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f13262,f1710])).
fof(f1710,plain,(
  ( ! [X76,X77] : (~sP4('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X76,X77) | p(s(bool,X77)) | p(s(bool,X76))) )),
  inference(resolution,[],[f678,f753])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f533])).
fof(f533,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f13262,plain,(
  ( ! [X103,X101,X102] : (sP4(X101,X102,X103) | s(bool,X101) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f756])).
fof(f756,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f534])).
fof(f13534,plain,(
  ( ! [X19,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X19)) | p(s(bool,X20))) )),
  inference(resolution,[],[f13262,f1089])).
fof(f1089,plain,(
  ( ! [X66,X67] : (~sP4('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X66,X67) | p(s(bool,X67)) | p(s(bool,X66))) )),
  inference(resolution,[],[f677,f753])).
fof(f13526,plain,(
  ( ! [X24,X23] : (s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X23)) | p(s(bool,X24))) )),
  inference(resolution,[],[f13246,f1696])).
fof(f1696,plain,(
  ( ! [X48,X49] : (~sP2('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X48,X49) | ~p(s(bool,X49)) | p(s(bool,X48))) )),
  inference(resolution,[],[f678,f736])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f525])).
fof(f525,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f524])).
fof(f524,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f13246,plain,(
  ( ! [X54,X55,X53] : (sP2(X53,X54,X55) | s(bool,X53) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f739])).
fof(f739,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f526])).
fof(f13525,plain,(
  ( ! [X21,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f13246,f1075])).
fof(f1075,plain,(
  ( ! [X39,X38] : (~sP2('const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)),X38,X39) | ~p(s(bool,X39)) | p(s(bool,X38))) )),
  inference(resolution,[],[f677,f736])).
fof(f13508,plain,(
  ( ! [X17,X18] : (s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f13245,f1697])).
fof(f1697,plain,(
  ( ! [X50,X51] : (~sP2(X50,X51,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X51)) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f678,f736])).
fof(f13245,plain,(
  ( ! [X52,X50,X51] : (sP2(X51,X52,X50) | s(bool,X50) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f3795,f737])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f526])).
fof(f13507,plain,(
  ( ! [X15,X16] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f13245,f1076])).
fof(f1076,plain,(
  ( ! [X41,X40] : (~sP2(X40,X41,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X41)) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f677,f736])).
fof(f13450,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1710,f8325])).
fof(f8325,plain,(
  ( ! [X103,X101,X102] : (sP4(X101,X102,X103) | s(bool,X101) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f756])).
fof(f13333,plain,(
  ( ! [X304,X306,X305] : (s(bool,X304) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X305)) | ~p(s(bool,X306)) | s(bool,X304) != s(bool,X306)) )),
  inference(resolution,[],[f3795,f1048])).
fof(f13238,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f3795,f728])).
fof(f13237,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f3795,f727])).
fof(f13235,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f3795,f726])).
fof(f13233,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f3795,f724])).
fof(f13231,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f3795,f722])).
fof(f13058,plain,(
  ( ! [X23,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f12169,f1121])).
fof(f12169,plain,(
  ( ! [X191,X192,X193] : (sP9(X191,X192,X193) | s(bool,X191) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f794])).
fof(f1778,plain,(
  ( ! [X202] : (~p(s(bool,X202)) | s(bool,X202) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f865])).
fof(f12322,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl44),
  inference(cnf_transformation,[],[f12322_D])).
fof(f12322_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f13050,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(resolution,[],[f12169,f9443])).
fof(f12996,plain,(
  ( ! [X21,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f12168,f1122])).
fof(f12168,plain,(
  ( ! [X189,X190,X188] : (sP9(X189,X188,X190) | s(bool,X188) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f793])).
fof(f12986,plain,(
  ( ! [X23,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f12166,f1118])).
fof(f12166,plain,(
  ( ! [X182,X184,X183] : (sP10(X182,X183,X184) | s(bool,X182) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f790])).
fof(f12974,plain,(
  ( ! [X21,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f12165,f1119])).
fof(f12165,plain,(
  ( ! [X180,X181,X179] : (sP10(X180,X179,X181) | s(bool,X179) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f789])).
fof(f12976,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X3) = s(bool,X4) | p(s(bool,X4))) ) | $spl48),
  inference(cnf_transformation,[],[f12976_D])).
fof(f12976_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X3) = s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12968,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | ~p(s(bool,X5))) )),
  inference(resolution,[],[f12165,f9156])).
fof(f12965,plain,(
  ( ! [X19,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X19)) | ~p(s(bool,X20))) )),
  inference(resolution,[],[f12153,f1105])).
fof(f12153,plain,(
  ( ! [X144,X145,X143] : (sP6(X143,X144,X145) | s(bool,X143) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f773])).
fof(f12954,plain,(
  ( ! [X17,X18] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f12152,f1106])).
fof(f12152,plain,(
  ( ! [X142,X140,X141] : (sP6(X141,X140,X142) | s(bool,X140) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f772])).
fof(f12907,plain,(
  ( ! [X19,X20] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X19)) | p(s(bool,X20))) )),
  inference(resolution,[],[f12139,f1089])).
fof(f12139,plain,(
  ( ! [X103,X101,X102] : (sP4(X101,X102,X103) | s(bool,X101) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f756])).
fof(f12898,plain,(
  ( ! [X21,X22] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f12123,f1075])).
fof(f12123,plain,(
  ( ! [X54,X55,X53] : (sP2(X53,X54,X55) | s(bool,X53) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f739])).
fof(f12884,plain,(
  ( ! [X15,X16] : (s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f12122,f1076])).
fof(f12122,plain,(
  ( ! [X52,X50,X51] : (sP2(X51,X52,X50) | s(bool,X50) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1778,f737])).
fof(f12831,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1697,f8308])).
fof(f8308,plain,(
  ( ! [X52,X50,X51] : (sP2(X51,X52,X50) | s(bool,X50) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f737])).
fof(f12802,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1696,f8309])).
fof(f8309,plain,(
  ( ! [X54,X55,X53] : (sP2(X53,X54,X55) | s(bool,X53) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1157,f739])).
fof(f12210,plain,(
  ( ! [X304,X306,X305] : (s(bool,X304) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X305)) | ~p(s(bool,X306)) | s(bool,X304) != s(bool,X306)) )),
  inference(resolution,[],[f1778,f1048])).
fof(f12115,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1778,f728])).
fof(f12114,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f1778,f727])).
fof(f12112,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1778,f726])).
fof(f12110,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1778,f724])).
fof(f12108,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1778,f722])).
fof(f8507,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f8507_D])).
fof(f8507_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9772,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X17) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f9443,f8355])).
fof(f9572,plain,(
  $false | $spl6),
  inference(resolution,[],[f1221,f955])).
fof(f955,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1))))))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X1))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X25,X18] : p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X25),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X18))),s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X18))))))),
  file('thm.mod_to_conProof.mod_to_conProof_58_prove_ltb','thm.mod_to_conProof.sv_rel_rules|split|1')).
fof(f9574,plain,(
  $false | $spl6),
  inference(resolution,[],[f1221,f677])).
fof(f9578,plain,(
  $false | $spl6),
  inference(resolution,[],[f1221,f678])).
fof(f9588,plain,(
  $false | ($spl6 | $spl12)),
  inference(resolution,[],[f1221,f6298])).
fof(f9642,plain,(
  $false | $spl6),
  inference(resolution,[],[f1221,f1844])).
fof(f9463,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) | $spl42),
  inference(cnf_transformation,[],[f9463_D])).
fof(f9463_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9456,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X28))) )),
  inference(resolution,[],[f9442,f3758])).
fof(f3758,plain,(
  ( ! [X125,X124] : (sP10(X124,X125,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f788])).
fof(f9462,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl40),
  inference(cnf_transformation,[],[f9462_D])).
fof(f9462_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f9455,plain,(
  ( ! [X26,X25] : (~p(s(bool,X25)) | s(bool,X25) = s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f9442,f6799])).
fof(f6799,plain,(
  ( ! [X132,X133] : (sP10(X132,X133,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f6298,f788])).
fof(f9461,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | s(bool,X23) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) | $spl38),
  inference(cnf_transformation,[],[f9461_D])).
fof(f9461_D,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | s(bool,X23) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9454,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | s(bool,X23) = s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f9442,f1741])).
fof(f1741,plain,(
  ( ! [X138,X139] : (sP10(X138,X139,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f788])).
fof(f9460,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,X19) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X18) = s(bool,X19)) ) | $spl36),
  inference(cnf_transformation,[],[f9460_D])).
fof(f9460_D,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,X19) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X18) = s(bool,X19)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f9452,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,X19) | ~p(s(bool,X20)) | s(bool,X19) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f9442,f8774])).
fof(f8774,plain,(
  ( ! [X182,X184,X183] : (sP10(X183,X184,X182) | s(bool,X182) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1318,f788])).
fof(f8941,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f8941_D])).
fof(f8941_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9425,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) )),
  inference(resolution,[],[f8780,f9058])).
fof(f9418,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f8779,f9057])).
fof(f8779,plain,(
  ( ! [X198,X199,X197] : (sP8(X198,X197,X199) | s(bool,X197) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f1318,f797])).
fof(f9216,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X4) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X3) = s(bool,X4)) ) | $spl34),
  inference(cnf_transformation,[],[f9216_D])).
fof(f9216_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X4) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X3) = s(bool,X4)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9204,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f9156,f8351])).
fof(f3848,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) | $spl26),
  inference(cnf_transformation,[],[f3848_D])).
fof(f3848_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9132,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | s(bool,X35) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X35))) )),
  inference(resolution,[],[f9058,f3764])).
fof(f3764,plain,(
  ( ! [X136,X137] : (sP8('const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X136,X137)) )),
  inference(resolution,[],[f1844,f798])).
fof(f5728,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl28),
  inference(cnf_transformation,[],[f5728_D])).
fof(f5728_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9131,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,X33) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | p(s(bool,X33))) ) | $spl12),
  inference(resolution,[],[f9058,f6805])).
fof(f6805,plain,(
  ( ! [X144,X145] : (sP8('const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))),X144,X145)) ) | $spl12),
  inference(resolution,[],[f6298,f798])).
fof(f1831,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) | $spl24),
  inference(cnf_transformation,[],[f1831_D])).
fof(f1831_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9130,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,X31) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X31))) )),
  inference(resolution,[],[f9058,f1747])).
fof(f1747,plain,(
  ( ! [X151,X150] : (sP8('const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X150,X151)) )),
  inference(resolution,[],[f678,f798])).
fof(f9112,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | s(bool,X27) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X27))) )),
  inference(resolution,[],[f9057,f3763])).
fof(f3763,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X135)) )),
  inference(resolution,[],[f1844,f797])).
fof(f9111,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f9057,f6804])).
fof(f6804,plain,(
  ( ! [X142,X143] : (sP8(X142,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))),X143)) ) | $spl12),
  inference(resolution,[],[f6298,f797])).
fof(f9110,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X23) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | p(s(bool,X23))) )),
  inference(resolution,[],[f9057,f1746])).
fof(f1746,plain,(
  ( ! [X149,X148] : (sP8(X148,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)),X149)) )),
  inference(resolution,[],[f678,f797])).
fof(f8819,plain,(
  ( ! [X308,X309,X307] : (s(bool,X307) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X308)) | p(s(bool,X309)) | s(bool,X307) != s(bool,X309)) )),
  inference(resolution,[],[f1318,f1048])).
fof(f8723,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1318,f728])).
fof(f8720,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1318,f727])).
fof(f8719,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1318,f726])).
fof(f8718,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1318,f724])).
fof(f8716,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1318,f722])).
fof(f8396,plain,(
  ( ! [X304,X306,X305] : (s(bool,X304) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X305)) | ~p(s(bool,X306)) | s(bool,X304) != s(bool,X306)) )),
  inference(resolution,[],[f1157,f1048])).
fof(f8301,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1157,f728])).
fof(f8300,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f1157,f727])).
fof(f8298,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1157,f726])).
fof(f8296,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1157,f724])).
fof(f8294,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1157,f722])).
fof(f6844,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f6298,f1048])).
fof(f6748,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | s(bool,X30) != s(bool,X31)) ) | $spl12),
  inference(resolution,[],[f6298,f728])).
fof(f6745,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f6298,f727])).
fof(f6744,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f6298,f726])).
fof(f6743,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | s(bool,X15) != s(bool,X16)) ) | $spl12),
  inference(resolution,[],[f6298,f724])).
fof(f6741,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f6298,f722])).
fof(f6137,plain,(
  $false | ($spl14 | $spl18)),
  inference(resolution,[],[f6082,f679])).
fof(f6082,plain,(
  ( ! [X0] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) ) | ($spl14 | $spl18)),
  inference(forward_demodulation,[],[f5964,f1388])).
fof(f5964,plain,(
  ( ! [X0] : (p(s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),X0),s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))))) ) | $spl18),
  inference(superposition,[],[f955,f1392])).
fof(f1392,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl18),
  inference(cnf_transformation,[],[f1392_D])).
fof(f1392_D,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5767,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5760,f1060])).
fof(f5760,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f943])).
fof(f1385,plain,(
  ~sP17(sK19,sK21,sK20) | $spl13),
  inference(cnf_transformation,[],[f1385_D])).
fof(f1385_D,plain,(
  ~sP17(sK19,sK21,sK20) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1388,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl14),
  inference(cnf_transformation,[],[f1388_D])).
fof(f1388_D,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5766,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5759,f1060])).
fof(f5759,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f942])).
fof(f5765,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5758,f1060])).
fof(f5758,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f941])).
fof(f5764,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5757,f1060])).
fof(f5757,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f940])).
fof(f5763,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5756,f1060])).
fof(f5756,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f939])).
fof(f5762,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | $spl13),
  inference(subsumption_resolution,[],[f5755,f1060])).
fof(f5755,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | ~sP18(sK19,sK21,sK20) | $spl13),
  inference(resolution,[],[f1385,f938])).
fof(f5667,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f1834,f1048])).
fof(f1834,plain,(
  p(s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) | $spl12),
  inference(resolution,[],[f1386,f949])).
fof(f5571,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | s(bool,X30) != s(bool,X31)) ) | $spl12),
  inference(resolution,[],[f1834,f728])).
fof(f5568,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f1834,f727])).
fof(f5567,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f1834,f726])).
fof(f5566,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19)))) | s(bool,X15) != s(bool,X16)) ) | $spl12),
  inference(resolution,[],[f1834,f724])).
fof(f5564,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.mod_to_conProof.v_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.modSem.v',sK66(sK20,sK21,sK19)),s('type.conSem.v',sK67(sK20,sK21,sK19))))) ) | $spl12),
  inference(resolution,[],[f1834,f722])).
fof(f3803,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f1048])).
fof(f3707,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1844,f728])).
fof(f3704,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f727])).
fof(f3703,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f726])).
fof(f3702,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1844,f724])).
fof(f3700,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f1844,f722])).
fof(f2323,plain,(
  $false | $spl4),
  inference(resolution,[],[f2211,f955])).
fof(f2211,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(equality_resolution,[],[f1220])).
fof(f2325,plain,(
  $false | $spl4),
  inference(resolution,[],[f2211,f677])).
fof(f2329,plain,(
  $false | $spl4),
  inference(resolution,[],[f2211,f678])).
fof(f2020,plain,(
  $false | $spl2),
  inference(resolution,[],[f1900,f679])).
fof(f1900,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl2),
  inference(equality_resolution,[],[f1219])).
fof(f1786,plain,(
  ( ! [X218,X219] : (p(s(bool,X218)) | p(s(bool,X219)) | s(bool,X219) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f1048])).
fof(f1690,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f678,f728])).
fof(f1687,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f727])).
fof(f1686,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f726])).
fof(f1685,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f678,f724])).
fof(f1683,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.mod_to_conProof.gtagenv_weak_2'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22)))) )),
  inference(resolution,[],[f678,f722])).
fof(f1519,plain,(
  $false | $spl8),
  inference(resolution,[],[f1222,f679])).
fof(f1383,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f943])).
fof(f1382,plain,(
  p(s(bool,'const.mod_to_conProof.vs_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)),s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19))))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f942])).
fof(f1381,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f941])).
fof(f1380,plain,(
  s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.conSem.v'),sK65(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f940])).
fof(f1379,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21) = s('type.semanticPrimitives.store_v'('type.conSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f939])).
fof(f1378,plain,(
  s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.Varray_1'(s('type.list.list'('type.modSem.v'),sK64(sK20,sK21,sK19)))) | s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20) = s('type.semanticPrimitives.store_v'('type.modSem.v'),'const.semanticPrimitives.W8array_1'(s('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),sK63(sK20,sK21,sK19)))) | sP17(sK19,sK21,sK20)),
  inference(resolution,[],[f1060,f938])).
fof(f1377,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) | $spl10),
  inference(cnf_transformation,[],[f1377_D])).
fof(f1377_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1327,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f679,f1048])).
fof(f1232,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f679,f728])).
fof(f1231,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f679,f727])).
fof(f1229,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f679,f726])).
fof(f1227,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f679,f724])).
fof(f1225,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK22),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f679,f722])).
fof(f1218,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) | $spl0),
  inference(cnf_transformation,[],[f1218_D])).
fof(f1218_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1165,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f677,f1048])).
fof(f1069,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f677,f728])).
fof(f1066,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f677,f727])).
fof(f1065,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f677,f726])).
fof(f1064,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f677,f724])).
fof(f1062,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.mod_to_conProof.sv_rel_3'(s('type.finite_map.fmap'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.pair.prod'('type.num.num','type.num.num')),sK19),s('type.semanticPrimitives.store_v'('type.modSem.v'),sK20),s('type.semanticPrimitives.store_v'('type.conSem.v'),sK21)))) )),
  inference(resolution,[],[f677,f722])).
% SZS output end Proof for thm.mod_to_conProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 740
% Passive clauses: 14706
% Generated clauses: 23500
% Final active clauses: 687
% Final passive clauses: 12761
% Input formulas: 260
% Initial clauses: 371
% 
% Selected by SInE selection: 107
% SInE iterations: 1
% 
% Duplicate literals: 659
% Trivial inequalities: 14
% Fw subsumption resolutions: 395
% Fw demodulations: 890
% 
% Simple tautologies: 1055
% Equational tautologies: 39
% Forward subsumptions: 4870
% 
% Binary resolution: 9238
% Forward superposition: 10677
% Backward superposition: 1416
% Equality resolution: 8
% 
% Split clauses: 299
% Split components: 684
% SAT solver clauses: 318
% SAT solver unit clauses: 10
% SAT solver binary clauses: 145
% 
% Sat splits: 299
% Sat splitting refutations: 14
% 
% Memory used [KB]: 17142
% Time elapsed: 0.690 s
% ------------------------------
% ------------------------------
% Success in time 12.517 s
