Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvlSem
% (127391)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 138
% Passive clauses: 3113
% Generated clauses: 24181
% Final active clauses: 123
% Final passive clauses: 2819
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 862
% Trivial inequalities: 304
% Fw subsumption resolutions: 2585
% Fw demodulations: 4277
% Bw demodulations: 10
% 
% Simple tautologies: 292
% Equational tautologies: 62
% Forward subsumptions: 12122
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 147
% Factoring: 1
% Forward superposition: 8262
% Backward superposition: 7262
% Self superposition: 268
% Equality factoring: 150
% Equality resolution: 27
% 
% Split clauses: 241
% Split components: 482
% SAT solver clauses: 686
% SAT solver unit clauses: 12
% SAT solver binary clauses: 52
% 
% Sat splits: 241
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 686
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvlSem
% (127407)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 144
% Passive clauses: 1361
% Generated clauses: 17960
% Final active clauses: 86
% Final passive clauses: 686
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 665
% Trivial inequalities: 86
% Fw subsumption resolutions: 909
% Fw demodulations: 2206
% Bw demodulations: 548
% Condensations: 1374
% 
% Simple tautologies: 130
% Equational tautologies: 45
% Forward subsumptions: 9533
% Bw demodulations to eq. taut.: 30
% 
% Binary resolution: 258
% Factoring: 1
% Forward superposition: 5662
% Backward superposition: 4580
% Self superposition: 181
% Equality factoring: 86
% Equality resolution: 10
% 
% Split clauses: 75
% Split components: 150
% SAT solver clauses: 1614
% SAT solver unit clauses: 130
% SAT solver binary clauses: 39
% 
% Sat splits: 75
% Sat splitting refutations: 32
% 
% TWLsolver clauses: 1614
% TWLsolver calls for satisfiability: 245
% 
% Memory used [KB]: 17398
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvlSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvlSem
% SZS output start Proof for thm.bvlSem
fof(f31726,plain,(
  $false),
  inference(subsumption_resolution,[],[f31725,f214])).
fof(f214,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3',aHL_TRUTH)).
fof(f31725,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f31723,f2248])).
fof(f2248,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2163,f295])).
fof(f295,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f216,f217])).
fof(f217,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f216,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f2163,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2088,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6,X14] : ! [X16,X21] : s(X14,X21) = s(X14,'const.pair.SND_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X16),s(X14,X21)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3','thm.pair.SND')).
fof(f2088,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f259,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X49,X50,X51] : (p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),'=_0'),s(X49,X50))),s(X49,X51)))) <=> s(X49,X50) = s(X49,X51))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f31723,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),fun(fun(sK0,'type.bvlSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f31663,f10177])).
fof(f10177,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),fun(fun(sK0,'type.bvlSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f213,f265])).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f213,plain,(
  s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) != s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) = s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) = s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))))),s(fun(X6,'type.bvlSem.state'(X7)),X8))) = s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10))),s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))),s(fun(X6,'type.bvlSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))))),s(fun(X6,'type.bvlSem.state'(X7)),X8))) = s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10))),s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))),s(fun(X6,'type.bvlSem.state'(X7)),X8)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3','thm.bvlSem.state_fupdfupds_comp|split|3')).
fof(f31663,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))))),s(fun(X3,'type.bvlSem.state'(X0)),X4))) = s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1))),s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s(fun(X3,'type.bvlSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f272,f226])).
fof(f226,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))))) = s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10))),s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_refs_fupd_1'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3','thm.bvlSem.state_fupdfupds_comp|split|2')).
fof(f272,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14,X6,X22,X15] : ! [X10,X9,X8] : s(fun(X22,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X22,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X22,X15),X8))))) = s(fun(X22,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X22,X15),X8)))),
  file('thm.bvlSem.state_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.bvlSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 197
% Passive clauses: 2772
% Generated clauses: 32868
% Final active clauses: 98
% Final passive clauses: 1202
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 1102
% Trivial inequalities: 108
% Fw subsumption resolutions: 1000
% Fw demodulations: 5019
% Bw demodulations: 1245
% Condensations: 1661
% Global subsumptions: 2
% 
% Simple tautologies: 154
% Equational tautologies: 232
% Forward subsumptions: 20619
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 25
% 
% Unit resulting resolution: 1240
% Factoring: 1
% Forward superposition: 9998
% Backward superposition: 9378
% Self superposition: 326
% Equality factoring: 144
% Equality resolution: 8
% 
% Split clauses: 51
% Split components: 102
% SAT solver clauses: 2975
% SAT solver unit clauses: 486
% SAT solver binary clauses: 1321
% 
% Sat splits: 51
% Sat splitting refutations: 97
% 
% Memory used [KB]: 38762
% Time elapsed: 1.333 s
% ------------------------------
% ------------------------------
% Success in time 2.68 s
