Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (102548)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 3429
% Generated clauses: 17588
% Final active clauses: 98
% Final passive clauses: 3043
% Input formulas: 105
% Initial clauses: 117
% 
% Duplicate literals: 631
% Trivial inequalities: 204
% Fw subsumption resolutions: 1399
% Fw demodulations: 2121
% Bw demodulations: 2
% 
% Simple tautologies: 281
% Equational tautologies: 83
% Forward subsumptions: 8995
% 
% Binary resolution: 89
% Factoring: 1
% Forward superposition: 5921
% Backward superposition: 6522
% Self superposition: 291
% Equality factoring: 179
% Equality resolution: 18
% 
% Split clauses: 170
% Split components: 345
% SAT solver clauses: 272
% SAT solver unit clauses: 18
% SAT solver binary clauses: 58
% 
% Sat splits: 170
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 272
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (102573)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 2043
% Generated clauses: 16891
% Final active clauses: 84
% Final passive clauses: 1056
% Input formulas: 105
% Initial clauses: 117
% 
% Duplicate literals: 602
% Trivial inequalities: 173
% Fw subsumption resolutions: 596
% Fw demodulations: 817
% Bw demodulations: 838
% Condensations: 1161
% 
% Simple tautologies: 173
% Equational tautologies: 59
% Forward subsumptions: 7830
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 23
% 
% Binary resolution: 88
% Factoring: 1
% Forward superposition: 5042
% Backward superposition: 4067
% Self superposition: 212
% Equality factoring: 110
% Equality resolution: 7
% 
% Split clauses: 172
% Split components: 349
% SAT solver clauses: 3026
% SAT solver unit clauses: 196
% SAT solver binary clauses: 56
% 
% Sat splits: 172
% Sat splitting refutations: 41
% 
% TWLsolver clauses: 3026
% TWLsolver calls for satisfiability: 3584
% 
% Memory used [KB]: 24818
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% (102613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 254
% Passive clauses: 3453
% Generated clauses: 64352
% Final active clauses: 207
% Final passive clauses: 657
% Input formulas: 105
% Initial clauses: 117
% Discarded non-redundant clauses: 190671
% 
% Duplicate literals: 1466
% Trivial inequalities: 437
% Fw subsumption resolutions: 2030
% Fw demodulations: 6469
% Bw demodulations: 595
% Condensations: 3554
% Global subsumptions: 26
% 
% Simple tautologies: 415
% Equational tautologies: 211
% Forward subsumptions: 39529
% Fw demodulations to eq. taut.: 171
% Bw demodulations to eq. taut.: 91
% 
% Unit resulting resolution: 3058
% Factoring: 5
% Forward superposition: 28333
% Backward superposition: 17133
% Self superposition: 565
% Equality factoring: 158
% Equality resolution: 34
% 
% Split clauses: 74
% Split components: 153
% SAT solver clauses: 3633
% SAT solver unit clauses: 220
% SAT solver binary clauses: 481
% 
% Sat splits: 74
% Sat splitting refutations: 4
% 
% Memory used [KB]: 100296
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f2321,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f130,f138,f129,f142,f140,f137,f128,f127,f144,f126,f141,f125,f146,f139,f136,f162,f165,f178,f176,f179,f194,f195,f200,f214,f217,f216,f242,f259,f262,f265,f282,f199,f286,f292,f312,f310,f313,f316,f284,f317,f388,f390,f389,f414,f447,f459,f460,f479,f489,f526,f543,f548,f547,f551,f563,f590,f630,f669,f670,f261,f192,f684,f687,f701,f704,f197,f705,f720,f718,f721,f726,f725,f723,f753,f772,f770,f914,f937,f988,f1085,f1088,f1089,f1090,f1087,f1091,f1093,f1244,f264,f1247,f315,f1248,f686,f1249,f703,f1246,f1250,f1095,f1256,f1258,f1252,f1295,f1325,f1323,f1361,f1377,f1362,f1380,f1365,f1379,f1381,f1481,f1196,f1484,f1110,f608,f1332,f1486,f1294,f1488,f1135,f1489,f917,f488,f730,f545,f729,f1491,f1274,f1494,f486,f1497,f504,f502,f1539,f541,f1545,f1546,f542,f540,f1578,f1594,f1592,f1619,f1640,f1638,f1671,f1010,f1008,f1759,f1781,f1780,f2029,f2031,f2030,f2193,f2320])).
fof(f2320,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f2318])).
fof(f2318,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f2303,f136])).
fof(f2303,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f2276,f462])).
fof(f462,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl5),
  inference(trivial_inequality_removal,[],[f461])).
fof(f461,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.NIL_0') | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl5),
  inference(superposition,[],[f140,f132])).
fof(f132,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK8(X1,X0)),s('type.list.list'(X0),sK9(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK8(X1,X0)),s('type.list.list'(X0),sK9(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X6] : ! [X73] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X73) | ? [X12,X13] : s('type.list.list'(X6),X73) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))))),
  file('thm.closProps.eqs|split|1','thm.list.list_CASES')).
fof(f2276,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f138,f491])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK8(sK2,sK0)))),s('type.list.list'(sK0),sK9(sK2,sK0))))) ) | $spl5),
  inference(superposition,[],[f131,f462])).
fof(f131,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X7,X6] : ! [X58,X59,X21,X19] : s(X7,happ(s(fun('type.list.list'(X6),X7),happ(s(fun(X6,fun('type.list.list'(X6),X7)),X19),s(X6,X58))),s('type.list.list'(X6),X59))) = s(X7,'const.list.list_CASE_3'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X58),s('type.list.list'(X6),X59))),s(X7,X21),s(fun(X6,fun('type.list.list'(X6),X7)),X19)))),
  file('thm.closProps.eqs|split|1','thm.list.list_case_def|split|1')).
fof(f2193,plain,(
  $false | ($spl0 | $spl2 | $spl50)),
  inference(trivial_inequality_removal,[],[f2191])).
fof(f2191,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl0 | $spl2 | $spl50)),
  inference(superposition,[],[f2032,f136])).
fof(f2032,plain,(
  ( ! [X8] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X8),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl2 | $spl50)),
  inference(subsumption_resolution,[],[f2003,f504])).
fof(f2003,plain,(
  ( ! [X8] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X8),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | ($spl2 | $spl50)),
  inference(superposition,[],[f138,f1522])).
fof(f1522,plain,(
  ( ! [X2,X3,X1] : (s(X1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(X1,X3),s(fun(sK0,fun('type.list.list'(sK0),X1)),X2))) = s(X1,happ(s(fun('type.list.list'(sK0),X1),happ(s(fun(sK0,fun('type.list.list'(sK0),X1)),X2),s(sK0,sK8(sK2,sK0)))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl50),
  inference(superposition,[],[f131,f504])).
fof(f2030,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) | $spl132),
  inference(cnf_transformation,[],[f2030_D])).
fof(f2030_D,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f2031,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | $spl134),
  inference(cnf_transformation,[],[f2031_D])).
fof(f2031_D,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2029,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) | ($spl2 | $spl50)),
  inference(subsumption_resolution,[],[f2002,f1518])).
fof(f1518,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0)))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),X0)))) ) | $spl50),
  inference(superposition,[],[f504,f132])).
fof(f2002,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) | ($spl2 | $spl50)),
  inference(superposition,[],[f961,f1522])).
fof(f961,plain,(
  ( ! [X21,X20] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X21))),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X21),s('type.list.list'(sK0),X20))) | s('type.list.list'(sK0),X20) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X20,sK0)),s('type.list.list'(sK0),sK9(X20,sK0))))) ) | $spl2),
  inference(superposition,[],[f138,f132])).
fof(f1780,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) | $spl128),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1781,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | $spl130),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1759,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),X7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X7,sK0)),s('type.list.list'(sK0),sK9(X7,sK0))))) ) | ($spl2 | $spl18)),
  inference(superposition,[],[f961,f204])).
fof(f204,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK8(sK7,sK0)))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl18),
  inference(superposition,[],[f131,f194])).
fof(f1008,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | $spl80),
  inference(cnf_transformation,[],[f1008_D])).
fof(f1008_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1010,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl82),
  inference(cnf_transformation,[],[f1010_D])).
fof(f1010_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1671,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | ($spl5 | $spl50)),
  inference(superposition,[],[f1535,f132])).
fof(f1535,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl5 | $spl50)),
  inference(forward_demodulation,[],[f1534,f504])).
fof(f1534,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl5 | $spl50)),
  inference(subsumption_resolution,[],[f1520,f140])).
fof(f1520,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl50),
  inference(superposition,[],[f504,f132])).
fof(f1638,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)))) | $spl124),
  inference(cnf_transformation,[],[f1638_D])).
fof(f1638_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1640,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl126),
  inference(cnf_transformation,[],[f1640_D])).
fof(f1640_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1619,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0),sK0)))) | $spl5),
  inference(superposition,[],[f506,f132])).
fof(f506,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)))) | $spl5),
  inference(forward_demodulation,[],[f505,f462])).
fof(f505,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)))) | $spl5),
  inference(subsumption_resolution,[],[f490,f140])).
fof(f490,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0))),sK0)))) | $spl5),
  inference(superposition,[],[f462,f132])).
fof(f1592,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | $spl120),
  inference(cnf_transformation,[],[f1592_D])).
fof(f1592_D,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1594,plain,(
  s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | $spl123),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f1578,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | ($spl2 | $spl70)),
  inference(superposition,[],[f960,f720])).
fof(f960,plain,(
  ( ! [X19,X18] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X19),s('type.list.list'(sK0),'const.list.NIL_0'))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X19))),s('type.list.list'(sK0),X18))) | s('type.list.list'(sK0),X18) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X18,sK0)),s('type.list.list'(sK0),sK9(X18,sK0))))) ) | $spl2),
  inference(superposition,[],[f138,f132])).
fof(f540,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) | $spl52),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f542,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl54),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1546,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) | $spl48),
  inference(superposition,[],[f502,f132])).
fof(f1545,plain,(
  $false | ($spl49 | $spl57)),
  inference(subsumption_resolution,[],[f1543,f544])).
fof(f544,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl57),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK9(sK2,sK0)) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f1543,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl49),
  inference(trivial_inequality_removal,[],[f1541])).
fof(f1541,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),sK9(sK2,sK0)) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl49),
  inference(superposition,[],[f501,f132])).
fof(f501,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) | $spl49),
  inference(cnf_transformation,[],[f501_D])).
fof(f501_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f541,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl55),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f1539,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) | $spl49),
  inference(superposition,[],[f501,f132])).
fof(f502,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) | $spl48),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f504,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl50),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1497,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) | $spl5),
  inference(superposition,[],[f462,f132])).
fof(f486,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1494,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl118),
  inference(cnf_transformation,[],[f1494_D])).
fof(f1494_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1274,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl92),
  inference(superposition,[],[f1252,f132])).
fof(f1491,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl116),
  inference(cnf_transformation,[],[f1491_D])).
fof(f1491_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f729,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | ($spl18 | $spl72)),
  inference(superposition,[],[f201,f723])).
fof(f201,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),X0)))) ) | $spl18),
  inference(superposition,[],[f194,f132])).
fof(f545,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl56),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f730,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl72),
  inference(superposition,[],[f132,f723])).
fof(f488,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | $spl46),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f917,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl0),
  inference(superposition,[],[f136,f132])).
fof(f1489,plain,(
  ( ! [X1] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X1))),s('type.list.list'(sK0),sK9(sK2,sK0))))) ) | $spl114),
  inference(cnf_transformation,[],[f1489_D])).
fof(f1489_D,plain,(
  ( ! [X1] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X1))),s('type.list.list'(sK0),sK9(sK2,sK0))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1135,plain,(
  ( ! [X1] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X1),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X1))),s('type.list.list'(sK0),sK9(sK2,sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0))) ) | ($spl2 | $spl72)),
  inference(superposition,[],[f959,f723])).
fof(f959,plain,(
  ( ! [X17,X16] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X17),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X16,sK0)),s('type.list.list'(sK0),sK9(X16,sK0)))))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X17))),s('type.list.list'(sK0),X16))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),X16)) ) | $spl2),
  inference(superposition,[],[f138,f132])).
fof(f1488,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl112),
  inference(cnf_transformation,[],[f1488_D])).
fof(f1488_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1294,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | ($spl5 | $spl86)),
  inference(superposition,[],[f1125,f132])).
fof(f1125,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | ($spl5 | $spl86)),
  inference(forward_demodulation,[],[f1124,f1093])).
fof(f1124,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | ($spl5 | $spl86)),
  inference(subsumption_resolution,[],[f1111,f140])).
fof(f1111,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | $spl86),
  inference(superposition,[],[f1093,f132])).
fof(f1486,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl110),
  inference(cnf_transformation,[],[f1486_D])).
fof(f1486_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1332,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl98),
  inference(superposition,[],[f1325,f132])).
fof(f608,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl62),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1110,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl86),
  inference(superposition,[],[f1093,f132])).
fof(f1484,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl108),
  inference(cnf_transformation,[],[f1484_D])).
fof(f1484_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1196,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | ($spl70 | $spl72)),
  inference(superposition,[],[f778,f723])).
fof(f778,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | ($spl70 | $spl72)),
  inference(superposition,[],[f777,f132])).
fof(f777,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl70 | $spl72)),
  inference(forward_demodulation,[],[f720,f723])).
fof(f1481,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl56)),
  inference(trivial_inequality_removal,[],[f1479])).
fof(f1479,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl0 | $spl2 | $spl5 | $spl56)),
  inference(superposition,[],[f1470,f136])).
fof(f1470,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl2 | $spl5 | $spl56)),
  inference(subsumption_resolution,[],[f1448,f550])).
fof(f550,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl5 | $spl56)),
  inference(backward_demodulation,[],[f545,f462])).
fof(f1448,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X6),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | ($spl2 | $spl5 | $spl56)),
  inference(superposition,[],[f138,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK8(sK2,sK0)))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | ($spl5 | $spl56)),
  inference(superposition,[],[f131,f550])).
fof(f1381,plain,(
  ( ! [X4] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK7,sK0)))),s('type.list.list'(sK0),X4))) | s('type.list.list'(sK0),X4) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X4,sK0)),s('type.list.list'(sK0),sK9(X4,sK0))))) ) | $spl106),
  inference(cnf_transformation,[],[f1381_D])).
fof(f1381_D,plain,(
  ( ! [X4] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK7,sK0)))),s('type.list.list'(sK0),X4))) | s('type.list.list'(sK0),X4) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X4,sK0)),s('type.list.list'(sK0),sK9(X4,sK0))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1379,plain,(
  s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK7) | $spl103),
  inference(cnf_transformation,[],[f1379_D])).
fof(f1379_D,plain,(
  s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK7) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f1365,plain,(
  ( ! [X4] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK7) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8(sK7,sK0)))),s('type.list.list'(sK0),X4))) | s('type.list.list'(sK0),X4) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X4,sK0)),s('type.list.list'(sK0),sK9(X4,sK0))))) ) | ($spl2 | $spl18)),
  inference(superposition,[],[f960,f194])).
fof(f1380,plain,(
  ( ! [X1] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))),s('type.list.list'(sK0),X1))) | s('type.list.list'(sK0),X1) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X1,sK0)),s('type.list.list'(sK0),sK9(X1,sK0))))) ) | $spl104),
  inference(cnf_transformation,[],[f1380_D])).
fof(f1380_D,plain,(
  ( ! [X1] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))),s('type.list.list'(sK0),X1))) | s('type.list.list'(sK0),X1) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X1,sK0)),s('type.list.list'(sK0),sK9(X1,sK0))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1362,plain,(
  ( ! [X1] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK7) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))),s('type.list.list'(sK0),X1))) | s('type.list.list'(sK0),X1) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X1,sK0)),s('type.list.list'(sK0),sK9(X1,sK0))))) ) | ($spl2 | $spl28)),
  inference(superposition,[],[f960,f261])).
fof(f1377,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | $spl100),
  inference(cnf_transformation,[],[f1377_D])).
fof(f1377_D,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1361,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),sK7) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | ($spl2 | $spl36)),
  inference(superposition,[],[f960,f312])).
fof(f1323,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)))) | $spl96),
  inference(cnf_transformation,[],[f1323_D])).
fof(f1323_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1325,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl98),
  inference(cnf_transformation,[],[f1325_D])).
fof(f1325_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1295,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0),sK0)))) | ($spl5 | $spl86)),
  inference(superposition,[],[f1125,f132])).
fof(f1252,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) | $spl92),
  inference(cnf_transformation,[],[f1252_D])).
fof(f1252_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1258,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | $spl94),
  inference(cnf_transformation,[],[f1258_D])).
fof(f1258_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1256,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) | ($spl70 | $spl72 | $spl86)),
  inference(forward_demodulation,[],[f1177,f1093])).
fof(f1177,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2)),sK0)))) | ($spl70 | $spl72 | $spl86)),
  inference(superposition,[],[f778,f1093])).
fof(f1095,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl88),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1250,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) | ($spl5 | $spl56 | $spl70 | $spl72)),
  inference(forward_demodulation,[],[f1169,f587])).
fof(f587,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl5 | $spl56)),
  inference(forward_demodulation,[],[f586,f550])).
fof(f586,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl5 | $spl56)),
  inference(subsumption_resolution,[],[f574,f140])).
fof(f574,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl5 | $spl56)),
  inference(superposition,[],[f550,f132])).
fof(f1169,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK2))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | ($spl5 | $spl56 | $spl70 | $spl72)),
  inference(superposition,[],[f778,f587])).
fof(f1246,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | $spl90),
  inference(cnf_transformation,[],[f1246_D])).
fof(f1246_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f703,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl66),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1249,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl28 | $spl70 | $spl72)),
  inference(forward_demodulation,[],[f1168,f261])).
fof(f1168,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl28 | $spl70 | $spl72)),
  inference(superposition,[],[f778,f261])).
fof(f686,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl64),
  inference(cnf_transformation,[],[f686_D])).
fof(f686_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1248,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl36 | $spl70 | $spl72)),
  inference(forward_demodulation,[],[f1167,f312])).
fof(f1167,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl36 | $spl70 | $spl72)),
  inference(superposition,[],[f778,f312])).
fof(f315,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | $spl38),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1247,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl22 | $spl70 | $spl72)),
  inference(forward_demodulation,[],[f1166,f199])).
fof(f1166,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | ($spl22 | $spl70 | $spl72)),
  inference(superposition,[],[f778,f199])).
fof(f264,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl30),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1244,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl24 | $spl70 | $spl72)),
  inference(forward_demodulation,[],[f1165,f216])).
fof(f1165,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | ($spl24 | $spl70 | $spl72)),
  inference(superposition,[],[f778,f216])).
fof(f1093,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) | $spl86),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1091,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) | ($spl5 | $spl48 | $spl56 | $spl72)),
  inference(forward_demodulation,[],[f1023,f587])).
fof(f1023,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK2))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | ($spl5 | $spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f741,f587])).
fof(f741,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | ($spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f728,f132])).
fof(f728,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl48 | $spl56 | $spl72)),
  inference(backward_demodulation,[],[f723,f549])).
fof(f549,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) | ($spl48 | $spl56)),
  inference(backward_demodulation,[],[f545,f502])).
fof(f1087,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | $spl84),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1090,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl28 | $spl48 | $spl56 | $spl72)),
  inference(forward_demodulation,[],[f1022,f261])).
fof(f1022,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl28 | $spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f741,f261])).
fof(f1089,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl36 | $spl48 | $spl56 | $spl72)),
  inference(forward_demodulation,[],[f1021,f312])).
fof(f1021,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl36 | $spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f741,f312])).
fof(f1088,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl22 | $spl48 | $spl56 | $spl72)),
  inference(forward_demodulation,[],[f1020,f199])).
fof(f1020,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | ($spl22 | $spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f741,f199])).
fof(f1085,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | ($spl24 | $spl48 | $spl56 | $spl72)),
  inference(forward_demodulation,[],[f1019,f216])).
fof(f1019,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | ($spl24 | $spl48 | $spl56 | $spl72)),
  inference(superposition,[],[f741,f216])).
fof(f988,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | ($spl5 | $spl56)),
  inference(superposition,[],[f587,f132])).
fof(f937,plain,(
  $false | ($spl0 | $spl9 | $spl10)),
  inference(trivial_inequality_removal,[],[f935])).
fof(f935,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl0 | $spl9 | $spl10)),
  inference(superposition,[],[f926,f136])).
fof(f926,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X0),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl9 | $spl10)),
  inference(superposition,[],[f143,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK6))),s('type.list.list'(sK0),sK7)))) ) | $spl10),
  inference(superposition,[],[f131,f146])).
fof(f143,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) | $spl9),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f914,plain,(
  $false | ($spl1 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f906])).
fof(f906,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl1 | $spl8 | $spl10)),
  inference(superposition,[],[f137,f863])).
fof(f863,plain,(
  ( ! [X7] : (s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,X7),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f465,f144])).
fof(f770,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)))) | $spl76),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f772,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl78),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f753,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0),sK0)))) | ($spl5 | $spl10)),
  inference(superposition,[],[f476,f132])).
fof(f476,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)))) | ($spl5 | $spl10)),
  inference(forward_demodulation,[],[f475,f146])).
fof(f475,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)))) | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f464,f140])).
fof(f464,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7)),sK0)))) | $spl10),
  inference(superposition,[],[f146,f132])).
fof(f723,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | $spl72),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f725,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) | $spl74),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f726,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | $spl52),
  inference(forward_demodulation,[],[f710,f540])).
fof(f710,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) | $spl52),
  inference(superposition,[],[f132,f540])).
fof(f721,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | $spl52),
  inference(forward_demodulation,[],[f706,f540])).
fof(f706,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0))),sK0)))) | $spl52),
  inference(superposition,[],[f540,f132])).
fof(f718,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)))) | $spl68),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f720,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl70),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f705,plain,(
  s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK9(sK2,sK0),sK0),sK0),sK0)))) | $spl52),
  inference(superposition,[],[f540,f132])).
fof(f197,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl20),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f704,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl28),
  inference(forward_demodulation,[],[f694,f261])).
fof(f694,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl28),
  inference(superposition,[],[f132,f261])).
fof(f701,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl28),
  inference(forward_demodulation,[],[f690,f261])).
fof(f690,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl28),
  inference(superposition,[],[f261,f132])).
fof(f687,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl36),
  inference(forward_demodulation,[],[f677,f312])).
fof(f677,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl36),
  inference(superposition,[],[f132,f312])).
fof(f684,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl36),
  inference(forward_demodulation,[],[f673,f312])).
fof(f673,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl36),
  inference(superposition,[],[f312,f132])).
fof(f192,plain,(
  s('type.list.list'(sK0),sK9(sK7,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK7,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK7,sK0),sK0)))) | $spl16),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s('type.list.list'(sK0),sK9(sK7,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK7,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK7,sK0),sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f261,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl28),
  inference(cnf_transformation,[],[f261_D])).
fof(f261_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f670,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK7,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK7,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK7,sK0),sK0)))) | $spl36),
  inference(superposition,[],[f312,f132])).
fof(f669,plain,(
  $false | ($spl2 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f652,f146])).
fof(f652,plain,(
  s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f651])).
fof(f651,plain,(
  s(sK1,sK3) != s(sK1,sK3) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | ($spl2 | $spl8)),
  inference(superposition,[],[f138,f144])).
fof(f630,plain,(
  $false | ($spl8 | $spl42)),
  inference(trivial_inequality_removal,[],[f629])).
fof(f629,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl8 | $spl42)),
  inference(superposition,[],[f524,f144])).
fof(f524,plain,(
  ( ! [X0,X1] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X0))),s('type.list.list'(sK0),X1)))) ) | $spl42),
  inference(superposition,[],[f390,f131])).
fof(f590,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) | ($spl48 | $spl56)),
  inference(superposition,[],[f549,f132])).
fof(f563,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK2))) | $spl60),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK2))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f551,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK2))) | ($spl18 | $spl56)),
  inference(superposition,[],[f201,f545])).
fof(f547,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) | $spl58),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f548,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl48),
  inference(forward_demodulation,[],[f532,f502])).
fof(f532,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) | $spl48),
  inference(superposition,[],[f132,f502])).
fof(f543,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | $spl48),
  inference(forward_demodulation,[],[f527,f502])).
fof(f527,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK9(sK2,sK0)) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0))),sK0)))) | $spl48),
  inference(superposition,[],[f502,f132])).
fof(f526,plain,(
  s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK9(sK2,sK0),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK9(sK2,sK0),sK0),sK0)))) | $spl48),
  inference(superposition,[],[f502,f132])).
fof(f489,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK2,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK2,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK2,sK0),sK0)))) | $spl5),
  inference(superposition,[],[f462,f132])).
fof(f479,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK0)),s('type.list.list'(sK0),sK9(sK2,sK0)))) | $spl0),
  inference(superposition,[],[f136,f132])).
fof(f460,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f452,f142])).
fof(f452,plain,(
  s(sK1,sK3) = s(sK1,sK4) | ($spl0 | $spl4)),
  inference(superposition,[],[f133,f416])).
fof(f416,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f136,f139])).
fof(f133,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X6,X7] : ! [X21,X19] : s(X7,X21) = s(X7,'const.list.list_CASE_3'(s('type.list.list'(X6),'const.list.NIL_0'),s(X7,X21),s(fun(X6,fun('type.list.list'(X6),X7)),X19)))),
  file('thm.closProps.eqs|split|1','thm.list.list_case_def|split|0')).
fof(f459,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f451,f142])).
fof(f451,plain,(
  s(sK1,sK3) = s(sK1,sK4) | ($spl0 | $spl4)),
  inference(superposition,[],[f416,f133])).
fof(f447,plain,(
  $false | ($spl4 | $spl7 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f445])).
fof(f445,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl4 | $spl7 | $spl8 | $spl10)),
  inference(superposition,[],[f142,f402])).
fof(f402,plain,(
  ( ! [X2] : (s(sK1,sK3) = s(sK1,X2)) ) | ($spl4 | $spl8 | $spl10)),
  inference(superposition,[],[f347,f133])).
fof(f347,plain,(
  ( ! [X7] : (s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,X7),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | ($spl4 | $spl8 | $spl10)),
  inference(superposition,[],[f166,f144])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK6))),s('type.list.list'(sK0),sK7)))) ) | ($spl4 | $spl10)),
  inference(superposition,[],[f131,f164])).
fof(f164,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | ($spl4 | $spl10)),
  inference(forward_demodulation,[],[f146,f139])).
fof(f414,plain,(
  $false | ($spl1 | $spl4 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f403])).
fof(f403,plain,(
  s(sK1,sK3) != s(sK1,sK3) | ($spl1 | $spl4 | $spl8 | $spl10)),
  inference(superposition,[],[f163,f347])).
fof(f163,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f137,f139])).
fof(f389,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0))))) ) | $spl40),
  inference(cnf_transformation,[],[f389_D])).
fof(f389_D,plain,(
  ( ! [X7] : (s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f390,plain,(
  ( ! [X8] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),X8),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) | $spl42),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X8] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),X8),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f388,plain,(
  ( ! [X8,X7] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),X8),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0))))) ) | ($spl1 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f367,f133])).
fof(f367,plain,(
  ( ! [X8,X7] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X8),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7))),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0)),s('type.list.list'(sK0),sK9(happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X7),s(sK0,sK6))),s('type.list.list'(sK0),sK7)),sK0))))) ) | ($spl1 | $spl4 | $spl10)),
  inference(superposition,[],[f218,f166])).
fof(f218,plain,(
  ( ! [X0] : (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),X0),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(X0,sK0)),s('type.list.list'(sK0),sK9(X0,sK0))))) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f163,f132])).
fof(f317,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl22),
  inference(forward_demodulation,[],[f302,f199])).
fof(f302,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | $spl22),
  inference(superposition,[],[f132,f199])).
fof(f284,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | $spl33),
  inference(cnf_transformation,[],[f284_D])).
fof(f284_D,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f316,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl4 | $spl22)),
  inference(forward_demodulation,[],[f299,f199])).
fof(f299,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | ($spl1 | $spl4 | $spl22)),
  inference(superposition,[],[f218,f199])).
fof(f313,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl22),
  inference(forward_demodulation,[],[f295,f199])).
fof(f295,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0))),sK0)))) | $spl22),
  inference(superposition,[],[f199,f132])).
fof(f310,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)))) | $spl34),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f312,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl36),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f292,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0),sK0)))) | $spl22),
  inference(superposition,[],[f199,f132])).
fof(f286,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl4 | $spl24)),
  inference(forward_demodulation,[],[f271,f216])).
fof(f271,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | ($spl1 | $spl4 | $spl24)),
  inference(superposition,[],[f218,f216])).
fof(f199,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | $spl22),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f282,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl4 | $spl12)),
  inference(forward_demodulation,[],[f268,f176])).
fof(f268,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK7),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | ($spl1 | $spl4 | $spl12)),
  inference(superposition,[],[f218,f176])).
fof(f265,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl24),
  inference(forward_demodulation,[],[f251,f216])).
fof(f251,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl24),
  inference(superposition,[],[f132,f216])).
fof(f262,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl24),
  inference(forward_demodulation,[],[f245,f216])).
fof(f245,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl24),
  inference(superposition,[],[f216,f132])).
fof(f259,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | $spl26),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f242,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK9(sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | $spl24),
  inference(superposition,[],[f216,f132])).
fof(f216,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl24),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f217,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl18),
  inference(forward_demodulation,[],[f207,f194])).
fof(f207,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl18),
  inference(superposition,[],[f132,f194])).
fof(f214,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl18),
  inference(forward_demodulation,[],[f203,f194])).
fof(f203,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl18),
  inference(superposition,[],[f194,f132])).
fof(f200,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl12),
  inference(forward_demodulation,[],[f184,f176])).
fof(f184,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | $spl12),
  inference(superposition,[],[f132,f176])).
fof(f195,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | $spl12),
  inference(forward_demodulation,[],[f180,f176])).
fof(f180,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK7) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)),s('type.list.list'(sK0),sK9('const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0))),sK0)))) | $spl12),
  inference(superposition,[],[f176,f132])).
fof(f194,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl18),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f179,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK9(sK7,sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK9(sK7,sK0),sK0)),s('type.list.list'(sK0),sK9(sK9(sK7,sK0),sK0)))) | $spl12),
  inference(superposition,[],[f176,f132])).
fof(f176,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f178,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl14),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f165,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK7,sK0)),s('type.list.list'(sK0),sK9(sK7,sK0)))) | ($spl4 | $spl10)),
  inference(superposition,[],[f164,f132])).
fof(f162,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f161,f133])).
fof(f161,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK1,sK3),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f160,f139])).
fof(f160,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK3),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f137,f141])).
fof(f136,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | $spl0),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f139,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | $spl4),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f146,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | $spl10),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f125,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  (s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | (s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) & s(sK1,sK3) = s(sK1,sK4)) | (s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) & s(sK1,sK3) = s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))))) & (s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | ((s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK2) | s(sK1,sK3) != s(sK1,sK4)) & ! [X6,X7] : (s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X6))),s('type.list.list'(sK0),X7))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f122])).
fof(f122,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X8,X9] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X8),s('type.list.list'(X0),X9))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X8))),s('type.list.list'(X0),X9))))) & (s(X1,X3) != s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | s(X1,X3) != s(X1,X4)) & ! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | s(X1,X3) != s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7)))))))),
  inference(rectify,[],[f121])).
fof(f121,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7))))) & (s(X1,X3) != s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | s(X1,X3) != s(X1,X4)) & ! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | s(X1,X3) != s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7)))))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7)))))) & (s(X1,X3) != s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) | ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | s(X1,X3) != s(X1,X4)) & ! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | s(X1,X3) != s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7)))))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) <~> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7))))))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),X2),s(X1,X4),s(fun(X0,fun('type.list.list'(X0),X1)),X5))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) & s(X1,X3) = s(X1,X4)) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X3) = s(X1,happ(s(fun('type.list.list'(X0),X1),happ(s(fun(X0,fun('type.list.list'(X0),X1)),X5),s(X0,X6))),s('type.list.list'(X0),X7))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s(X7,X9) = s(X7,'const.list.list_CASE_3'(s('type.list.list'(X6),X8),s(X7,X10),s(fun(X6,fun('type.list.list'(X6),X7)),X11))) <=> ((s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X8) & s(X7,X9) = s(X7,X10)) | ? [X12,X13] : (s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & s(X7,X9) = s(X7,happ(s(fun('type.list.list'(X6),X7),happ(s(fun(X6,fun('type.list.list'(X6),X7)),X11),s(X6,X12))),s('type.list.list'(X6),X13))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s(X7,X9) = s(X7,'const.list.list_CASE_3'(s('type.list.list'(X6),X8),s(X7,X10),s(fun(X6,fun('type.list.list'(X6),X7)),X11))) <=> ((s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X8) & s(X7,X9) = s(X7,X10)) | ? [X12,X13] : (s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & s(X7,X9) = s(X7,happ(s(fun('type.list.list'(X6),X7),happ(s(fun(X6,fun('type.list.list'(X6),X7)),X11),s(X6,X12))),s('type.list.list'(X6),X13))))))),
  file('thm.closProps.eqs|split|1','thm.closProps.eqs|split|1')).
fof(f141,plain,(
  s(sK1,sK3) = s(sK1,sK4) | $spl6),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s(sK1,sK3) = s(sK1,sK4) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f126,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | s(sK1,sK3) = s(sK1,sK4) | s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))),
  inference(cnf_transformation,[],[f123])).
fof(f144,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) | $spl8),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f127,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK2) | s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))),
  inference(cnf_transformation,[],[f123])).
fof(f128,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,sK6))),s('type.list.list'(sK0),sK7))) | s(sK1,sK3) = s(sK1,sK4) | s(sK1,sK3) = s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))),
  inference(cnf_transformation,[],[f123])).
fof(f137,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) | $spl1),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f140,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK2) | $spl5),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK2) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f142,plain,(
  s(sK1,sK3) != s(sK1,sK4) | $spl7),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  s(sK1,sK3) != s(sK1,sK4) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f129,plain,(
  s(sK1,sK3) != s(sK1,sK4) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK2) | s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))),
  inference(cnf_transformation,[],[f123])).
fof(f138,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X6))),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f138_D])).
fof(f138_D,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X6))),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f130,plain,(
  ( ! [X6,X7] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.list.list'(sK0),sK1),happ(s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5),s(sK0,X6))),s('type.list.list'(sK0),X7))) | s('type.list.list'(sK0),sK2) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7))) | s(sK1,sK3) != s(sK1,'const.list.list_CASE_3'(s('type.list.list'(sK0),sK2),s(sK1,sK4),s(fun(sK0,fun('type.list.list'(sK0),sK1)),sK5)))) )),
  inference(cnf_transformation,[],[f123])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 87
% Passive clauses: 364
% Generated clauses: 2144
% Final active clauses: 43
% Final passive clauses: 121
% Input formulas: 105
% Initial clauses: 11
% 
% Selected by SInE selection: 6
% SInE iterations: 1
% 
% Duplicate literals: 10
% Trivial inequalities: 18
% Fw subsumption resolutions: 43
% Fw demodulations: 198
% Bw demodulations: 5
% 
% Forward subsumptions: 1435
% 
% Forward superposition: 790
% Backward superposition: 940
% Self superposition: 57
% 
% Split clauses: 71
% Split components: 147
% SAT solver clauses: 85
% SAT solver binary clauses: 3
% 
% Sat splits: 71
% Sat splitting refutations: 13
% 
% Memory used [KB]: 4477
% Time elapsed: 0.134 s
% ------------------------------
% ------------------------------
% Success in time 11.887 s
