Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.conSem
% (51268)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 158
% Passive clauses: 3367
% Generated clauses: 22347
% Final active clauses: 128
% Final passive clauses: 3027
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 774
% Trivial inequalities: 110
% Fw subsumption resolutions: 1579
% Fw demodulations: 3320
% Bw demodulations: 26
% 
% Simple tautologies: 122
% Equational tautologies: 47
% Forward subsumptions: 12606
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 470
% Factoring: 2
% Forward superposition: 8111
% Backward superposition: 7462
% Self superposition: 215
% Equality factoring: 94
% Equality resolution: 7
% 
% Split clauses: 211
% Split components: 422
% SAT solver clauses: 596
% SAT solver unit clauses: 40
% SAT solver binary clauses: 129
% 
% Sat splits: 211
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 596
% TWLsolver calls for satisfiability: 145
% 
% Memory used [KB]: 20852
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.conSem
% (51288)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 1279
% Generated clauses: 18325
% Final active clauses: 97
% Final passive clauses: 817
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 637
% Trivial inequalities: 105
% Fw subsumption resolutions: 1061
% Fw demodulations: 2175
% Bw demodulations: 353
% Condensations: 1638
% 
% Simple tautologies: 118
% Equational tautologies: 44
% Forward subsumptions: 9148
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 454
% Factoring: 1
% Forward superposition: 5485
% Backward superposition: 4721
% Self superposition: 185
% Equality factoring: 98
% Equality resolution: 8
% 
% Split clauses: 110
% Split components: 220
% SAT solver clauses: 1914
% SAT solver unit clauses: 131
% SAT solver binary clauses: 42
% 
% Sat splits: 110
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1914
% TWLsolver calls for satisfiability: 276
% 
% Memory used [KB]: 17142
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.conSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.conSem
% SZS output start Proof for thm.conSem
fof(f37743,plain,(
  $false),
  inference(subsumption_resolution,[],[f37742,f216])).
fof(f216,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.conSem.state_fupdfupds_comp|split|3',aHL_TRUTH)).
fof(f37742,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f37740,f3579])).
fof(f3579,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f3578])).
fof(f3578,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f3514,f300])).
fof(f300,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f219])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.conSem.state_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f218,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.conSem.state_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f3514,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f3003,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f219,f219])).
fof(f3003,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2914,f263])).
fof(f263,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X14] : ! [X17,X22] : s(X14,X22) = s(X14,'const.pair.SND_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X17),s(X14,X22)))))),
  file('thm.conSem.state_fupdfupds_comp|split|3','thm.pair.SND')).
fof(f2914,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f263,f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.conSem.state_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f37740,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),fun(fun(sK0,'type.conSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f37673,f13497])).
fof(f13497,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.conSem.state'(sK1)),fun(fun(sK0,'type.conSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f215,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f215,plain,(
  s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))) != s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))))),s(fun(sK0,'type.conSem.state'(sK1)),sK2))) != s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK4))),s(fun(sK0,'type.conSem.state'(sK1)),'const.combin.o_2'(s(fun('type.conSem.state'(sK1),'type.conSem.state'(sK1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),sK3))),s(fun(sK0,'type.conSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) != s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) = s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))))),s(fun(X0,'type.conSem.state'(X1)),X2))) = s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X4))),s(fun(X0,'type.conSem.state'(X1)),'const.combin.o_2'(s(fun('type.conSem.state'(X1),'type.conSem.state'(X1)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X3))),s(fun(X0,'type.conSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9))))),s(fun(X6,'type.conSem.state'(X7)),X8))) = s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10))),s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9))),s(fun(X6,'type.conSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9))))),s(fun(X6,'type.conSem.state'(X7)),X8))) = s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10))),s(fun(X6,'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9))),s(fun(X6,'type.conSem.state'(X7)),X8)))))),
  file('thm.conSem.state_fupdfupds_comp|split|3','thm.conSem.state_fupdfupds_comp|split|3')).
fof(f37673,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2))))),s(fun(X3,'type.conSem.state'(X0)),X4))) = s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1))),s(fun(X3,'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2))),s(fun(X3,'type.conSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f277,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1))))) = s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.combin.o_2'(s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X2))),s(fun('type.conSem.state'(X0),'type.conSem.state'(X0)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9))))) = s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.combin.o_2'(s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X10))),s(fun('type.conSem.state'(X7),'type.conSem.state'(X7)),'const.conSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v'))),X9)))))),
  file('thm.conSem.state_fupdfupds_comp|split|3','thm.conSem.state_fupdfupds_comp|split|2')).
fof(f277,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14,X6,X21,X15] : ! [X10,X9,X8] : s(fun(X21,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X21,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X21,X15),X8))))) = s(fun(X21,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X21,X15),X8)))),
  file('thm.conSem.state_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.conSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 242
% Passive clauses: 3813
% Generated clauses: 41089
% Final active clauses: 112
% Final passive clauses: 1401
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 1057
% Trivial inequalities: 63
% Fw subsumption resolutions: 1128
% Fw demodulations: 6653
% Bw demodulations: 2110
% Condensations: 1533
% Global subsumptions: 354
% 
% Simple tautologies: 127
% Equational tautologies: 180
% Forward subsumptions: 25634
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 71
% 
% Unit resulting resolution: 1144
% Factoring: 1
% Forward superposition: 12507
% Backward superposition: 10296
% Self superposition: 281
% Equality factoring: 94
% Equality resolution: 2
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 4689
% SAT solver unit clauses: 1083
% SAT solver binary clauses: 2091
% 
% Sat splits: 39
% Sat splitting refutations: 442
% 
% Memory used [KB]: 44775
% Time elapsed: 2.0000 s
% ------------------------------
% ------------------------------
% Success in time 3.33 s
