Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (42811)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 73
% Passive clauses: 2450
% Generated clauses: 8097
% Final active clauses: 69
% Final passive clauses: 2272
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 384
% Trivial inequalities: 57
% Fw subsumption resolutions: 912
% Fw demodulations: 201
% Bw demodulations: 1
% 
% Simple tautologies: 134
% Equational tautologies: 71
% Forward subsumptions: 3376
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 191
% Factoring: 4
% Forward superposition: 2894
% Backward superposition: 3316
% Self superposition: 244
% Equality factoring: 178
% Equality resolution: 7
% 
% Split clauses: 163
% Split components: 326
% SAT solver clauses: 610
% SAT solver unit clauses: 17
% SAT solver binary clauses: 164
% 
% Sat splits: 163
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 610
% TWLsolver calls for satisfiability: 107
% 
% Memory used [KB]: 8315
% Time elapsed: 0.608 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f5126,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5125,f4347,f4348,f93,f121,f70,f98,f96,f4559,f5089,f4990,f4991,f85,f115,f77,f107,f75,f105,f69,f97,f95])).
fof(f95,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) = s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) | $spl0),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) = s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f97,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK2) | $spl2),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f69,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK2) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) = s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2)))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) = s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK2)) & (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) != s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) | s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2] : ((s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) | s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2)) & (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) != s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) | s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X2)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ? [X0,X1,X2] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) <~> s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) <=> s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) <=> s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X7))) = s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X8))) <=> s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X7))) = s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X8))) <=> s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X8))),
  file('thm.stackSem.result_|49||49||split|0','thm.stackSem.result_|49||49||split|0')).
fof(f105,plain,(
  ( ! [X0,X1] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.stackSem.result'(X0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X0),X1)))) ) | $spl12),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X0,X1] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.stackSem.result'(X0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X0),X1)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f75,plain,(
  ( ! [X0,X1] : (s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.stackSem.result'(X0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X0),X1)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.stackSem.result'(X0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.stackSem.result'(X0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9] : s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X9))) = s('type.stackSem.result'(X6),happ(s(fun('type.wordSem.word_loc'(X6),'type.stackSem.result'(X6)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(X6),X9)))),
  file('thm.stackSem.result_|49||49||split|0','ARITY_0_1_const.stackSem.Result')).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f77,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ! [X7,X10] : s(X11,X7) = s(X11,'const.ind_type.FCONS_3'(s(X11,X7),s(fun('type.num.num',X11),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.stackSem.result_|49||49||split|0','thm.ind_type.FCONS|split|0')).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) | $spl32),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : ! [X9,X31,X32] : s(X11,'const.ind_type.FCONS_3'(s(X11,X9),s(fun('type.num.num',X11),X31),s('type.num.num',X32))) = s(X11,happ(s(fun('type.num.num',X11),happ(s(fun(fun('type.num.num',X11),fun('type.num.num',X11)),happ(s(fun(X11,fun(fun('type.num.num',X11),fun('type.num.num',X11))),'const.ind_type.FCONS_0'),s(X11,X9))),s(fun('type.num.num',X11),X31))),s('type.num.num',X32)))),
  file('thm.stackSem.result_|49||49||split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f4991,plain,(
  ( ! [X4,X5] : (s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK1),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5))) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK2),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5)))) ) | $spl778),
  inference(cnf_transformation,[],[f4991_D])).
fof(f4991_D,plain,(
  ( ! [X4,X5] : (s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK1),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5))) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK2),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5)))) ) <=> ~$spl778),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl778])])).
fof(f4990,plain,(
  ( ! [X4,X5] : (s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK1),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5))) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK2),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5)))) ) | ($spl32 | $spl656)),
  inference(forward_demodulation,[],[f4848,f115])).
fof(f4848,plain,(
  ( ! [X4,X5] : (s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK2),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4),s('type.num.num',X5))) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.num.num','type.wordSem.word_loc'(sK0)),happ(s(fun(fun('type.num.num','type.wordSem.word_loc'(sK0)),fun('type.num.num','type.wordSem.word_loc'(sK0))),happ(s(fun('type.wordSem.word_loc'(sK0),fun(fun('type.num.num','type.wordSem.word_loc'(sK0)),fun('type.num.num','type.wordSem.word_loc'(sK0)))),'const.ind_type.FCONS_0'),s('type.wordSem.word_loc'(sK0),sK1))),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X4))),s('type.num.num',X5)))) ) | ($spl32 | $spl656)),
  inference(superposition,[],[f115,f4348])).
fof(f5089,plain,(
  $false | ($spl3 | $spl16 | $spl778)),
  inference(subsumption_resolution,[],[f5088,f98])).
fof(f5088,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK2) | ($spl16 | $spl778)),
  inference(forward_demodulation,[],[f5065,f107])).
fof(f5065,plain,(
  ( ! [X2] : (s('type.wordSem.word_loc'(sK0),sK2) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),sK1),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X2),s('type.num.num','const.num.|48|_0')))) ) | ($spl16 | $spl778)),
  inference(superposition,[],[f4991,f107])).
fof(f4559,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) = s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) | ($spl12 | $spl656)),
  inference(forward_demodulation,[],[f4494,f105])).
fof(f4494,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) = s('type.stackSem.result'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.stackSem.result'(sK0)),'const.stackSem.Result_0'),s('type.wordSem.word_loc'(sK0),sK1))) | ($spl12 | $spl656)),
  inference(superposition,[],[f105,f4348])).
fof(f96,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) != s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) | $spl1),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) != s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f98,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK2) | $spl3),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f70,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK2) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) != s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2)))),
  inference(cnf_transformation,[],[f65])).
fof(f121,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) ) | $spl44),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f93,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X6] : ! [X7,X10,X12,X13,X14,X15] : s(X11,happ(s(fun('type.wordSem.word_loc'(X6),X11),X10),s('type.wordSem.word_loc'(X6),X7))) = s(X11,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),X11),X10),s(fun('type.wordSem.word_loc'(X6),X11),X12),s(fun('type.wordSem.word_loc'(X6),X11),X13),s(X11,X14),s(X11,X15)))),
  file('thm.stackSem.result_|49||49||split|0','thm.stackSem.result_case_def|split|0')).
fof(f4348,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl656),
  inference(cnf_transformation,[],[f4348_D])).
fof(f4348_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK2)))) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f4347,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK2)))) ) | ($spl0 | $spl44)),
  inference(forward_demodulation,[],[f4126,f121])).
fof(f4126,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(X0,X4),s(X0,X5)))) ) | ($spl0 | $spl44)),
  inference(superposition,[],[f121,f95])).
fof(f5125,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f5124])).
fof(f5124,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) != s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f96,f97])).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 54
% Passive clauses: 415
% Generated clauses: 4967
% Final active clauses: 47
% Final passive clauses: 331
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 366
% Trivial inequalities: 37
% Fw subsumption resolutions: 231
% Fw demodulations: 143
% Bw demodulations: 30
% Condensations: 819
% 
% Simple tautologies: 74
% Equational tautologies: 38
% Forward subsumptions: 2169
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 99
% Factoring: 2
% Forward superposition: 1395
% Backward superposition: 1177
% Self superposition: 129
% Equality factoring: 94
% Equality resolution: 6
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 712
% SAT solver unit clauses: 39
% SAT solver binary clauses: 32
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 7
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 710
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 6524
% Time elapsed: 0.355 s
% ------------------------------
% ------------------------------
% Success in time 0.98 s
