Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvlSem
% (67944)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 173
% Passive clauses: 3600
% Generated clauses: 19668
% Final active clauses: 163
% Final passive clauses: 3239
% Input formulas: 69
% Initial clauses: 143
% 
% Duplicate literals: 671
% Trivial inequalities: 475
% Fw subsumption resolutions: 1974
% Fw demodulations: 2965
% Bw demodulations: 7
% 
% Simple tautologies: 434
% Equational tautologies: 55
% Forward subsumptions: 9027
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 329
% Factoring: 1
% Forward superposition: 6482
% Backward superposition: 6250
% Self superposition: 159
% Equality factoring: 88
% Equality resolution: 41
% 
% Split clauses: 206
% Split components: 424
% SAT solver clauses: 486
% SAT solver unit clauses: 14
% SAT solver binary clauses: 105
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 206
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 483
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvlSem
% (67955)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1337
% Generated clauses: 14824
% Final active clauses: 135
% Final passive clauses: 1156
% Input formulas: 69
% Initial clauses: 143
% 
% Duplicate literals: 429
% Trivial inequalities: 396
% Fw subsumption resolutions: 949
% Fw demodulations: 2007
% Bw demodulations: 38
% Condensations: 730
% Global subsumptions: 8
% 
% Simple tautologies: 255
% Equational tautologies: 45
% Forward subsumptions: 6553
% 
% Binary resolution: 172
% Factoring: 1
% Forward superposition: 4484
% Backward superposition: 3866
% Self superposition: 130
% Equality factoring: 60
% Equality resolution: 34
% 
% Split clauses: 47
% Split components: 94
% SAT solver clauses: 2218
% SAT solver unit clauses: 152
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 47
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2184
% TWLsolver calls for satisfiability: 602
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvlSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvlSem
% SZS output start Proof for thm.bvlSem
fof(f20265,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f253,f417,f415,f252,f416,f414,f420,f1557,f1568,f2402,f2408,f2405,f5234,f5241,f6291,f6293,f6295,f6299,f6301,f6303,f6305,f6309,f6311,f6313,f6315,f6318,f6320,f6322,f6324,f6327,f6329,f6331,f6333,f6335,f6337,f6339,f6361,f6363,f6393,f6397,f6399,f6497,f6499,f6501,f6503,f6600,f6601,f6602,f6603,f6604,f7066,f7067,f7070,f7071,f9666,f9785,f9787,f9789,f9790,f9786,f9794,f9795,f9788,f6985,f7005,f4897,f4917,f9833,f1567,f9834,f10125,f10127,f10136,f10139,f10140,f10141,f10196,f10198,f10200,f10202,f10206,f10208,f10210,f10212,f10214,f10216,f10218,f10220,f12481,f2403,f12488,f12482,f2409,f13255,f13272,f13270,f16401,f16476,f16819,f9733,f17519,f17596,f9732,f20156,f20169,f20172,f20220,f20246])).
fof(f20246,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20245,f929])).
fof(f929,plain,(
  s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X6))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_|49||49||split|3')).
fof(f20245,plain,(
  s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK1))) | $spl0),
  inference(forward_demodulation,[],[f19978,f260])).
fof(f260,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X8))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X8)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_1_const.bvlSem.CodePtr')).
fof(f19978,plain,(
  s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK1))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f260,f19593])).
fof(f19593,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f19214,f323])).
fof(f323,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_case_def|split|4')).
fof(f19214,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X1)))) ) | $spl0),
  inference(superposition,[],[f323,f414])).
fof(f20220,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20219,f2224])).
fof(f2224,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',sK0),s('type.list.list'('type.bvlSem.v'),X0))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.bvlSem.v'),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ((s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.bvlSem.v'),X1) = s('type.list.list'('type.bvlSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.bvlSem.v'),X1) != s('type.list.list'('type.bvlSem.v'),X3) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : ((s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.bvlSem.v'),X1) = s('type.list.list'('type.bvlSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.bvlSem.v'),X1) != s('type.list.list'('type.bvlSem.v'),X3)) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3)))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.bvlSem.v'),X1) = s('type.list.list'('type.bvlSem.v'),X3)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X15,X16,X17,X18] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X15),s('type.list.list'('type.bvlSem.v'),X16))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X17),s('type.list.list'('type.bvlSem.v'),X18))) <=> (s('type.num.num',X15) = s('type.num.num',X17) & s('type.list.list'('type.bvlSem.v'),X16) = s('type.list.list'('type.bvlSem.v'),X18)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_|49||49||split|2')).
fof(f20219,plain,(
  ( ! [X183] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',sK0),s('type.list.list'('type.bvlSem.v'),X183))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.bvlSem.v'),X183)))) ) | $spl0),
  inference(forward_demodulation,[],[f19932,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.bvlSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.bvlSem.v'),X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X19] : s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X8),s('type.list.list'('type.bvlSem.v'),X19))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X8))),s('type.list.list'('type.bvlSem.v'),X19)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_2_const.bvlSem.Block')).
fof(f19932,plain,(
  ( ! [X183] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.bvlSem.v'),X183))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',sK0))),s('type.list.list'('type.bvlSem.v'),X183)))) ) | $spl0),
  inference(superposition,[],[f324,f19593])).
fof(f20172,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20171,f417])).
fof(f20171,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f19851,f305])).
fof(f305,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X23] : ! [X25] : s(X23,X25) = s(X23,happ(s(fun(X23,X23),'_i_0_0'),s(X23,X25)))),
  file('thm.bvlSem.v_|49||49||split|4','_i0_ax')).
fof(f19851,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f305,f19593])).
fof(f20169,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20168,f929])).
fof(f20168,plain,(
  s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK1))) | $spl0),
  inference(forward_demodulation,[],[f19849,f260])).
fof(f19849,plain,(
  s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK1))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f19593,f260])).
fof(f20156,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20155,f417])).
fof(f20155,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f19828,f305])).
fof(f19828,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f19593,f305])).
fof(f9732,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl20),
  inference(cnf_transformation,[],[f9732_D])).
fof(f9732_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f17596,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) | $spl38),
  inference(cnf_transformation,[],[f17596_D])).
fof(f17596_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f17519,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f17348])).
fof(f17348,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f7425,f16022])).
fof(f16022,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f16007])).
fof(f16007,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f15924,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f266,f266])).
fof(f266,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvlSem.v_|49||49||split|4',aHL_BOOL_CASES)).
fof(f15924,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl10),
  inference(forward_demodulation,[],[f15842,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X27,X9] : ! [X28] : s(X9,X28) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X27),'const.sum.INL_1'(s(X9,X28)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.sum.OUTL')).
fof(f15842,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f9793,f8659])).
fof(f8659,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f8579,f256])).
fof(f256,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvlSem.v_|49||49||split|4',aHL_FALSITY)).
fof(f8579,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(superposition,[],[f267,f3047])).
fof(f3047,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2928,f2434])).
fof(f2434,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f421,f2405,f2219])).
fof(f2219,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2218,f305])).
fof(f2218,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2186,f254])).
fof(f254,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.v_|49||49||split|4',aHL_TRUTH)).
fof(f2186,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f329,f461])).
fof(f461,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f305,f266])).
fof(f329,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X32,X33] : ((p(s(bool,X32)) => p(s(bool,X33))) => ((p(s(bool,X33)) => p(s(bool,X32))) => s(bool,X32) = s(bool,X33)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f421,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f256,f267])).
fof(f2928,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2517,f305])).
fof(f2517,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f2492,f254])).
fof(f2492,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f2219,f2434])).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X20] : (s(bool,'T_0') = s(bool,X20) <=> p(s(bool,X20)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f9793,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2405,f2434,f2428,f2516,f2517,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3047,f2937,f3083,f3099,f3454,f3455,f3456,f3506,f3507,f3509,f3510,f3512,f3513,f3514,f3549,f3578,f3580,f3581,f3583,f3584,f3585,f3692,f3693,f3766,f3767,f4207,f4214,f3862,f3863,f4241,f4274,f4275,f4277,f4278,f4280,f4281,f4282,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4816,f4646,f4647,f4700,f4702,f4703,f4818,f4718,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5074,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f5004,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f3085,f2407,f346,f345,f317,f316,f313,f4212,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f3682,f3852,f5235,f6605,f4899,f4511,f4393,f4276,f3579,f3508,f3049,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f4279,f3582,f3511,f3052,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7066,f7068,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7953,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8328,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8556,f8559,f8563,f8659,f8667,f8753,f8758,f8760,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8889,f8890,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9045,f9046,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9151,f9152,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9785,f9787,f9789,f9790,f9792])).
fof(f9792,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9791])).
fof(f9791,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f8266,f3767])).
fof(f8266,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f7806,f2937])).
fof(f9117,plain,(
  ( ! [X99,X97,X95,X98,X96] : (s(bool,'F_0') = s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f266,f365])).
fof(f9203,plain,(
  ( ! [X648,X650,X644,X646,X649,X645,X647,X643] : (s('type.list.list'('type.bvlSem.v'),X648) != s('type.list.list'('type.bvlSem.v'),happ(s(fun(X644,'type.list.list'('type.bvlSem.v')),X645),s(X644,happ(s(fun(X643,X644),X646),s(X643,X647))))) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X650),s('type.list.list'('type.bvlSem.v'),X648))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X649),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X643,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X644,'type.list.list'('type.bvlSem.v')),X645),s(fun(X643,X644),X646))),s(X643,X647))))) | s('type.num.num',X649) != s('type.num.num',X650)) )),
  inference(superposition,[],[f392,f365])).
fof(f9201,plain,(
  ( ! [X633,X629,X631,X627,X632,X634,X628,X630] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X633),s('type.list.list'('type.bvlSem.v'),X634))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X632),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X628,'type.list.list'('type.bvlSem.v')),X629),s(X628,happ(s(fun(X627,X628),X630),s(X627,X631))))))) | s('type.list.list'('type.bvlSem.v'),X634) = s('type.list.list'('type.bvlSem.v'),happ(s(fun(X627,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X628,'type.list.list'('type.bvlSem.v')),X629),s(fun(X627,X628),X630))),s(X627,X631)))) )),
  inference(superposition,[],[f391,f365])).
fof(f9192,plain,(
  ( ! [X573,X575,X571,X572,X574,X570] : (s('type.integer.int',X575) != s('type.integer.int',happ(s(fun(X571,'type.integer.int'),X572),s(X571,happ(s(fun(X570,X571),X573),s(X570,X574))))) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X575))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X570,'type.integer.int'),'_dst_x0x1_2'(s(fun(X571,'type.integer.int'),X572),s(fun(X570,X571),X573))),s(X570,X574)))))) )),
  inference(superposition,[],[f333,f365])).
fof(f9190,plain,(
  ( ! [X561,X563,X559,X560,X562,X558] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X563))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X559,'type.integer.int'),X560),s(X559,happ(s(fun(X558,X559),X561),s(X558,X562))))))) | s('type.integer.int',X563) = s('type.integer.int',happ(s(fun(X558,'type.integer.int'),'_dst_x0x1_2'(s(fun(X559,'type.integer.int'),X560),s(fun(X558,X559),X561))),s(X558,X562)))) )),
  inference(superposition,[],[f332,f365])).
fof(f9183,plain,(
  ( ! [X511,X516,X518,X512,X514,X517,X513,X515] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X517),s('type.list.list'('type.bvlSem.v'),X518))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X512,'type.num.num'),X513),s(X512,happ(s(fun(X511,X512),X514),s(X511,X515))))),s('type.list.list'('type.bvlSem.v'),X516))) | s('type.num.num',X517) = s('type.num.num',happ(s(fun(X511,'type.num.num'),'_dst_x0x1_2'(s(fun(X512,'type.num.num'),X513),s(fun(X511,X512),X514))),s(X511,X515)))) )),
  inference(superposition,[],[f390,f365])).
fof(f9181,plain,(
  ( ! [X499,X500,X502,X504,X501,X503] : (s('type.num.num',X504) != s('type.num.num',happ(s(fun(X500,'type.num.num'),X501),s(X500,happ(s(fun(X499,X500),X502),s(X499,X503))))) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X504))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X499,'type.num.num'),'_dst_x0x1_2'(s(fun(X500,'type.num.num'),X501),s(fun(X499,X500),X502))),s(X499,X503)))))) )),
  inference(superposition,[],[f331,f365])).
fof(f9179,plain,(
  ( ! [X492,X488,X490,X487,X489,X491] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X492))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X488,'type.num.num'),X489),s(X488,happ(s(fun(X487,X488),X490),s(X487,X491))))))) | s('type.num.num',X492) = s('type.num.num',happ(s(fun(X487,'type.num.num'),'_dst_x0x1_2'(s(fun(X488,'type.num.num'),X489),s(fun(X487,X488),X490))),s(X487,X491)))) )),
  inference(superposition,[],[f330,f365])).
fof(f9170,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X426,X425),X432),s(fun(X425,X426),happ(s(fun(X428,fun(X425,X426)),X429),s(X428,happ(s(fun(X427,X428),X430),s(X427,X431)))))))) | s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) )),
  inference(superposition,[],[f354,f365])).
fof(f9169,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X416,X417),happ(s(fun(X419,fun(X416,X417)),X420),s(X419,happ(s(fun(X418,X419),X421),s(X418,X422))))),s(fun(X417,X416),X423)))) | s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) )),
  inference(superposition,[],[f354,f365])).
fof(f9168,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X408,X407),X414),s(fun(X407,X408),happ(s(fun(X410,fun(X407,X408)),X411),s(X410,happ(s(fun(X409,X410),X412),s(X409,X413)))))))) | s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) )),
  inference(superposition,[],[f353,f365])).
fof(f9167,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X398,X399),happ(s(fun(X401,fun(X398,X399)),X402),s(X401,happ(s(fun(X400,X401),X403),s(X400,X404))))),s(fun(X399,X398),X405)))) | s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) )),
  inference(superposition,[],[f353,f365])).
fof(f9165,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (s(fun(X380,X381),X387) != s(fun(X380,X381),happ(s(fun(X383,fun(X380,X381)),X384),s(X383,happ(s(fun(X382,X383),X385),s(X382,X386))))) | s(X381,happ(s(fun(X380,X381),X387),s(X380,X388))) = s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))) )),
  inference(superposition,[],[f351,f365])).
fof(f9164,plain,(
  ( ! [X377,X379,X374,X376,X378,X375] : (p(s(bool,happ(s(fun('type.list.list'(X374),bool),happ(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(X376,happ(s(fun(X375,X376),X378),s(X375,X379))))),s('type.list.list'(X374),sK2(happ(s(fun(X375,fun('type.list.list'(X374),bool)),'_dst_x0x1_2'(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(fun(X375,X376),X378))),s(X375,X379)),X374))))) | ~p(s(bool,happ(s(fun('type.list.list'(X374),bool),happ(s(fun(X375,fun('type.list.list'(X374),bool)),'_dst_x0x1_2'(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(fun(X375,X376),X378))),s(X375,X379))),s('type.list.list'(X374),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f365])).
fof(f9163,plain,(
  ( ! [X372,X368,X370,X367,X373,X369,X371] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X367,bool),happ(s(fun(X369,fun(X367,bool)),X370),s(X369,happ(s(fun(X368,X369),X371),s(X368,X372))))))) | s(X367,X373) = s(X367,sK4(X373,happ(s(fun(X368,fun(X367,bool)),'_dst_x0x1_2'(s(fun(X369,fun(X367,bool)),X370),s(fun(X368,X369),X371))),s(X368,X372)),X367))) )),
  inference(superposition,[],[f866,f365])).
fof(f9162,plain,(
  ( ! [X364,X366,X360,X362,X365,X361,X363] : (p(s(bool,'|dollar|exists_1'(s(fun(X360,bool),happ(s(fun(X362,fun(X360,bool)),X363),s(X362,happ(s(fun(X361,X362),X364),s(X361,X365)))))))) | s(X360,X366) = s(X360,sK4(X366,happ(s(fun(X361,fun(X360,bool)),'_dst_x0x1_2'(s(fun(X362,fun(X360,bool)),X363),s(fun(X361,X362),X364))),s(X361,X365)),X360))) )),
  inference(superposition,[],[f314,f365])).
fof(f9161,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (~p(s(bool,happ(s(fun(X353,bool),happ(s(fun(X355,fun(X353,bool)),X356),s(X355,happ(s(fun(X354,X355),X357),s(X354,X358))))),s(X353,X359)))) | p(s(bool,happ(s(fun(X353,bool),happ(s(fun(X354,fun(X353,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X353,bool)),X356),s(fun(X354,X355),X357))),s(X354,X358))),s(X353,'const.min.@_1'(s(fun(X353,bool),happ(s(fun(X354,fun(X353,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X353,bool)),X356),s(fun(X354,X355),X357))),s(X354,X358))))))))) )),
  inference(superposition,[],[f313,f365])).
fof(f9160,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (~p(s(bool,happ(s(fun(X346,bool),happ(s(fun(X348,fun(X346,bool)),X349),s(X348,happ(s(fun(X347,X348),X350),s(X347,X351))))),s(X346,X352)))) | p(s(bool,'|dollar|exists_1'(s(fun(X346,bool),happ(s(fun(X347,fun(X346,bool)),'_dst_x0x1_2'(s(fun(X348,fun(X346,bool)),X349),s(fun(X347,X348),X350))),s(X347,X351))))))) )),
  inference(superposition,[],[f4817,f365])).
fof(f9157,plain,(
  ( ! [X325,X327,X329,X326,X328,X330] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X330),'const.sum.INL_1'(s(bool,happ(s(fun(X326,bool),X327),s(X326,happ(s(fun(X325,X326),X328),s(X325,X329))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X325,bool),'_dst_x0x1_2'(s(fun(X326,bool),X327),s(fun(X325,X326),X328))),s(X325,X329)))) )),
  inference(superposition,[],[f7806,f365])).
fof(f9156,plain,(
  ( ! [X321,X323,X324,X320,X322,X319] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X324,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X320,bool),X321),s(X320,happ(s(fun(X319,X320),X322),s(X319,X323))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X319,bool),'_dst_x0x1_2'(s(fun(X320,bool),X321),s(fun(X319,X320),X322))),s(X319,X323)))) )),
  inference(superposition,[],[f7662,f365])).
fof(f9153,plain,(
  ( ! [X300,X302,X304,X306,X301,X303,X305] : (~p(s(bool,happ(s(fun(X301,bool),X302),s(X301,happ(s(fun(X300,X301),X303),s(X300,X304)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X305),bool),sK16(X306,happ(s(fun(X300,bool),'_dst_x0x1_2'(s(fun(X301,bool),X302),s(fun(X300,X301),X303))),s(X300,X304)),X305,bool))))) )),
  inference(superposition,[],[f4999,f365])).
fof(f9152,plain,(
  ( ! [X294,X296,X298,X295,X297,X299] : (s(bool,'T_0') != s(bool,happ(s(fun(X295,bool),X296),s(X295,happ(s(fun(X294,X295),X297),s(X294,X298))))) | s(bool,X299) = s(bool,happ(s(fun(X294,bool),'_dst_x0x1_2'(s(fun(X295,bool),X296),s(fun(X294,X295),X297))),s(X294,X298))) | s(bool,'F_0') = s(bool,X299)) ) | $spl10),
  inference(superposition,[],[f3767,f365])).
fof(f9151,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X290,bool),X291),s(X290,happ(s(fun(X289,X290),X292),s(X289,X293))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X289,bool),'_dst_x0x1_2'(s(fun(X290,bool),X291),s(fun(X289,X290),X292))),s(X289,X293)))) ) | $spl10),
  inference(superposition,[],[f3047,f365])).
fof(f9140,plain,(
  ( ! [X231,X229,X227,X230,X228,X226] : (p(s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X226,X227),X229),s(X226,X230)))))) | s(bool,X231) != s(bool,happ(s(fun(X226,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X226,X227),X229))),s(X226,X230))) | s(bool,'F_0') = s(bool,X231)) )),
  inference(superposition,[],[f2221,f365])).
fof(f9139,plain,(
  ( ! [X222,X220,X225,X223,X221,X224] : (~p(s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X220,X221),X223),s(X220,X224)))))) | s(bool,X225) = s(bool,happ(s(fun(X220,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X220,X221),X223))),s(X220,X224))) | s(bool,'F_0') = s(bool,X225)) )),
  inference(superposition,[],[f2219,f365])).
fof(f9138,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X219),'const.sum.INL_1'(s(bool,happ(s(fun(X215,bool),X216),s(X215,happ(s(fun(X214,X215),X217),s(X214,X218))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X214,bool),'_dst_x0x1_2'(s(fun(X215,bool),X216),s(fun(X214,X215),X217))),s(X214,X218)))) )),
  inference(superposition,[],[f2217,f365])).
fof(f9137,plain,(
  ( ! [X212,X210,X208,X213,X211,X209] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X213,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X209,bool),X210),s(X209,happ(s(fun(X208,X209),X211),s(X208,X212))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X208,bool),'_dst_x0x1_2'(s(fun(X209,bool),X210),s(fun(X208,X209),X211))),s(X208,X212)))) )),
  inference(superposition,[],[f2213,f365])).
fof(f9134,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X196),'const.sum.INL_1'(s(bool,happ(s(fun(X192,bool),X193),s(X192,happ(s(fun(X191,X192),X194),s(X191,X195))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X191,bool),'_dst_x0x1_2'(s(fun(X192,bool),X193),s(fun(X191,X192),X194))),s(X191,X195)))) )),
  inference(superposition,[],[f529,f365])).
fof(f9133,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X190,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X186,bool),X187),s(X186,happ(s(fun(X185,X186),X188),s(X185,X189))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X185,bool),'_dst_x0x1_2'(s(fun(X186,bool),X187),s(fun(X185,X186),X188))),s(X185,X189)))) )),
  inference(superposition,[],[f502,f365])).
fof(f9132,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (p(s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183)))))) | s(bool,X184) = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f483,f365])).
fof(f9131,plain,(
  ( ! [X177,X175,X178,X176,X174] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X175,bool),X176),s(X175,happ(s(fun(X174,X175),X177),s(X174,X178))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X174,bool),'_dst_x0x1_2'(s(fun(X175,bool),X176),s(fun(X174,X175),X177))),s(X174,X178)))) )),
  inference(superposition,[],[f461,f365])).
fof(f9130,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (~p(s(bool,happ(s(fun(X169,bool),X170),s(X169,happ(s(fun(X168,X169),X171),s(X168,X172)))))) | s(bool,X173) != s(bool,happ(s(fun(X168,bool),'_dst_x0x1_2'(s(fun(X169,bool),X170),s(fun(X168,X169),X171))),s(X168,X172))) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f436,f365])).
fof(f9125,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (~p(s(bool,happ(s(fun(X139,bool),X140),s(X139,happ(s(fun(X138,X139),X141),s(X138,X142)))))) | s(bool,X143) != s(bool,happ(s(fun(X138,bool),'_dst_x0x1_2'(s(fun(X139,bool),X140),s(fun(X138,X139),X141))),s(X138,X142))) | p(s(bool,X143))) )),
  inference(superposition,[],[f406,f365])).
fof(f9124,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(X133,bool),X134),s(X133,happ(s(fun(X132,X133),X135),s(X132,X136)))))) | ~p(s(bool,X137)) | s(bool,X137) = s(bool,happ(s(fun(X132,bool),'_dst_x0x1_2'(s(fun(X133,bool),X134),s(fun(X132,X133),X135))),s(X132,X136)))) )),
  inference(superposition,[],[f329,f365])).
fof(f9122,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | p(s(bool,X125)) | s(bool,X125) = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f326,f365])).
fof(f9121,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119)))) )),
  inference(superposition,[],[f272,f365])).
fof(f9120,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'F_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114))))) | ~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) )),
  inference(superposition,[],[f271,f365])).
fof(f9119,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f268,f365])).
fof(f9118,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,happ(s(fun(X101,bool),X102),s(X101,happ(s(fun(X100,X101),X103),s(X100,X104)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X100,bool),'_dst_x0x1_2'(s(fun(X101,bool),X102),s(fun(X100,X101),X103))),s(X100,X104)))) )),
  inference(superposition,[],[f267,f365])).
fof(f9116,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun('type.list.list'(X91),X92),X94),s('type.list.list'(X91),sK2('_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun('type.list.list'(X91),X92),X94)),X91))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X91),bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun('type.list.list'(X91),X92),X94))),s('type.list.list'(X91),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f365])).
fof(f9115,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,happ(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))),s(X86,'const.min.@_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))))) )),
  inference(superposition,[],[f313,f365])).
fof(f9114,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,X85)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f4817,f365])).
fof(f9098,plain,(
  ( ! [X681,X683,X677,X679,X684,X680,X682,X678] : (s('type.list.list'('type.bvlSem.v'),X682) != s('type.list.list'('type.bvlSem.v'),happ(s(fun(X679,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X677,'type.list.list'('type.bvlSem.v')),X678),s(fun(X679,X677),X680))),s(X679,X681))) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X684),s('type.list.list'('type.bvlSem.v'),X682))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X683),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X677,'type.list.list'('type.bvlSem.v')),X678),s(X677,happ(s(fun(X679,X677),X680),s(X679,X681))))))) | s('type.num.num',X683) != s('type.num.num',X684)) )),
  inference(superposition,[],[f392,f365])).
fof(f9096,plain,(
  ( ! [X666,X668,X662,X664,X667,X661,X663,X665] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X667),s('type.list.list'('type.bvlSem.v'),X668))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X666),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X663,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X661,'type.list.list'('type.bvlSem.v')),X662),s(fun(X663,X661),X664))),s(X663,X665))))) | s('type.list.list'('type.bvlSem.v'),X668) = s('type.list.list'('type.bvlSem.v'),happ(s(fun(X661,'type.list.list'('type.bvlSem.v')),X662),s(X661,happ(s(fun(X663,X661),X664),s(X663,X665)))))) )),
  inference(superposition,[],[f391,f365])).
fof(f9087,plain,(
  ( ! [X606,X609,X605,X607,X608,X604] : (s('type.integer.int',X609) != s('type.integer.int',happ(s(fun(X606,'type.integer.int'),'_dst_x0x1_2'(s(fun(X604,'type.integer.int'),X605),s(fun(X606,X604),X607))),s(X606,X608))) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X609))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X604,'type.integer.int'),X605),s(X604,happ(s(fun(X606,X604),X607),s(X606,X608)))))))) )),
  inference(superposition,[],[f333,f365])).
fof(f9085,plain,(
  ( ! [X596,X592,X594,X597,X593,X595] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X597))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X594,'type.integer.int'),'_dst_x0x1_2'(s(fun(X592,'type.integer.int'),X593),s(fun(X594,X592),X595))),s(X594,X596))))) | s('type.integer.int',X597) = s('type.integer.int',happ(s(fun(X592,'type.integer.int'),X593),s(X592,happ(s(fun(X594,X592),X595),s(X594,X596)))))) )),
  inference(superposition,[],[f332,f365])).
fof(f9078,plain,(
  ( ! [X546,X549,X551,X545,X547,X552,X548,X550] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X551),s('type.list.list'('type.bvlSem.v'),X552))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X547,'type.num.num'),'_dst_x0x1_2'(s(fun(X545,'type.num.num'),X546),s(fun(X547,X545),X548))),s(X547,X549))),s('type.list.list'('type.bvlSem.v'),X550))) | s('type.num.num',X551) = s('type.num.num',happ(s(fun(X545,'type.num.num'),X546),s(X545,happ(s(fun(X547,X545),X548),s(X547,X549)))))) )),
  inference(superposition,[],[f390,f365])).
fof(f9076,plain,(
  ( ! [X536,X538,X534,X537,X533,X535] : (s('type.num.num',X538) != s('type.num.num',happ(s(fun(X535,'type.num.num'),'_dst_x0x1_2'(s(fun(X533,'type.num.num'),X534),s(fun(X535,X533),X536))),s(X535,X537))) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X538))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X533,'type.num.num'),X534),s(X533,happ(s(fun(X535,X533),X536),s(X535,X537)))))))) )),
  inference(superposition,[],[f331,f365])).
fof(f9074,plain,(
  ( ! [X524,X526,X522,X525,X521,X523] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X526))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X523,'type.num.num'),'_dst_x0x1_2'(s(fun(X521,'type.num.num'),X522),s(fun(X523,X521),X524))),s(X523,X525))))) | s('type.num.num',X526) = s('type.num.num',happ(s(fun(X521,'type.num.num'),X522),s(X521,happ(s(fun(X523,X521),X524),s(X523,X525)))))) )),
  inference(superposition,[],[f330,f365])).
fof(f9065,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X447,X446),X453),s(fun(X446,X447),happ(s(fun(X450,fun(X446,X447)),'_dst_x0x1_2'(s(fun(X448,fun(X446,X447)),X449),s(fun(X450,X448),X451))),s(X450,X452)))))) | s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) )),
  inference(superposition,[],[f354,f365])).
fof(f9064,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X437,X438),happ(s(fun(X441,fun(X437,X438)),'_dst_x0x1_2'(s(fun(X439,fun(X437,X438)),X440),s(fun(X441,X439),X442))),s(X441,X443))),s(fun(X438,X437),X444)))) | s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) )),
  inference(superposition,[],[f354,f365])).
fof(f9063,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X429,X428),X435),s(fun(X428,X429),happ(s(fun(X432,fun(X428,X429)),'_dst_x0x1_2'(s(fun(X430,fun(X428,X429)),X431),s(fun(X432,X430),X433))),s(X432,X434)))))) | s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) )),
  inference(superposition,[],[f353,f365])).
fof(f9062,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X419,X420),happ(s(fun(X423,fun(X419,X420)),'_dst_x0x1_2'(s(fun(X421,fun(X419,X420)),X422),s(fun(X423,X421),X424))),s(X423,X425))),s(fun(X420,X419),X426)))) | s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) )),
  inference(superposition,[],[f353,f365])).
fof(f9060,plain,(
  ( ! [X408,X405,X407,X401,X403,X409,X404,X406,X402] : (s(fun(X401,X402),X408) != s(fun(X401,X402),happ(s(fun(X405,fun(X401,X402)),'_dst_x0x1_2'(s(fun(X403,fun(X401,X402)),X404),s(fun(X405,X403),X406))),s(X405,X407))) | s(X402,happ(s(fun(X401,X402),X408),s(X401,X409))) = s(X402,happ(s(fun(X401,X402),happ(s(fun(X403,fun(X401,X402)),X404),s(X403,happ(s(fun(X405,X403),X406),s(X405,X407))))),s(X401,X409)))) )),
  inference(superposition,[],[f351,f365])).
fof(f9059,plain,(
  ( ! [X397,X399,X395,X396,X398,X400] : (p(s(bool,happ(s(fun('type.list.list'(X395),bool),happ(s(fun(X398,fun('type.list.list'(X395),bool)),'_dst_x0x1_2'(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(fun(X398,X396),X399))),s(X398,X400))),s('type.list.list'(X395),sK2(happ(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400)))),X395))))) | ~p(s(bool,happ(s(fun('type.list.list'(X395),bool),happ(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400))))),s('type.list.list'(X395),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f365])).
fof(f9058,plain,(
  ( ! [X389,X391,X393,X388,X390,X392,X394] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X388,bool),happ(s(fun(X391,fun(X388,bool)),'_dst_x0x1_2'(s(fun(X389,fun(X388,bool)),X390),s(fun(X391,X389),X392))),s(X391,X393))))) | s(X388,X394) = s(X388,sK4(X394,happ(s(fun(X389,fun(X388,bool)),X390),s(X389,happ(s(fun(X391,X389),X392),s(X391,X393)))),X388))) )),
  inference(superposition,[],[f866,f365])).
fof(f9057,plain,(
  ( ! [X383,X385,X387,X382,X384,X386,X381] : (p(s(bool,'|dollar|exists_1'(s(fun(X381,bool),happ(s(fun(X384,fun(X381,bool)),'_dst_x0x1_2'(s(fun(X382,fun(X381,bool)),X383),s(fun(X384,X382),X385))),s(X384,X386)))))) | s(X381,X387) = s(X381,sK4(X387,happ(s(fun(X382,fun(X381,bool)),X383),s(X382,happ(s(fun(X384,X382),X385),s(X384,X386)))),X381))) )),
  inference(superposition,[],[f314,f365])).
fof(f9056,plain,(
  ( ! [X377,X379,X374,X380,X376,X378,X375] : (~p(s(bool,happ(s(fun(X374,bool),happ(s(fun(X377,fun(X374,bool)),'_dst_x0x1_2'(s(fun(X375,fun(X374,bool)),X376),s(fun(X377,X375),X378))),s(X377,X379))),s(X374,X380)))) | p(s(bool,happ(s(fun(X374,bool),happ(s(fun(X375,fun(X374,bool)),X376),s(X375,happ(s(fun(X377,X375),X378),s(X377,X379))))),s(X374,'const.min.@_1'(s(fun(X374,bool),happ(s(fun(X375,fun(X374,bool)),X376),s(X375,happ(s(fun(X377,X375),X378),s(X377,X379))))))))))) )),
  inference(superposition,[],[f313,f365])).
fof(f9055,plain,(
  ( ! [X372,X368,X370,X367,X373,X369,X371] : (~p(s(bool,happ(s(fun(X367,bool),happ(s(fun(X370,fun(X367,bool)),'_dst_x0x1_2'(s(fun(X368,fun(X367,bool)),X369),s(fun(X370,X368),X371))),s(X370,X372))),s(X367,X373)))) | p(s(bool,'|dollar|exists_1'(s(fun(X367,bool),happ(s(fun(X368,fun(X367,bool)),X369),s(X368,happ(s(fun(X370,X368),X371),s(X370,X372))))))))) )),
  inference(superposition,[],[f4817,f365])).
fof(f9051,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X341),'const.sum.INL_1'(s(bool,happ(s(fun(X338,bool),'_dst_x0x1_2'(s(fun(X336,bool),X337),s(fun(X338,X336),X339))),s(X338,X340))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X336,bool),X337),s(X336,happ(s(fun(X338,X336),X339),s(X338,X340)))))) )),
  inference(superposition,[],[f7806,f365])).
fof(f9050,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X335,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X332,bool),'_dst_x0x1_2'(s(fun(X330,bool),X331),s(fun(X332,X330),X333))),s(X332,X334))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X330,bool),X331),s(X330,happ(s(fun(X332,X330),X333),s(X332,X334)))))) )),
  inference(superposition,[],[f7662,f365])).
fof(f9047,plain,(
  ( ! [X313,X315,X316,X312,X314,X311,X317] : (~p(s(bool,happ(s(fun(X313,bool),'_dst_x0x1_2'(s(fun(X311,bool),X312),s(fun(X313,X311),X314))),s(X313,X315)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X316),bool),sK16(X317,happ(s(fun(X311,bool),X312),s(X311,happ(s(fun(X313,X311),X314),s(X313,X315)))),X316,bool))))) )),
  inference(superposition,[],[f4999,f365])).
fof(f9046,plain,(
  ( ! [X308,X310,X306,X309,X305,X307] : (s(bool,'T_0') != s(bool,happ(s(fun(X307,bool),'_dst_x0x1_2'(s(fun(X305,bool),X306),s(fun(X307,X305),X308))),s(X307,X309))) | s(bool,X310) = s(bool,happ(s(fun(X305,bool),X306),s(X305,happ(s(fun(X307,X305),X308),s(X307,X309))))) | s(bool,'F_0') = s(bool,X310)) ) | $spl10),
  inference(superposition,[],[f3767,f365])).
fof(f9045,plain,(
  ( ! [X300,X302,X304,X301,X303] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X302,bool),'_dst_x0x1_2'(s(fun(X300,bool),X301),s(fun(X302,X300),X303))),s(X302,X304))))) | s(bool,'T_0') = s(bool,happ(s(fun(X300,bool),X301),s(X300,happ(s(fun(X302,X300),X303),s(X302,X304)))))) ) | $spl10),
  inference(superposition,[],[f3047,f365])).
fof(f9034,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (p(s(bool,happ(s(fun(X239,bool),'_dst_x0x1_2'(s(fun(X237,bool),X238),s(fun(X239,X237),X240))),s(X239,X241)))) | s(bool,X242) != s(bool,happ(s(fun(X237,bool),X238),s(X237,happ(s(fun(X239,X237),X240),s(X239,X241))))) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f2221,f365])).
fof(f9033,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X233,bool),'_dst_x0x1_2'(s(fun(X231,bool),X232),s(fun(X233,X231),X234))),s(X233,X235)))) | s(bool,X236) = s(bool,happ(s(fun(X231,bool),X232),s(X231,happ(s(fun(X233,X231),X234),s(X233,X235))))) | s(bool,'F_0') = s(bool,X236)) )),
  inference(superposition,[],[f2219,f365])).
fof(f9032,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X230),'const.sum.INL_1'(s(bool,happ(s(fun(X227,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X227,X225),X228))),s(X227,X229))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X227,X225),X228),s(X227,X229)))))) )),
  inference(superposition,[],[f2217,f365])).
fof(f9031,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X224,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X221,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X221,X219),X222))),s(X221,X223))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X221,X219),X222),s(X221,X223)))))) )),
  inference(superposition,[],[f2213,f365])).
fof(f9028,plain,(
  ( ! [X206,X204,X202,X207,X205,X203] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X207),'const.sum.INL_1'(s(bool,happ(s(fun(X204,bool),'_dst_x0x1_2'(s(fun(X202,bool),X203),s(fun(X204,X202),X205))),s(X204,X206))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X202,bool),X203),s(X202,happ(s(fun(X204,X202),X205),s(X204,X206)))))) )),
  inference(superposition,[],[f529,f365])).
fof(f9027,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X198,bool),'_dst_x0x1_2'(s(fun(X196,bool),X197),s(fun(X198,X196),X199))),s(X198,X200))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),X197),s(X196,happ(s(fun(X198,X196),X199),s(X198,X200)))))) )),
  inference(superposition,[],[f502,f365])).
fof(f9026,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (p(s(bool,happ(s(fun(X192,bool),'_dst_x0x1_2'(s(fun(X190,bool),X191),s(fun(X192,X190),X193))),s(X192,X194)))) | s(bool,X195) = s(bool,happ(s(fun(X190,bool),X191),s(X190,happ(s(fun(X192,X190),X193),s(X192,X194))))) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f483,f365])).
fof(f9025,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X187,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X187,X185),X188))),s(X187,X189))))) | s(bool,'F_0') = s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X187,X185),X188),s(X187,X189)))))) )),
  inference(superposition,[],[f461,f365])).
fof(f9024,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun(X181,bool),'_dst_x0x1_2'(s(fun(X179,bool),X180),s(fun(X181,X179),X182))),s(X181,X183)))) | s(bool,X184) != s(bool,happ(s(fun(X179,bool),X180),s(X179,happ(s(fun(X181,X179),X182),s(X181,X183))))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f436,f365])).
fof(f9019,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,happ(s(fun(X151,bool),'_dst_x0x1_2'(s(fun(X149,bool),X150),s(fun(X151,X149),X152))),s(X151,X153)))) | s(bool,X154) != s(bool,happ(s(fun(X149,bool),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))) | p(s(bool,X154))) )),
  inference(superposition,[],[f406,f365])).
fof(f9018,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X145,bool),'_dst_x0x1_2'(s(fun(X143,bool),X144),s(fun(X145,X143),X146))),s(X145,X147)))) | ~p(s(bool,X148)) | s(bool,X148) = s(bool,happ(s(fun(X143,bool),X144),s(X143,happ(s(fun(X145,X143),X146),s(X145,X147)))))) )),
  inference(superposition,[],[f329,f365])).
fof(f9016,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (p(s(bool,happ(s(fun(X133,bool),'_dst_x0x1_2'(s(fun(X131,bool),X132),s(fun(X133,X131),X134))),s(X133,X135)))) | p(s(bool,X136)) | s(bool,X136) = s(bool,happ(s(fun(X131,bool),X132),s(X131,happ(s(fun(X133,X131),X134),s(X133,X135)))))) )),
  inference(superposition,[],[f326,f365])).
fof(f9015,plain,(
  ( ! [X127,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X128,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X128,X126),X129))),s(X128,X130)))) | s(bool,'F_0') = s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130)))))) )),
  inference(superposition,[],[f272,f365])).
fof(f9014,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'F_0') != s(bool,happ(s(fun(X123,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X123,X121),X124))),s(X123,X125))) | ~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) )),
  inference(superposition,[],[f271,f365])).
fof(f9013,plain,(
  ( ! [X118,X116,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X118,X116),X119))),s(X118,X120)))) | s(bool,'T_0') = s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X118,X116),X119),s(X118,X120)))))) )),
  inference(superposition,[],[f268,f365])).
fof(f9012,plain,(
  ( ! [X111,X114,X112,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X113,X111),X114))),s(X113,X115)))) | s(bool,'T_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X113,X111),X114),s(X113,X115)))))) )),
  inference(superposition,[],[f267,f365])).
fof(f9010,plain,(
  ( ! [X103,X101,X105,X102,X104] : (~p(s(bool,happ(s(fun(X103,bool),'_dst_x0x1_2'(s(fun(X101,bool),X102),s(fun(X103,X101),X104))),s(X103,X105)))) | p(s(bool,happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102))))))) )),
  inference(superposition,[],[f313,f365])).
fof(f9009,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,happ(s(fun(X98,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X98,X96),X99))),s(X98,X100)))) | p(s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97))))) )),
  inference(superposition,[],[f4817,f365])).
fof(f8894,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s(fun(X104,X105),'_dst_x0x1_2'(s(fun(X106,X105),happ(s(fun(X108,fun(X106,X105)),X109),s(X108,X110))),s(fun(X104,X106),X107))) = s(fun(X104,X105),happ(s(fun(X108,fun(X104,X105)),'_dst_x0x1_2'(s(fun(fun(X106,X105),fun(X104,X105)),'_dstsel_0_x01_1'(s(fun(X104,X106),X107))),s(fun(X108,fun(X106,X105)),X109))),s(X108,X110)))) )),
  inference(superposition,[],[f365,f364])).
fof(f8890,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X85) = s(bool,happ(s(fun(X82,bool),'_dst_x0x1_2'(s(fun(X80,bool),X81),s(fun(X82,X80),X83))),s(X82,X84))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,happ(s(fun(X82,X80),X83),s(X82,X84)))))) ) | $spl10),
  inference(superposition,[],[f365,f2937])).
fof(f8889,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(X76,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X76,X74),X77))),s(X76,X78))) | s(bool,X79) = s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X76,X74),X77),s(X76,X78))))) | s(bool,'F_0') = s(bool,X79)) ) | $spl10),
  inference(superposition,[],[f365,f2937])).
fof(f8883,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,X45) = s(bool,happ(s(fun(X42,bool),'_dst_x0x1_2'(s(fun(X40,bool),X41),s(fun(X42,X40),X43))),s(X42,X44))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,happ(s(fun(X42,X40),X43),s(X42,X44)))))) )),
  inference(superposition,[],[f365,f434])).
fof(f8882,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),'_dst_x0x1_2'(s(fun(X34,bool),X35),s(fun(X36,X34),X37))),s(X36,X38))) | s(bool,X39) = s(bool,happ(s(fun(X34,bool),X35),s(X34,happ(s(fun(X36,X34),X37),s(X36,X38))))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f365,f434])).
fof(f8881,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X29,bool),X30),s(fun(X31,X29),X32))),s(X31,X33))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f365,f266])).
fof(f8880,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f365,f338])).
fof(f8879,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f365,f365])).
fof(f8878,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X113] : (s(X113,happ(s(fun(fun(X108,X109),X113),X114),s(fun(X108,X109),'_dst_x0x1_2'(s(fun(X110,X109),X112),s(fun(X108,X110),X111))))) = s(X113,happ(s(fun(fun(X110,X109),X113),'_dst_x0x1_2'(s(fun(fun(X108,X109),X113),X114),s(fun(fun(X110,X109),fun(X108,X109)),'_dstsel_0_x01_1'(s(fun(X108,X110),X111))))),s(fun(X110,X109),X112)))) )),
  inference(superposition,[],[f365,f364])).
fof(f8877,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK16(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f365,f356])).
fof(f8876,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f365,f365])).
fof(f8875,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f365,f365])).
fof(f8818,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f365,f351])).
fof(f9246,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),X5))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),X5)))) )),
  inference(forward_demodulation,[],[f8835,f337])).
fof(f8835,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),X5))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.bvlSem.v'),X6),'const.sum.INL_1'(s('type.list.list'('type.bvlSem.v'),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f337,f365,f392])).
fof(f9245,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X5),X6))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X5),X6)))))) )),
  inference(forward_demodulation,[],[f8836,f338])).
fof(f8836,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X5),X6))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.bvlSem.v'),X5),'type.list.list'('type.bvlSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X5),X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f365,f392])).
fof(f9244,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X5),s('type.num.num',X6))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X5),s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f8837,f320])).
fof(f8837,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X5),s('type.num.num',X6))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X7),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(unit_resulting_resolution,[],[f320,f365,f392])).
fof(f8839,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f365,f331])).
fof(f8840,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X0,'type.integer.int'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X2,'type.integer.int'),'_dst_x0x1_2'(s(fun(X0,'type.integer.int'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f365,f333])).
fof(f8841,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X3,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f261,f365,f392])).
fof(f8856,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X0,X1),X3))),s(X0,X4))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X5,'type.list.list'('type.bvlSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X7,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.bvlSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f365,f365,f392])).
fof(f8857,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X0,'type.list.list'('type.bvlSem.v')),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X5))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X2,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X0,'type.list.list'('type.bvlSem.v')),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f261,f365,f392])).
fof(f8860,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X3,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f813,f365,f392])).
fof(f9218,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X3,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(forward_demodulation,[],[f8869,f337])).
fof(f8869,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X6),'const.sum.INL_1'(s('type.num.num',X0))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X3,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.bvlSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f337,f365,f392])).
fof(f9217,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X4,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f8870,f338])).
fof(f8870,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X0),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X4,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f365,f392])).
fof(f9216,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X4,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f8871,f320])).
fof(f8871,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X7),s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X4,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.bvlSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f320,f365,f392])).
fof(f8872,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X5,'type.list.list'('type.bvlSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.bvlSem.v'),happ(s(fun(X7,'type.list.list'('type.bvlSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.bvlSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f365,f365,f392])).
fof(f8760,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f8725,f254])).
fof(f8725,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(superposition,[],[f4999,f7068])).
fof(f8758,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f8702,f254])).
fof(f8702,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f436,f7068])).
fof(f8753,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f8684,f421])).
fof(f8684,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f7068,f2937])).
fof(f8667,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8601,f256])).
fof(f8601,plain,(
  ( ! [X43,X42] : (p(s(bool,'F_0')) | s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f2221,f3047])).
fof(f8563,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,sK4(X75,X76,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76)))) ) | $spl10),
  inference(superposition,[],[f3047,f866])).
fof(f8559,plain,(
  ( ! [X66,X64,X65] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X64),s(fun('type.num.num',bool),X65),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))))) ) | $spl10),
  inference(superposition,[],[f3047,f320])).
fof(f8556,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) ) | $spl10),
  inference(superposition,[],[f3047,f338])).
fof(f8508,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),sK2(happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34)),X31)))))) )),
  inference(forward_demodulation,[],[f8488,f364])).
fof(f8488,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),sK2(happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34)),X31))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34))),s('type.list.list'(X31),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f364])).
fof(f8487,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0x1_2'(s(fun(X27,bool),X29),s(fun(X26,X27),X28))))) | s(X26,X30) = s(X26,sK4(X30,happ(s(fun(fun(X27,bool),fun(X26,bool)),'_dstsel_0_x01_1'(s(fun(X26,X27),X28))),s(fun(X27,bool),X29)),X26))) )),
  inference(superposition,[],[f866,f364])).
fof(f8486,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0x1_2'(s(fun(X22,bool),X24),s(fun(X21,X22),X23)))))) | s(X21,X25) = s(X21,sK4(X25,happ(s(fun(fun(X22,bool),fun(X21,bool)),'_dstsel_0_x01_1'(s(fun(X21,X22),X23))),s(fun(X22,bool),X24)),X21))) )),
  inference(superposition,[],[f314,f364])).
fof(f8298,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f7806,f866])).
fof(f8297,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f7806,f337])).
fof(f8296,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f7806,f341])).
fof(f8295,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f7806,f340])).
fof(f8294,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f7806,f320])).
fof(f8293,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7806,f312])).
fof(f8292,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f7806,f336])).
fof(f8291,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f7806,f338])).
fof(f8290,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f7806,f356])).
fof(f8328,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8270,f421])).
fof(f8270,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f7806,f2937])).
fof(f8156,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X31),s('type.list.list'('type.bvlSem.v'),X32))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.bvlSem.v'),X29))))),s('type.list.list'('type.bvlSem.v'),X30))) | s('type.num.num',X31) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X29)))))))) )),
  inference(superposition,[],[f390,f325])).
fof(f8152,plain,(
  ( ! [X17,X18,X16] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X18))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.bvlSem.v'),X17))))))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X17)))))))) )),
  inference(superposition,[],[f330,f325])).
fof(f8145,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X3)))))))))) )),
  inference(superposition,[],[f325,f263])).
fof(f8246,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0))))))))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X0)))))) )),
  inference(forward_demodulation,[],[f8144,f263])).
fof(f8144,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))))) )),
  inference(superposition,[],[f325,f263])).
fof(f8115,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.bvlSem.v'),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X2)))))))) )),
  inference(superposition,[],[f325,f261])).
fof(f8178,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),X2))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))))),s('type.list.list'('type.bvlSem.v'),X2)))) )),
  inference(forward_demodulation,[],[f8177,f263])).
fof(f8177,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),X2))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),X2)))) )),
  inference(forward_demodulation,[],[f8111,f337])).
fof(f8111,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),X2))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.bvlSem.v'),X3),'const.sum.INL_1'(s('type.list.list'('type.bvlSem.v'),X2)))))))) )),
  inference(unit_resulting_resolution,[],[f337,f325,f392])).
fof(f8176,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3)))))) )),
  inference(forward_demodulation,[],[f8175,f263])).
fof(f8175,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3)))))) )),
  inference(forward_demodulation,[],[f8112,f338])).
fof(f8112,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),'type.list.list'('type.bvlSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X2),X3)))))) )),
  inference(unit_resulting_resolution,[],[f338,f325,f392])).
fof(f8174,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num',X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f8173,f263])).
fof(f8173,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num',X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f8113,f320])).
fof(f8113,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num',X3))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1))))))),s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X4),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) )),
  inference(unit_resulting_resolution,[],[f320,f325,f392])).
fof(f8172,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1)))))))))))) )),
  inference(forward_demodulation,[],[f8114,f263])).
fof(f8114,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f325,f331])).
fof(f7924,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f7662,f866])).
fof(f7923,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f7662,f337])).
fof(f7922,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f7662,f341])).
fof(f7921,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f7662,f340])).
fof(f7920,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f7662,f320])).
fof(f7919,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7662,f312])).
fof(f7918,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f7662,f336])).
fof(f7917,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f7662,f338])).
fof(f7916,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f7662,f356])).
fof(f7953,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f7896,f421])).
fof(f7896,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f7662,f2937])).
fof(f7814,plain,(
  ( ! [X61,X62,X63] : (s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f7757,f256])).
fof(f7757,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'F_0')) | s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f2221,f2217])).
fof(f7806,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7736,f256])).
fof(f7736,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f267,f2217])).
fof(f7722,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f2217,f866])).
fof(f7720,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f2217,f341])).
fof(f7719,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f2217,f340])).
fof(f7718,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f2217,f320])).
fof(f7717,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2217,f312])).
fof(f7715,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f2217,f338])).
fof(f7714,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2217,f356])).
fof(f7670,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f7613,f256])).
fof(f7613,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f2221,f2213])).
fof(f7662,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7593,f256])).
fof(f7593,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f267,f2213])).
fof(f7579,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f2213,f866])).
fof(f7577,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X108),s(bool,X109)))))) )),
  inference(superposition,[],[f2213,f341])).
fof(f7576,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X106,bool),'const.sum.INR_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f2213,f340])).
fof(f7575,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f2213,f320])).
fof(f7574,plain,(
  ( ! [X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2213,f312])).
fof(f7572,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f2213,f338])).
fof(f7571,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X86),bool),sK16(X87,X88,X86,bool)),s('type.list.list'(X86),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2213,f356])).
fof(f7052,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK16(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4999,f866])).
fof(f7051,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK16(X145,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X142),'const.sum.INL_1'(s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4999,f337])).
fof(f7050,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,X139)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK16(X141,'const.pair.SND_1'(s('type.pair.prod'(X137,bool),'const.pair.,_2'(s(X137,X138),s(bool,X139)))),X140,bool))))) )),
  inference(superposition,[],[f4999,f341])).
fof(f7049,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK16(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4999,f340])).
fof(f7048,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,happ(s(fun('type.num.num',bool),X128),s('type.num.num',X129)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK16(X131,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129)))),X130,bool))))) )),
  inference(superposition,[],[f4999,f320])).
fof(f7047,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,X123)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X125),bool),sK16(X126,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.|48|_0')),X125,bool))))) )),
  inference(superposition,[],[f4999,f312])).
fof(f7046,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK16(X122,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X120)))),X121,bool))))) )),
  inference(superposition,[],[f4999,f336])).
fof(f7045,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK16(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4999,f338])).
fof(f7044,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK16(X114,happ(s(fun('type.list.list'(X110),bool),sK16(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4999,f356])).
fof(f7043,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK16(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4999,f305])).
fof(f7076,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f7040,f254])).
fof(f7040,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4999,f529])).
fof(f7073,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f7038,f254])).
fof(f7038,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f4999,f5028])).
fof(f7069,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f7035,f254])).
fof(f7035,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4999,f866])).
fof(f7068,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f7034,f254])).
fof(f7034,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4999,f2434])).
fof(f7065,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(subsumption_resolution,[],[f7032,f254])).
fof(f7032,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(superposition,[],[f4999,f429])).
fof(f7064,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f7031,f254])).
fof(f7031,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f4999,f502])).
fof(f7063,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f7030,f254])).
fof(f7030,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4999,f461])).
fof(f7053,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7014,f254])).
fof(f7014,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4999,f434])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : ! [X29,X20] : (! [X28] : (s(X9,X20) = s(X9,X28) => p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X29)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.JRH_INDUCT_UTIL')).
fof(f4820,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4768,f254])).
fof(f4768,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f329,f866])).
fof(f4821,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4769,f254])).
fof(f4769,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f406,f866])).
fof(f859,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f314,f305])).
fof(f1093,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK16(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f314,f356])).
fof(f1690,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f314,f338])).
fof(f860,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X6) = s(X3,sK4(X6,'const.sum.OUTR_1'(s('type.sum.sum'(X4,fun(X3,bool)),'const.sum.INR_1'(s(fun(X3,bool),X5)))),X3))) )),
  inference(superposition,[],[f314,f336])).
fof(f861,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8)))) | s(X7,X10) = s(X7,sK4(X10,'const.ind_type.FCONS_3'(s(fun(X7,bool),X8),s(fun('type.num.num',fun(X7,bool)),X9),s('type.num.num','const.num.|48|_0')),X7))) )),
  inference(superposition,[],[f314,f312])).
fof(f2619,plain,(
  ( ! [X134,X132,X136,X135,X133] : (p(s(bool,'|dollar|exists_1'(s(fun(X132,bool),happ(s(fun('type.num.num',fun(X132,bool)),X134),s('type.num.num',X135)))))) | s(X132,X136) = s(X132,sK4(X136,'const.ind_type.FCONS_3'(s(fun(X132,bool),X133),s(fun('type.num.num',fun(X132,bool)),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))),X132))) )),
  inference(superposition,[],[f314,f320])).
fof(f862,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f314,f340])).
fof(f863,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X19)))) | s(X16,X20) = s(X16,sK4(X20,'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))) )),
  inference(superposition,[],[f314,f341])).
fof(f864,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f314,f337])).
fof(f4723,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f314,f866])).
fof(f4687,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f267,f866])).
fof(f4688,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f268,f866])).
fof(f4689,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f271,f866])).
fof(f4690,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f272,f866])).
fof(f4691,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f326,f866])).
fof(f4693,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f329,f866])).
fof(f4694,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f406,f866])).
fof(f4699,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f436,f866])).
fof(f4701,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f483,f866])).
fof(f4706,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f2219,f866])).
fof(f4707,plain,(
  ( ! [X111,X109,X110] : (p(s(bool,X109)) | s(bool,X111) != s(bool,sK4(X109,X110,bool)) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f2221,f866])).
fof(f4833,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4817,f866])).
fof(f4721,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f313,f866])).
fof(f4724,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f408,f866])).
fof(f5094,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f866])).
fof(f5084,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f354,f866])).
fof(f4945,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f353,f866])).
fof(f4935,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f353,f866])).
fof(f4725,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (s(fun(X166,X167),X168) != s(fun(X166,X167),X170) | s(X167,happ(s(fun(X166,X167),X170),s(X166,X171))) = s(X167,happ(s(fun(X166,X167),sK4(X168,X169,fun(X166,X167))),s(X166,X171))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X166,X167),bool),X169)))) )),
  inference(superposition,[],[f351,f866])).
fof(f4734,plain,(
  ( ! [X198,X200,X199] : (s('type.num.num',X198) != s('type.num.num',X200) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X200))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK4(X198,X199,'type.num.num')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X199)))) )),
  inference(superposition,[],[f331,f866])).
fof(f4745,plain,(
  ( ! [X237,X238,X236] : (s('type.integer.int',X236) != s('type.integer.int',X238) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X238))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK4(X236,X237,'type.integer.int')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X237)))) )),
  inference(superposition,[],[f333,f866])).
fof(f5757,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.list.list'('type.bvlSem.v'),X48) != s('type.list.list'('type.bvlSem.v'),X50) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X52),s('type.list.list'('type.bvlSem.v'),X50))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X51),s('type.list.list'('type.bvlSem.v'),sK4(X48,X49,'type.list.list'('type.bvlSem.v'))))) | s('type.num.num',X51) != s('type.num.num',X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.bvlSem.v'),bool),X49)))) )),
  inference(superposition,[],[f392,f866])).
fof(f2583,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f320,f434])).
fof(f2585,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f320,f434])).
fof(f2611,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f461,f320])).
fof(f3052,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2933,f2434])).
fof(f2933,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))) ) | $spl10),
  inference(superposition,[],[f2517,f320])).
fof(f3511,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X78) = s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) ) | $spl10),
  inference(superposition,[],[f2937,f320])).
fof(f3582,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X63),s('type.num.num',X64))) | s(bool,X65) = s(bool,'const.ind_type.FCONS_3'(s(bool,X62),s(fun('type.num.num',bool),X63),s('type.num.num','const.num.SUC_1'(s('type.num.num',X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2937,f320])).
fof(f4279,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X78),s('type.num.num',X79))) | s(bool,X80) = s(bool,'const.ind_type.FCONS_3'(s(bool,X77),s(fun('type.num.num',bool),X78),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3767,f320])).
fof(f4396,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f502,f320])).
fof(f4514,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X84),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X82),s('type.num.num',X83))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X81),s(fun('type.num.num',bool),X82),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83)))))) )),
  inference(superposition,[],[f529,f320])).
fof(f1659,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f338,f434])).
fof(f1661,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f338,f434])).
fof(f1663,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f338,f266])).
fof(f2167,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f461,f338])).
fof(f2410,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29))))) )),
  inference(forward_demodulation,[],[f2388,f338])).
fof(f2388,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f313,f338])).
fof(f3049,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2930,f2434])).
fof(f2930,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2517,f338])).
fof(f3508,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2937,f338])).
fof(f3579,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2937,f338])).
fof(f4276,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3767,f338])).
fof(f4393,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f502,f338])).
fof(f4511,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f529,f338])).
fof(f4899,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4817,f338])).
fof(f6605,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) )),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f314,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f859,f860,f861,f862,f863,f864,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1093,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1659,f1661,f1663,f1672,f1716,f1690,f1693,f1700,f1832,f1849,f2062,f2212,f2167,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2410,f2583,f2585,f2611,f2663,f2619,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4393,f4394,f4395,f4396,f4397,f4398,f4399,f4478,f4510,f4511,f4512,f4513,f4514,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4687,f4688,f4689,f4690,f4691,f4693,f4694,f4699,f4700,f4701,f4702,f4703,f4706,f4707,f4721,f4722,f4723,f4724,f4725,f4732,f4734,f4736,f4743,f4745,f4753,f4820,f4821,f4822,f4833,f4899,f4935,f4945,f4985,f5028,f5022,f5084,f5094,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5757,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235])).
fof(f5235,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5223,f254])).
fof(f5223,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4899,f434])).
fof(f3852,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'F_0') != s(bool,f(s(X148,sK12(X149,X150,bool,X148)))) | s(fun(X148,bool),X149) = s(fun(X148,bool),X150) | s(bool,X151) = s(bool,g(s(X148,sK12(X149,X150,bool,X148)))) | s(bool,'F_0') = s(bool,X151)) ) | $spl10),
  inference(superposition,[],[f350,f2937])).
fof(f3682,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X164) != s(bool,f(s(X161,sK12(X162,X163,bool,X161)))) | s(fun(X161,bool),X162) = s(fun(X161,bool),X163) | s(bool,'F_0') = s(bool,X164) | s(bool,'F_0') = s(bool,g(s(X161,sK12(X162,X163,bool,X161))))) ) | $spl10),
  inference(superposition,[],[f350,f2937])).
fof(f3464,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK12(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK12(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f350,f434])).
fof(f3463,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK12(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK12(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f350,f434])).
fof(f3462,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f350,f266])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : ! [X29] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X29),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X20] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X29),s('type.list.list'(X9),X20)))) => ! [X30] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X29),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X30),s('type.list.list'(X9),X20)))))))) => ! [X31] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X29),s('type.list.list'(X9),X31)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.list.list_INDUCT')).
fof(f408,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f3128,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f3112,f305])).
fof(f3112,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f305])).
fof(f3129,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f3113,f356])).
fof(f3113,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f356])).
fof(f3131,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9)))))) )),
  inference(forward_demodulation,[],[f3115,f336])).
fof(f3115,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11))))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f336])).
fof(f3132,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f3116,f312])).
fof(f3116,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f312])).
fof(f3134,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f3118,f340])).
fof(f3118,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f340])).
fof(f3135,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23)))))) )),
  inference(forward_demodulation,[],[f3119,f341])).
fof(f3119,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f341])).
fof(f3136,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27)))))) )),
  inference(forward_demodulation,[],[f3120,f337])).
fof(f3120,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29))))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f337])).
fof(f3137,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f3121,f256])).
fof(f3121,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f408,f266])).
fof(f4212,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(subsumption_resolution,[],[f3830,f256])).
fof(f3830,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(superposition,[],[f408,f2937])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ! [X29,X28] : (p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28)))) => p(s(bool,happ(s(fun(X9,bool),X29),s(X9,'const.min.@_1'(s(fun(X9,bool),X29)))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.SELECT_AX')).
fof(f316,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X54,X29] : (! [X28] : (p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28)))) => p(s(bool,happ(s(fun(X9,bool),X54),s(X9,X28))))) => (? [X28] : p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28)))) => ? [X28] : p(s(bool,happ(s(fun(X9,bool),X54),s(X9,X28))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.MONO_EXISTS')).
fof(f317,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f345,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X27] : ! [X50,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X27),X10),s(fun(X27,X9),X50)))) => (! [X29] : (! [X28] : p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28)))) <=> ! [X28] : p(s(bool,happ(s(fun(X9,bool),X29),s(X9,happ(s(fun(X27,X9),X50),s(X27,X28))))))) & ! [X29] : (? [X28] : p(s(bool,happ(s(fun(X9,bool),X29),s(X9,X28)))) <=> ? [X28] : p(s(bool,happ(s(fun(X9,bool),X29),s(X9,happ(s(fun(X27,X9),X50),s(X27,X28))))))) & ! [X6,X52] : (s(X9,X6) = s(X9,happ(s(fun(X27,X9),X50),s(X27,X52))) <=> s(X27,X52) = s(X27,happ(s(fun(X9,X27),X10),s(X9,X6))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.ISO_USAGE')).
fof(f346,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2407,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2381,f254])).
fof(f2381,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f313,f434])).
fof(f3085,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3084,f254])).
fof(f3084,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2969,f2434])).
fof(f2969,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f313,f2517])).
fof(f4817,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4671,f267])).
fof(f4671,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f315,f866])).
fof(f4919,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4886,f254])).
fof(f4886,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4817,f434])).
fof(f344,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f347,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f3133,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15)))))) )),
  inference(forward_demodulation,[],[f3117,f320])).
fof(f3117,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15))))) | ~p(s(bool,happ(s(fun('type.list.list'(X15),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.list.list'(X15),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f320])).
fof(f3130,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f3114,f338])).
fof(f3114,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f338])).
fof(f2387,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.min.@_1'(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)))))))) )),
  inference(superposition,[],[f313,f356])).
fof(f348,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f236])).
fof(f349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f236])).
fof(f353,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0)))))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X27,X9] : ! [X10,X50] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X27),X10),s(fun(X27,X9),X50)))) <=> (! [X28] : s(X27,X28) = s(X27,happ(s(fun(X9,X27),X10),s(X9,happ(s(fun(X27,X9),X50),s(X27,X28))))) & ! [X35] : s(X9,X35) = s(X9,happ(s(fun(X27,X9),X50),s(X27,happ(s(fun(X9,X27),X10),s(X9,X35)))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.ISO')).
fof(f354,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f355,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f394,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X27,X47,X9,X48] : ! [X49,X50,X51,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X48),X10),s(fun(X48,X9),X51)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X27,X47),X50),s(fun(X47,X27),X49))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X27),fun(X48,X47)),'_dst_x00x11_2'(s(fun(X27,X47),X50),s(fun(fun(X9,X27),fun(X48,X27)),'_dstsel_0_x01_1'(s(fun(X48,X9),X51))))),s(fun(fun(X48,X47),fun(X9,X27)),'_dst_x00x11_2'(s(fun(X47,X27),X49),s(fun(fun(X48,X47),fun(X9,X47)),'_dstsel_0_x01_1'(s(fun(X9,X48),X10)))))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.ISO_FUN')).
fof(f4979,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4948,f254])).
fof(f4948,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f353,f434])).
fof(f5128,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5097,f254])).
fof(f5097,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f354,f434])).
fof(f4650,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f866,f305])).
fof(f4651,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f866,f356])).
fof(f4653,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))) | s(X12,X15) = s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))) )),
  inference(superposition,[],[f866,f336])).
fof(f4654,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),X17))) | s(X16,X19) = s(X16,sK4(X19,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0')),X16))) )),
  inference(superposition,[],[f866,f312])).
fof(f4656,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f866,f340])).
fof(f4657,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X33))) | s(X30,X34) = s(X30,sK4(X34,'const.pair.SND_1'(s('type.pair.prod'(X31,fun(X30,bool)),'const.pair.,_2'(s(X31,X32),s(fun(X30,bool),X33)))),X30))) )),
  inference(superposition,[],[f866,f341])).
fof(f4658,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X37))) | s(X35,X38) = s(X35,sK4(X38,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X35,bool),X36),'const.sum.INL_1'(s(fun(X35,bool),X37)))),X35))) )),
  inference(superposition,[],[f866,f337])).
fof(f4659,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f866,f866])).
fof(f4655,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),happ(s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num',X23))))) | s(X20,X24) = s(X20,sK4(X24,'const.ind_type.FCONS_3'(s(fun(X20,bool),X21),s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X20))) )),
  inference(superposition,[],[f866,f320])).
fof(f4652,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f866,f338])).
fof(f5004,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4995,f256])).
fof(f4995,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4898,f2937])).
fof(f4999,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4987,f256])).
fof(f4987,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4898,f266])).
fof(f4898,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4817,f356])).
fof(f5029,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4985,f4898])).
fof(f5768,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X1))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(equality_resolution,[],[f392])).
fof(f5838,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X0))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))))),s('type.list.list'('type.bvlSem.v'),X0)))) )),
  inference(forward_demodulation,[],[f5676,f337])).
fof(f5676,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X0))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))))),s('type.list.list'('type.bvlSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.bvlSem.v'),X2),'const.sum.INL_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))) )),
  inference(unit_resulting_resolution,[],[f261,f337,f392])).
fof(f5835,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),X1)))) )),
  inference(forward_demodulation,[],[f5679,f337])).
fof(f5679,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.bvlSem.v'),X3),'const.sum.INL_1'(s('type.list.list'('type.bvlSem.v'),X1)))))))) )),
  inference(unit_resulting_resolution,[],[f813,f337,f392])).
fof(f5810,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X0),X1))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X0),X1)))))) )),
  inference(forward_demodulation,[],[f5704,f338])).
fof(f5704,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X0),X1))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.bvlSem.v'),X0),'type.list.list'('type.bvlSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X0),X1)))))) )),
  inference(unit_resulting_resolution,[],[f261,f338,f392])).
fof(f5807,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X1),X2))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X1),X2)))))) )),
  inference(forward_demodulation,[],[f5707,f338])).
fof(f5707,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X1),X2))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.bvlSem.v'),X1),'type.list.list'('type.bvlSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X1),X2)))))) )),
  inference(unit_resulting_resolution,[],[f813,f338,f392])).
fof(f5782,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X0),s('type.num.num',X1))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X0),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f5732,f320])).
fof(f5732,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X0),s('type.num.num',X1))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X3),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f261,f320,f392])).
fof(f5779,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X1),s('type.num.num',X2))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f5735,f320])).
fof(f5735,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X1),s('type.num.num',X2))))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X4),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) )),
  inference(unit_resulting_resolution,[],[f813,f320,f392])).
fof(f5210,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5163,f254])).
fof(f5163,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f436,f5028])).
fof(f5074,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f5039,f256])).
fof(f5039,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4985,f2937])).
fof(f5022,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f356,f4985,f436])).
fof(f5028,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f421,f4985,f2219])).
fof(f4985,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f254,f4898])).
fof(f4822,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4772,f254])).
fof(f4772,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f436,f866])).
fof(f4753,plain,(
  ( ! [X263,X265,X267,X264,X266] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X265),s('type.list.list'('type.bvlSem.v'),X263))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X266),s('type.list.list'('type.bvlSem.v'),X267))) | s('type.list.list'('type.bvlSem.v'),X267) = s('type.list.list'('type.bvlSem.v'),sK4(X263,X264,'type.list.list'('type.bvlSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.bvlSem.v'),bool),X264)))) )),
  inference(superposition,[],[f391,f866])).
fof(f4743,plain,(
  ( ! [X231,X230,X232] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X230))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X232))) | s('type.integer.int',X232) = s('type.integer.int',sK4(X230,X231,'type.integer.int')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X231)))) )),
  inference(superposition,[],[f332,f866])).
fof(f4736,plain,(
  ( ! [X208,X206,X204,X207,X205] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X204),s('type.list.list'('type.bvlSem.v'),X206))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X207),s('type.list.list'('type.bvlSem.v'),X208))) | s('type.num.num',X207) = s('type.num.num',sK4(X204,X205,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X205)))) )),
  inference(superposition,[],[f390,f866])).
fof(f4732,plain,(
  ( ! [X194,X192,X193] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X192))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X194))) | s('type.num.num',X194) = s('type.num.num',sK4(X192,X193,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X193)))) )),
  inference(superposition,[],[f330,f866])).
fof(f4722,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f315,f866])).
fof(f4718,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3767,f866])).
fof(f4818,plain,(
  ( ! [X114,X115] : (s(bool,'T_0') != s(bool,X114) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) ) | $spl10),
  inference(forward_demodulation,[],[f4709,f2434])).
fof(f4709,plain,(
  ( ! [X114,X115] : (s(bool,X114) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) ) | $spl10),
  inference(superposition,[],[f2516,f866])).
fof(f4703,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f529,f866])).
fof(f4702,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f502,f866])).
fof(f4700,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f461,f866])).
fof(f4647,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f866,f2937])).
fof(f4646,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f866,f2937])).
fof(f4816,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X26))) | s(bool,'F_0') = s(bool,sK4(X25,X26,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4645,f2434])).
fof(f4645,plain,(
  ( ! [X26,X25] : (s(bool,X25) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X26))) | s(bool,'F_0') = s(bool,sK4(X25,X26,bool))) ) | $spl10),
  inference(superposition,[],[f866,f2517])).
fof(f4640,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f866,f434])).
fof(f4639,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f866,f434])).
fof(f4638,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f866,f266])).
fof(f4597,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4540,f254])).
fof(f4540,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f436,f529])).
fof(f4517,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,X94)))))) )),
  inference(superposition,[],[f529,f337])).
fof(f4516,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f529,f341])).
fof(f4515,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f529,f340])).
fof(f4513,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,X78))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f529,f312])).
fof(f4512,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X77),'const.sum.INL_1'(s(bool,X76))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76)))))) )),
  inference(superposition,[],[f529,f336])).
fof(f4510,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK16(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f529,f356])).
fof(f4478,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4422,f254])).
fof(f4422,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f436,f502])).
fof(f4399,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f502,f337])).
fof(f4398,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) )),
  inference(superposition,[],[f502,f341])).
fof(f4397,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X85,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f502,f340])).
fof(f4395,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f502,f312])).
fof(f4394,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,X73))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,X73)))))) )),
  inference(superposition,[],[f502,f336])).
fof(f4392,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X68,bool),'const.sum.INR_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK16(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f502,f356])).
fof(f4282,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3767,f337])).
fof(f4281,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') != s(bool,X87) | s(bool,X88) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87))))) | s(bool,'F_0') = s(bool,X88)) ) | $spl10),
  inference(superposition,[],[f3767,f341])).
fof(f4280,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3767,f340])).
fof(f4278,plain,(
  ( ! [X76,X74,X75] : (s(bool,'T_0') != s(bool,X74) | s(bool,X76) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f3767,f312])).
fof(f4277,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | s(bool,X73) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3767,f336])).
fof(f4275,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK16(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3767,f356])).
fof(f4274,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3767,f305])).
fof(f4241,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3767,f434])).
fof(f3863,plain,(
  ( ! [X187,X185,X188,X186] : (s(X187,f(s(bool,'F_0'))) != s(X187,g(s(bool,'F_0'))) | s(fun(bool,X187),X185) = s(fun(bool,X187),X186) | s(bool,X188) = s(bool,sK12(X185,X186,X187,bool)) | s(bool,'F_0') = s(bool,X188)) ) | $spl10),
  inference(superposition,[],[f350,f2937])).
fof(f3862,plain,(
  ( ! [X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,X184) = s(bool,sK4(X182,X183,bool)) | s(bool,'F_0') = s(bool,X184)) ) | $spl10),
  inference(superposition,[],[f315,f2937])).
fof(f4214,plain,(
  ( ! [X94,X95,X93,X96] : (s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(subsumption_resolution,[],[f3835,f256])).
fof(f3835,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,'F_0')) | s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(superposition,[],[f314,f2937])).
fof(f4207,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3802,f256])).
fof(f3802,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f326,f2937])).
fof(f3767,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f421,f2937])).
fof(f3766,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f271,f2937])).
fof(f3693,plain,(
  ( ! [X198,X200,X199,X201] : (s(X200,f(s(bool,X201))) != s(X200,g(s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,'F_0') = s(bool,X201) | s(bool,'F_0') = s(bool,sK12(X198,X199,X200,bool))) ) | $spl10),
  inference(superposition,[],[f350,f2937])).
fof(f3692,plain,(
  ( ! [X196,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X196),s(bool,X197)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X196)))) | s(bool,'F_0') = s(bool,X197) | s(bool,'F_0') = s(bool,sK4(X195,X196,bool))) ) | $spl10),
  inference(superposition,[],[f315,f2937])).
fof(f3585,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2937,f337])).
fof(f3584,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f2937,f341])).
fof(f3583,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2937,f340])).
fof(f3581,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X61)) ) | $spl10),
  inference(superposition,[],[f2937,f312])).
fof(f3580,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X57) | s(bool,X58) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X56,bool),'const.sum.INR_1'(s(bool,X57))))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2937,f336])).
fof(f3578,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK16(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2937,f356])).
fof(f3549,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2937,f2937])).
fof(f3514,plain,(
  ( ! [X88,X87,X89] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X87),'const.sum.INL_1'(s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2937,f337])).
fof(f3513,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,X85) = s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X85)))))) ) | $spl10),
  inference(superposition,[],[f2937,f341])).
fof(f3512,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2937,f340])).
fof(f3510,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2937,f312])).
fof(f3509,plain,(
  ( ! [X70,X71,X69] : (s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,X70)))))) ) | $spl10),
  inference(superposition,[],[f2937,f336])).
fof(f3507,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK16(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2937,f356])).
fof(f3506,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2937,f305])).
fof(f3456,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK12(X7,X8,X9,bool))) )),
  inference(superposition,[],[f350,f434])).
fof(f3455,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK12(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f350,f434])).
fof(f3454,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,X2,bool))) )),
  inference(superposition,[],[f350,f266])).
fof(f3099,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f3000,f2434])).
fof(f3000,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f315,f2517])).
fof(f3083,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3082,f254])).
fof(f3082,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2968,f2434])).
fof(f2968,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f315,f2517])).
fof(f2937,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2517,f2517])).
fof(f2812,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.bvlSem.v'),X44))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X46),s('type.list.list'('type.bvlSem.v'),X47))) | s('type.list.list'('type.bvlSem.v'),X47) = s('type.list.list'('type.bvlSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.bvlSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.bvlSem.v'),X44)))))) )),
  inference(superposition,[],[f391,f337])).
fof(f2811,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X40),s('type.list.list'('type.bvlSem.v'),X39))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X41),s('type.list.list'('type.bvlSem.v'),X42))) | s('type.list.list'('type.bvlSem.v'),X42) = s('type.list.list'('type.bvlSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.bvlSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.bvlSem.v'),X39)))))) )),
  inference(superposition,[],[f391,f341])).
fof(f2810,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X34),s('type.list.list'('type.bvlSem.v'),X32))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X35),s('type.list.list'('type.bvlSem.v'),X36))) | s('type.list.list'('type.bvlSem.v'),X36) = s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.bvlSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f391,f340])).
fof(f2809,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X29),s('type.list.list'('type.bvlSem.v'),X30))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X26),s('type.num.num',X27))))) | s('type.list.list'('type.bvlSem.v'),X30) = s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X25),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) )),
  inference(superposition,[],[f391,f320])).
fof(f2808,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X22),s('type.list.list'('type.bvlSem.v'),X20))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X23),s('type.list.list'('type.bvlSem.v'),X24))) | s('type.list.list'('type.bvlSem.v'),X24) = s('type.list.list'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.bvlSem.v'),X20),s(fun('type.num.num','type.list.list'('type.bvlSem.v')),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f391,f312])).
fof(f2807,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X17),s('type.list.list'('type.bvlSem.v'),X16))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X18),s('type.list.list'('type.bvlSem.v'),X19))) | s('type.list.list'('type.bvlSem.v'),X19) = s('type.list.list'('type.bvlSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X15,'type.list.list'('type.bvlSem.v')),'const.sum.INR_1'(s('type.list.list'('type.bvlSem.v'),X16)))))) )),
  inference(superposition,[],[f391,f336])).
fof(f2806,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X13),s('type.list.list'('type.bvlSem.v'),X14))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X12),s('type.list.list'('type.bvlSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X10),X11))))) | s('type.list.list'('type.bvlSem.v'),X14) = s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.bvlSem.v'),X10),'type.list.list'('type.bvlSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.bvlSem.v'),X10),X11)))) )),
  inference(superposition,[],[f391,f338])).
fof(f2805,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X7),s('type.list.list'('type.bvlSem.v'),X6))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X8),s('type.list.list'('type.bvlSem.v'),X9))) | s('type.list.list'('type.bvlSem.v'),X9) = s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.bvlSem.v')),sK16(X5,X6,X4,'type.list.list'('type.bvlSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f391,f356])).
fof(f2804,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.bvlSem.v'),X0))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | s('type.list.list'('type.bvlSem.v'),X3) = s('type.list.list'('type.bvlSem.v'),happ(s(fun('type.list.list'('type.bvlSem.v'),'type.list.list'('type.bvlSem.v')),'_i_0_0'),s('type.list.list'('type.bvlSem.v'),X0)))) )),
  inference(superposition,[],[f391,f305])).
fof(f2632,plain,(
  ( ! [X185,X186,X184,X183] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X186))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun('type.num.num','type.integer.int'),X184),s('type.num.num',X185))))) | s('type.integer.int',X186) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X183),s(fun('type.num.num','type.integer.int'),X184),s('type.num.num','const.num.SUC_1'(s('type.num.num',X185)))))) )),
  inference(superposition,[],[f332,f320])).
fof(f2627,plain,(
  ( ! [X167,X165,X163,X166,X164,X162] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X166),s('type.list.list'('type.bvlSem.v'),X167))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X163),s('type.num.num',X164))),s('type.list.list'('type.bvlSem.v'),X165))) | s('type.num.num',X166) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X162),s(fun('type.num.num','type.num.num'),X163),s('type.num.num','const.num.SUC_1'(s('type.num.num',X164)))))) )),
  inference(superposition,[],[f390,f320])).
fof(f2623,plain,(
  ( ! [X146,X149,X147,X148] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X149))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X147),s('type.num.num',X148))))) | s('type.num.num',X149) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X146),s(fun('type.num.num','type.num.num'),X147),s('type.num.num','const.num.SUC_1'(s('type.num.num',X148)))))) )),
  inference(superposition,[],[f330,f320])).
fof(f2663,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130)))))) | ~p(s(bool,happ(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130))),s(X127,sK4(X131,'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))),X127)))))) )),
  inference(forward_demodulation,[],[f2618,f320])).
fof(f2618,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130))),s(X127,sK4(X131,'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))),X127))))) | p(s(bool,'|dollar|exists_1'(s(fun(X127,bool),'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))))))))) )),
  inference(superposition,[],[f315,f320])).
fof(f2516,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2488,f254])).
fof(f2488,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f436,f2434])).
fof(f2428,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f356,f2405,f436])).
fof(f2328,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2307,f254])).
fof(f2307,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2219,f429])).
fof(f2327,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2306,f254])).
fof(f2306,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2219,f461])).
fof(f2322,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2298,f254])).
fof(f2298,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2219,f434])).
fof(f2239,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X68),s('type.list.list'('type.bvlSem.v'),X69))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X70),s('type.list.list'('type.bvlSem.v'),X71))) | s('type.num.num',X70) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X67),'const.sum.INL_1'(s('type.num.num',X68)))))) )),
  inference(superposition,[],[f390,f337])).
fof(f2238,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X63),s('type.list.list'('type.bvlSem.v'),X64))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X65),s('type.list.list'('type.bvlSem.v'),X66))) | s('type.num.num',X65) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X61,'type.num.num'),'const.pair.,_2'(s(X61,X62),s('type.num.num',X63)))))) )),
  inference(superposition,[],[f390,f341])).
fof(f2237,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.bvlSem.v'),X58))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X59),s('type.list.list'('type.bvlSem.v'),X60))) | s('type.num.num',X59) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X55),'const.pair.,_2'(s('type.num.num',X56),s(X55,X57)))))) )),
  inference(superposition,[],[f390,f340])).
fof(f2236,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.bvlSem.v'),X52))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X53),s('type.list.list'('type.bvlSem.v'),X54))) | s('type.num.num',X53) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X50),s(fun('type.num.num','type.num.num'),X51),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f390,f312])).
fof(f2235,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X46),s('type.list.list'('type.bvlSem.v'),X47))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X48),s('type.list.list'('type.bvlSem.v'),X49))) | s('type.num.num',X48) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X45,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f390,f336])).
fof(f2234,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X43),s('type.list.list'('type.bvlSem.v'),X44))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X40),X41))),s('type.list.list'('type.bvlSem.v'),X42))) | s('type.num.num',X43) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X40),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X40),X41)))) )),
  inference(superposition,[],[f390,f338])).
fof(f2233,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X36),s('type.list.list'('type.bvlSem.v'),X37))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X38),s('type.list.list'('type.bvlSem.v'),X39))) | s('type.num.num',X38) = s('type.num.num',happ(s(fun('type.list.list'(X34),'type.num.num'),sK16(X35,X36,X34,'type.num.num')),s('type.list.list'(X34),'const.list.NIL_0')))) )),
  inference(superposition,[],[f390,f356])).
fof(f2232,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X30),s('type.list.list'('type.bvlSem.v'),X31))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X32),s('type.list.list'('type.bvlSem.v'),X33))) | s('type.num.num',X32) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X30)))) )),
  inference(superposition,[],[f390,f305])).
fof(f2230,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X24),s('type.list.list'('type.bvlSem.v'),X25))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X23))) | s('type.num.num',X24) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X22)))))) )),
  inference(superposition,[],[f390,f264])).
fof(f2228,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.bvlSem.v'),X17))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.bvlSem.v'),X15))) | s('type.num.num',X16) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f390,f261])).
fof(f2226,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X7),s('type.list.list'('type.bvlSem.v'),X8))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X5))))),s('type.list.list'('type.bvlSem.v'),X6))) | s('type.num.num',X7) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4)))))) )),
  inference(superposition,[],[f390,f813])).
fof(f2225,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.bvlSem.v'),X1))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f390,f261])).
fof(f2222,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2190,f254])).
fof(f2190,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f436,f461])).
fof(f2221,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2220,f305])).
fof(f2220,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2187,f254])).
fof(f2187,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f406,f461])).
fof(f2217,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f2172,f305])).
fof(f2172,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38)))))) )),
  inference(superposition,[],[f461,f337])).
fof(f2216,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(forward_demodulation,[],[f2171,f305])).
fof(f2171,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(superposition,[],[f461,f341])).
fof(f2215,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2170,f305])).
fof(f2170,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f461,f340])).
fof(f2214,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2169,f305])).
fof(f2169,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f461,f312])).
fof(f2213,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28))))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f2168,f305])).
fof(f2168,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28)))))) )),
  inference(superposition,[],[f461,f336])).
fof(f2212,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2166,f305])).
fof(f2166,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f461,f356])).
fof(f2062,plain,(
  ( ! [X4,X3] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X4))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num',X4) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3)))))) )),
  inference(superposition,[],[f330,f264])).
fof(f1849,plain,(
  ( ! [X2,X3] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X3))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X2))))))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f330,f263])).
fof(f1832,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f263,f261])).
fof(f1700,plain,(
  ( ! [X111,X112,X113] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X113))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X111),X112))))) | s('type.integer.int',X113) = s('type.integer.int',happ(s(fun('type.pair.prod'('type.integer.int',X111),'type.integer.int'),'const.pair.FST_0'),s('type.pair.prod'('type.integer.int',X111),X112)))) )),
  inference(superposition,[],[f332,f338])).
fof(f1693,plain,(
  ( ! [X94,X95,X93] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X95))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X93),X94))))) | s('type.num.num',X95) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X93),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X93),X94)))) )),
  inference(superposition,[],[f330,f338])).
fof(f1716,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1689,f338])).
fof(f1689,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f315,f338])).
fof(f1672,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f315,f338])).
fof(f1565,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1552,f254])).
fof(f1552,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f315,f434])).
fof(f1551,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f315,f266])).
fof(f1549,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f315,f434])).
fof(f1547,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f315,f434])).
fof(f1564,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26)))))) )),
  inference(forward_demodulation,[],[f1546,f337])).
fof(f1546,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28))))))))) )),
  inference(superposition,[],[f315,f337])).
fof(f1563,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24)))) | ~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21)))))) )),
  inference(forward_demodulation,[],[f1545,f341])).
fof(f1545,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24))))))))) )),
  inference(superposition,[],[f315,f341])).
fof(f1562,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1544,f340])).
fof(f1544,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f315,f340])).
fof(f1561,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13)))) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f1543,f312])).
fof(f1543,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f315,f312])).
fof(f1560,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8)))))) )),
  inference(forward_demodulation,[],[f1542,f336])).
fof(f1542,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10))))))))) )),
  inference(superposition,[],[f315,f336])).
fof(f1559,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1541,f356])).
fof(f1541,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f315,f356])).
fof(f1558,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1540,f305])).
fof(f1540,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f315,f305])).
fof(f1533,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1290,f254])).
fof(f1290,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f436,f434])).
fof(f1532,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1289,f254])).
fof(f1289,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f406,f434])).
fof(f1531,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1288,f254])).
fof(f1288,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f329,f434])).
fof(f1179,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f434,f337])).
fof(f1178,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f434,f341])).
fof(f1177,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f434,f340])).
fof(f1176,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f434,f312])).
fof(f1175,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X21,bool),'const.sum.INR_1'(s(bool,X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f434,f336])).
fof(f1174,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK16(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f434,f356])).
fof(f1173,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f434,f305])).
fof(f1163,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f434,f434])).
fof(f1146,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43)))))) )),
  inference(superposition,[],[f434,f337])).
fof(f1145,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))) )),
  inference(superposition,[],[f434,f341])).
fof(f1144,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f434,f340])).
fof(f1143,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f434,f312])).
fof(f1142,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X28,bool),'const.sum.INR_1'(s(bool,X29)))))) )),
  inference(superposition,[],[f434,f336])).
fof(f1141,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK16(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f434,f356])).
fof(f1140,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f434,f305])).
fof(f1103,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X125))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X126))) | s('type.integer.int',X126) = s('type.integer.int',happ(s(fun('type.list.list'(X123),'type.integer.int'),sK16(X124,X125,X123,'type.integer.int')),s('type.list.list'(X123),'const.list.NIL_0')))) )),
  inference(superposition,[],[f332,f356])).
fof(f1096,plain,(
  ( ! [X101,X99,X100,X98] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X100))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X101))) | s('type.num.num',X101) = s('type.num.num',happ(s(fun('type.list.list'(X98),'type.num.num'),sK16(X99,X100,X98,'type.num.num')),s('type.list.list'(X98),'const.list.NIL_0')))) )),
  inference(superposition,[],[f330,f356])).
fof(f1072,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool)),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f356,f266])).
fof(f1017,plain,(
  ( ! [X17,X18,X16] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))) | s('type.integer.int',X18) = s('type.integer.int','const.sum.OUTL_1'(s('type.sum.sum'('type.integer.int',X16),'const.sum.INL_1'(s('type.integer.int',X17)))))) )),
  inference(superposition,[],[f332,f337])).
fof(f1016,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X14))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))) | s('type.integer.int',X15) = s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'(X12,'type.integer.int'),'const.pair.,_2'(s(X12,X13),s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f332,f341])).
fof(f1015,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X8),'const.pair.,_2'(s('type.integer.int',X9),s(X8,X10)))))) )),
  inference(superposition,[],[f332,f340])).
fof(f1014,plain,(
  ( ! [X6,X7,X5] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X5))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X5),s(fun('type.num.num','type.integer.int'),X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f332,f312])).
fof(f1013,plain,(
  ( ! [X4,X2,X3] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.sum.OUTR_1'(s('type.sum.sum'(X2,'type.integer.int'),'const.sum.INR_1'(s('type.integer.int',X3)))))) )),
  inference(superposition,[],[f332,f336])).
fof(f1012,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int',happ(s(fun('type.integer.int','type.integer.int'),'_i_0_0'),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f332,f305])).
fof(f976,plain,(
  ( ! [X8,X9] : (s('type.num.num',X9) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X9))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8)))))))) )),
  inference(superposition,[],[f331,f261])).
fof(f974,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num',X4) != s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X4))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2)))))))) )),
  inference(superposition,[],[f331,f813])).
fof(f973,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f331,f261])).
fof(f964,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f261,f331])).
fof(f966,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f813,f331])).
fof(f940,plain,(
  ( ! [X30,X28,X29] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X29))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X30))) | s('type.num.num',X30) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X28),'const.sum.INL_1'(s('type.num.num',X29)))))) )),
  inference(superposition,[],[f330,f337])).
fof(f939,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X26))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X27))) | s('type.num.num',X27) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X24,'type.num.num'),'const.pair.,_2'(s(X24,X25),s('type.num.num',X26)))))) )),
  inference(superposition,[],[f330,f341])).
fof(f938,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X21))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X23))) | s('type.num.num',X23) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X20),'const.pair.,_2'(s('type.num.num',X21),s(X20,X22)))))) )),
  inference(superposition,[],[f330,f340])).
fof(f937,plain,(
  ( ! [X19,X17,X18] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X17))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X19))) | s('type.num.num',X19) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X17),s(fun('type.num.num','type.num.num'),X18),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f330,f312])).
fof(f936,plain,(
  ( ! [X14,X15,X16] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X15))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X14,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X15)))))) )),
  inference(superposition,[],[f330,f336])).
fof(f935,plain,(
  ( ! [X12,X13] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X12))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f330,f305])).
fof(f933,plain,(
  ( ! [X8,X9] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X9))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num',X9) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8)))))) )),
  inference(superposition,[],[f330,f261])).
fof(f931,plain,(
  ( ! [X4,X2,X3] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X4))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X3))))))) | s('type.num.num',X4) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X2)))))) )),
  inference(superposition,[],[f330,f813])).
fof(f930,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))) != s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f330,f261])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f865,f256])).
fof(f865,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f314,f266])).
fof(f813,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))))) )),
  inference(superposition,[],[f261,f261])).
fof(f806,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f793,f254])).
fof(f793,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f436,f429])).
fof(f745,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f341,f266])).
fof(f685,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f340,f266])).
fof(f629,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f312,f266])).
fof(f529,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f337,f266])).
fof(f502,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f336,f266])).
fof(f483,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f478,f256])).
fof(f478,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f326,f266])).
fof(f436,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f271,f266])).
fof(f429,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f257,f268])).
fof(f393,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X21,X22,X42,X23] : ! [X24,X25,X26] : s(fun(X22,X21),'_dst_x0x1_2'(s(fun(X42,X21),X24),s(fun(X22,X42),happ(s(fun(X23,fun(X22,X42)),X25),s(X23,X26))))) = s(fun(X22,X21),happ(s(fun(X23,fun(X22,X21)),'_dst_x00x11_2'(s(fun(X42,X21),X24),s(fun(X23,fun(X22,X42)),X25))),s(X23,X26)))),
  file('thm.bvlSem.v_|49||49||split|4','_dst_x00x11_ax')).
fof(f391,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) != s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | s('type.list.list'('type.bvlSem.v'),X1) = s('type.list.list'('type.bvlSem.v'),X3)) )),
  inference(cnf_transformation,[],[f251])).
fof(f392,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.bvlSem.v'),X1) != s('type.list.list'('type.bvlSem.v'),X3) | s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.bvlSem.v'),X3))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f406,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X28,X45,X35,X46] : ((s(bool,X28) = s(bool,X45) & (p(s(bool,X45)) => s(bool,X35) = s(bool,X46))) => ((p(s(bool,X28)) => p(s(bool,X35))) <=> (p(s(bool,X45)) => p(s(bool,X46)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.IMP_CONG')).
fof(f365,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X21,X22,X23] : ! [X24,X25,X26] : s(X21,happ(s(fun(X22,X21),X24),s(X22,happ(s(fun(X23,X22),X25),s(X23,X26))))) = s(X21,happ(s(fun(X23,X21),'_dst_x0x1_2'(s(fun(X22,X21),X24),s(fun(X23,X22),X25))),s(X23,X26)))),
  file('thm.bvlSem.v_|49||49||split|4','_dst_x0x1_ax')).
fof(f364,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X21,X22,X23] : ! [X25,X26] : s(fun(X22,X21),'_dst_x0x1_2'(s(fun(X23,X21),X26),s(fun(X22,X23),X25))) = s(fun(X22,X21),happ(s(fun(fun(X23,X21),fun(X22,X21)),'_dstsel_0_x01_1'(s(fun(X22,X23),X25))),s(fun(X23,X21),X26)))),
  file('thm.bvlSem.v_|49||49||split|4','_dstsel_0_x01_ax')).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X27,X9] : ! [X36,X11] : ? [X37] : (s(X27,X36) = s(X27,happ(s(fun('type.list.list'(X9),X27),X37),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X15,X16] : s(X27,happ(s(fun('type.list.list'(X9),X27),X37),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X15),s('type.list.list'(X9),X16))))) = s(X27,happ(s(fun(X27,X27),happ(s(fun('type.list.list'(X9),fun(X27,X27)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X27,X27))),X11),s(X9,X15))),s('type.list.list'(X9),X16))),s(X27,happ(s(fun('type.list.list'(X9),X27),X37),s('type.list.list'(X9),X16))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.list.list_Axiom')).
fof(f357,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f351,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X27,X9] : ! [X10,X50] : (s(fun(X9,X27),X10) = s(fun(X9,X27),X50) <=> ! [X28] : s(X27,happ(s(fun(X9,X27),X10),s(X9,X28))) = s(X27,happ(s(fun(X9,X27),X50),s(X9,X28))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.FUN_EQ_THM')).
fof(f352,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f240])).
fof(f350,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.bvlSem.v_|49||49||split|4',aHL_EXT)).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X27] : ! [X28,X35] : s(X27,X35) = s(X27,'const.pair.SND_1'(s('type.pair.prod'(X9,X27),'const.pair.,_2'(s(X9,X28),s(X27,X35)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.pair.SND')).
fof(f340,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X27,X9] : ! [X28,X35] : s(X9,X28) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X27),'const.pair.,_2'(s(X9,X28),s(X27,X35)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.pair.FST')).
fof(f339,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X27,X9] : ! [X38] : ? [X10] : ! [X39,X40,X41] : s(X27,happ(s(fun('type.ind_type.recspace'(X9),X27),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X39),s(X9,X40),s(fun('type.num.num','type.ind_type.recspace'(X9)),X41))))) = s(X27,happ(s(fun(fun('type.num.num',X27),X27),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X27),X27)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X27),X27))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X27),X27)))),X38),s('type.num.num',X39))),s(X9,X40))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X41))),s(fun('type.num.num',X27),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X27),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X41)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.CONSTR_REC')).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X27] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X27),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X27),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X27),X8)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_1_const.pair.FST')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9,X27] : ! [X28] : s(X27,X28) = s(X27,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X27),'const.sum.INR_1'(s(X27,X28)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.sum.OUTR')).
fof(f334,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) & (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_|49||49||split|1')).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))) )),
  inference(cnf_transformation,[],[f232])).
fof(f332,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_|49||49||split|0')).
fof(f333,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f231])).
fof(f331,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f230])).
fof(f326,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f325,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X1)))))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X15,X16] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X15),s('type.list.list'('type.bvlSem.v'),X16))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X16)))))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_size_def|split|2')).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : ! [X8,X19,X53] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X19),s('type.num.num',X53))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X19))),s('type.num.num',X53)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f321,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X19,X53] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X19),s(fun('type.num.num','type.ind_type.recspace'(X9)),X53))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X19))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X53)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_1_3_const.ind_type.CONSTR')).
fof(f320,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X6,X10,X34] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X34))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.FCONS|split|1')).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f271,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X20] : (s(bool,'F_0') = s(bool,X20) <=> ~p(s(bool,X20)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|2')).
fof(f272,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f265,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_1_const.bvlSem.Word64')).
fof(f264,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_size_def|split|1')).
fof(f418,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f263,f262])).
fof(f263,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_size_def|split|4')).
fof(f262,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_size_def|split|3')).
fof(f261,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6)))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_size_def|split|0')).
fof(f259,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X8)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_1_const.bvlSem.Number')).
fof(f258,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X8))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X8)))),
  file('thm.bvlSem.v_|49||49||split|4','ARITY_0_1_const.bvlSem.RefPtr')).
fof(f257,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.bvlSem.v_|49||49||split|4','thm.ind_type.ISO_REFL')).
fof(f7425,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6784,f6677])).
fof(f6677,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f6639,f256])).
fof(f6639,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f5241,f2937])).
fof(f6784,plain,(
  ( ! [X23,X22] : (s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(subsumption_resolution,[],[f6739,f254])).
fof(f6739,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(superposition,[],[f436,f6629])).
fof(f6629,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f421,f5241,f2219])).
fof(f9733,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl22),
  inference(cnf_transformation,[],[f9733_D])).
fof(f9733_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f16819,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f16658])).
fof(f16658,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f7425,f10662])).
fof(f10662,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10652])).
fof(f10652,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9798,f434])).
fof(f9798,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2405,f2434,f2428,f2516,f2517,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3047,f2937,f3083,f3099,f3454,f3455,f3456,f3506,f3507,f3509,f3510,f3512,f3513,f3514,f3549,f3578,f3580,f3581,f3583,f3584,f3585,f3692,f3693,f3766,f3767,f4207,f4214,f3862,f3863,f4241,f4274,f4275,f4277,f4278,f4280,f4281,f4282,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4816,f4646,f4647,f4700,f4702,f4703,f4818,f4718,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5074,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f5004,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f3085,f2407,f346,f345,f317,f316,f313,f4212,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f3682,f3852,f5235,f6605,f4899,f4511,f4393,f4276,f3579,f3508,f3049,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f4279,f3582,f3511,f3052,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7066,f7068,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7953,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8328,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8556,f8559,f8563,f8659,f8667,f8753,f8758,f8760,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8889,f8890,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9045,f9046,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9151,f9152,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9785,f9787,f9789,f9790,f9792,f9793,f9794,f9795,f9797])).
fof(f9797,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9796])).
fof(f9796,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f7892,f3767])).
fof(f7892,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f7662,f2937])).
fof(f16476,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) | $spl36),
  inference(cnf_transformation,[],[f16476_D])).
fof(f16476_D,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f16401,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f16244])).
fof(f16244,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f7425,f10548])).
fof(f10548,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10538])).
fof(f10538,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9793,f434])).
fof(f13270,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl32),
  inference(cnf_transformation,[],[f13270_D])).
fof(f13270_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13272,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl34),
  inference(cnf_transformation,[],[f13272_D])).
fof(f13272_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13255,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl30),
  inference(superposition,[],[f12482,f305])).
fof(f2409,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl12),
  inference(cnf_transformation,[],[f2409_D])).
fof(f2409_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12482,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl30),
  inference(cnf_transformation,[],[f12482_D])).
fof(f12482_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12488,plain,(
  ( ! [X59,X60] : (p(s(bool,sK5('_i_0_0',X60,bool))) | ~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool)))))) )),
  inference(forward_demodulation,[],[f12457,f305])).
fof(f12457,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X60,bool)))))) )),
  inference(superposition,[],[f316,f305])).
fof(f2403,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2403_D])).
fof(f2403_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f12481,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f12480,f305])).
fof(f12480,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f12444,f254])).
fof(f12444,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f461])).
fof(f10220,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209,f10211,f10213,f10215,f10217,f10219])).
fof(f10219,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10085,f254])).
fof(f10085,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (~p(s(bool,'T_0')) | s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9167])).
fof(f10217,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10084,f254])).
fof(f10084,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (~p(s(bool,'T_0')) | s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9169])).
fof(f10215,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10083,f254])).
fof(f10083,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (~p(s(bool,'T_0')) | s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9062])).
fof(f10213,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10082,f254])).
fof(f10082,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (~p(s(bool,'T_0')) | s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9064])).
fof(f10211,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10081,f254])).
fof(f10081,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (~p(s(bool,'T_0')) | s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9168])).
fof(f10209,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10080,f254])).
fof(f10080,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (~p(s(bool,'T_0')) | s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9170])).
fof(f10207,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10079,f254])).
fof(f10079,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (~p(s(bool,'T_0')) | s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9063])).
fof(f10205,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10078,f254])).
fof(f10078,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (~p(s(bool,'T_0')) | s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9065])).
fof(f10201,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10069,f254])).
fof(f10069,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f354])).
fof(f10199,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f10068,f254])).
fof(f10068,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f353])).
fof(f10197,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(subsumption_resolution,[],[f10067,f254])).
fof(f10067,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f349])).
fof(f10195,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(subsumption_resolution,[],[f10066,f254])).
fof(f10066,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f348])).
fof(f10126,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f9836,f9732])).
fof(f9836,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f271])).
fof(f2597,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f266,f320])).
fof(f10218,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209,f10211,f10213,f10215,f10217])).
fof(f10216,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209,f10211,f10213,f10215])).
fof(f10214,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209,f10211,f10213])).
fof(f10212,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209,f10211])).
fof(f10210,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207,f10209])).
fof(f10208,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205,f10207])).
fof(f10206,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201,f10205])).
fof(f10202,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199,f10201])).
fof(f10200,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197,f10199])).
fof(f10198,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195,f10197])).
fof(f10196,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126,f10195])).
fof(f10141,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9997,f9732])).
fof(f9997,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9732,f9798])).
fof(f10140,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9996,f9732])).
fof(f9996,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9732,f9793])).
fof(f10139,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f10138,f254])).
fof(f10138,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(forward_demodulation,[],[f10137,f9732])).
fof(f10137,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9985,f9732])).
fof(f9985,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114))))) | ~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9120])).
fof(f10136,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f10135,f254])).
fof(f10135,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(forward_demodulation,[],[f10134,f9732])).
fof(f10134,plain,(
  ( ! [X125,X123,X121,X124,X122] : (~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9973,f9732])).
fof(f9973,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X123,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X123,X121),X124))),s(X123,X125))) | ~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) ) | $spl20),
  inference(backward_demodulation,[],[f9732,f9014])).
fof(f10127,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f253,f252,f254,f256,f257,f258,f259,f260,f261,f262,f263,f418,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f329,f326,f331,f330,f333,f332,f335,f334,f336,f337,f338,f339,f340,f341,f350,f352,f351,f357,f356,f364,f365,f406,f392,f391,f390,f393,f421,f429,f434,f436,f461,f483,f502,f529,f629,f685,f745,f806,f813,f866,f930,f931,f933,f935,f936,f937,f938,f939,f940,f966,f964,f973,f974,f976,f1012,f1013,f1014,f1015,f1016,f1017,f1072,f1096,f1103,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1163,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1531,f1532,f1533,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1547,f1549,f1551,f1565,f1557,f1672,f1716,f1693,f1700,f1832,f1849,f2062,f2212,f2213,f2214,f2215,f2216,f2217,f2219,f2221,f2222,f2225,f2226,f2228,f2230,f2232,f2233,f2234,f2235,f2236,f2237,f2238,f2239,f2322,f2327,f2328,f2402,f2408,f2663,f2623,f2627,f2632,f2804,f2805,f2806,f2807,f2808,f2809,f2810,f2811,f2812,f3454,f3455,f3456,f4392,f4394,f4395,f4397,f4398,f4399,f4478,f4510,f4512,f4513,f4515,f4516,f4517,f4597,f4638,f4639,f4640,f4700,f4702,f4703,f4722,f4732,f4736,f4743,f4753,f4822,f4985,f5028,f5022,f5210,f5234,f5779,f5782,f5807,f5810,f5835,f5838,f5768,f5029,f4898,f4999,f4652,f4655,f4659,f4658,f4657,f4656,f4654,f4653,f4651,f4650,f5128,f4979,f394,f355,f354,f353,f349,f348,f2387,f3130,f3133,f347,f344,f4919,f4817,f2407,f346,f345,f317,f316,f313,f3137,f3136,f3135,f3134,f3132,f3131,f3129,f3128,f408,f307,f3462,f3463,f3464,f5235,f6605,f4899,f4511,f4393,f2410,f2167,f1663,f1661,f1659,f4514,f4396,f2611,f2585,f2583,f5757,f4745,f4734,f4725,f4935,f4945,f5084,f5094,f4724,f4721,f4833,f4707,f4706,f4701,f4699,f4694,f4693,f4691,f4690,f4689,f4688,f4687,f4723,f864,f863,f862,f2619,f861,f860,f1690,f1093,f859,f4821,f4820,f314,f7053,f7063,f7064,f7065,f7069,f7073,f7076,f7043,f7044,f7045,f7046,f7047,f7048,f7049,f7050,f7051,f7052,f7571,f7572,f7574,f7575,f7576,f7577,f7579,f7662,f7670,f7714,f7715,f7717,f7718,f7719,f7720,f7722,f7806,f7814,f7916,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f8172,f8174,f8176,f8178,f8115,f8246,f8145,f8152,f8156,f8290,f8291,f8292,f8293,f8294,f8295,f8296,f8297,f8298,f8486,f8487,f8508,f8872,f9216,f9217,f9218,f8860,f8857,f8856,f8841,f8840,f8839,f9244,f9245,f9246,f8818,f8875,f8876,f8877,f8878,f8879,f8880,f8881,f8882,f8883,f8894,f9009,f9010,f9012,f9013,f9014,f9015,f9016,f9018,f9019,f9024,f9025,f9026,f9027,f9028,f9031,f9032,f9033,f9034,f9047,f9050,f9051,f9055,f9056,f9057,f9058,f9059,f9060,f9062,f9063,f9064,f9065,f9074,f9076,f9078,f9085,f9087,f9096,f9098,f9114,f9115,f9116,f9118,f9119,f9120,f9121,f9122,f9124,f9125,f9130,f9131,f9132,f9133,f9134,f9137,f9138,f9139,f9140,f9153,f9156,f9157,f9160,f9161,f9162,f9163,f9164,f9165,f9167,f9168,f9169,f9170,f9179,f9181,f9183,f9190,f9192,f9201,f9203,f9117,f9789,f9790,f9794,f9795,f6985,f7005,f4897,f4917,f2597,f9833,f9732,f10126])).
fof(f10125,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f9835,f254])).
fof(f9835,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f9732,f256])).
fof(f9834,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl28),
  inference(cnf_transformation,[],[f9834_D])).
fof(f9834_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1567,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1567_D])).
fof(f1567_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9833,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2178,f254])).
fof(f2178,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f315,f461])).
fof(f4917,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4884,f254])).
fof(f4884,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4817,f461])).
fof(f4897,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4817,f305])).
fof(f7005,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6972,f254])).
fof(f6972,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4817,f461])).
fof(f6985,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4817,f305])).
fof(f9788,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f9788_D])).
fof(f9788_D,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9795,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f7900,f7662])).
fof(f7900,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f7662,f2213])).
fof(f9794,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f7911,f7806])).
fof(f7911,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f7662,f2217])).
fof(f9786,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl24),
  inference(cnf_transformation,[],[f9786_D])).
fof(f9786_D,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9790,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f8274,f7662])).
fof(f8274,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f7806,f2213])).
fof(f9789,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f8285,f7806])).
fof(f8285,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f7806,f2217])).
fof(f9787,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(subsumption_resolution,[],[f8614,f8659])).
fof(f8614,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f7662,f3047])).
fof(f9785,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(subsumption_resolution,[],[f8615,f8659])).
fof(f8615,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f7806,f3047])).
fof(f9666,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9529])).
fof(f9529,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2722,f8036])).
fof(f8036,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f8022])).
fof(f8022,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f7951,f434])).
fof(f7951,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f7879,f1582])).
fof(f1582,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1568,f272])).
fof(f7879,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1583,f7662])).
fof(f1583,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1568,f267])).
fof(f2722,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2721,f256])).
fof(f2721,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2687,f1582])).
fof(f2687,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2221,f1606])).
fof(f1606,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1584,f254])).
fof(f1584,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1568,f434])).
fof(f7071,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl18),
  inference(cnf_transformation,[],[f7071_D])).
fof(f7071_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f7070,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f7036,f271])).
fof(f7036,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4999,f2329])).
fof(f2329,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl6),
  inference(subsumption_resolution,[],[f2308,f254])).
fof(f2308,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl6),
  inference(superposition,[],[f2219,f1795])).
fof(f1795,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1742,f268])).
fof(f1742,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1583,f314])).
fof(f7067,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f7067_D])).
fof(f7067_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7066,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f7033,f271])).
fof(f7033,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4999,f2517])).
fof(f6604,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6302,f256])).
fof(f6302,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f421,f2409,f326])).
fof(f6603,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6304,f2409])).
fof(f6304,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1583,f2409,f326])).
fof(f6602,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6312,f256])).
fof(f6312,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f421,f2409,f326])).
fof(f6601,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6314,f2409])).
fof(f6314,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1583,f2409,f326])).
fof(f6600,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6592,f256])).
fof(f6592,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(backward_demodulation,[],[f6375,f5029])).
fof(f6375,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f337,f2409,f2221])).
fof(f6503,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6502,f421])).
fof(f6502,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f6378,f429])).
fof(f6378,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f429,f2409,f2221])).
fof(f6501,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6500,f421])).
fof(f6500,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f6379,f1795])).
fof(f6379,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1795,f2409,f2221])).
fof(f6499,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6498,f421])).
fof(f6498,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f6380,f5028])).
fof(f6380,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f5028,f2409,f2221])).
fof(f6497,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f6496,f421])).
fof(f6496,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f6381,f2434])).
fof(f6381,plain,(
  s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f2434,f2409,f2221])).
fof(f6399,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f2409,f4898])).
fof(f6397,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1583,f2409,f314])).
fof(f6393,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f429,f421,f2409,f2221])).
fof(f6363,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1583,f1583,f2409,f483])).
fof(f6361,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f421,f421,f2409,f483])).
fof(f6339,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f337,f2409,f406])).
fof(f6337,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f341,f2409,f406])).
fof(f6335,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f340,f2409,f406])).
fof(f6333,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f312,f2409,f406])).
fof(f6331,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f336,f2409,f406])).
fof(f6329,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f356,f2409,f406])).
fof(f6327,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f305,f2409,f406])).
fof(f6324,plain,(
  $false | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f254,f2434,f2409,f406])).
fof(f6322,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f5028,f2409,f406])).
fof(f6320,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f254,f1795,f2409,f406])).
fof(f6318,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f429,f2409,f406])).
fof(f6315,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2409,f1583,f2409,f326])).
fof(f6313,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f421,f2409,f326])).
fof(f6311,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2409,f1583,f2409,f326])).
fof(f6309,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2409,f421,f2409,f326])).
fof(f6305,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2409,f1583,f2409,f326])).
fof(f6303,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f421,f2409,f326])).
fof(f6301,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2409,f1583,f2409,f326])).
fof(f6299,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2409,f421,f2409,f326])).
fof(f6295,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f421,f2409,f272])).
fof(f6293,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f429,f2409,f267])).
fof(f6291,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f254,f2409])).
fof(f5241,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl14),
  inference(cnf_transformation,[],[f5241_D])).
fof(f5241_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5234,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4899,f340])).
fof(f2405,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2405_D])).
fof(f2405_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2408,plain,(
  ( ! [X24] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f2386,f305])).
fof(f2386,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f313,f305])).
fof(f2402,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2401,f305])).
fof(f2401,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2379,f254])).
fof(f2379,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f313,f461])).
fof(f1568,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1557,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f315,f305])).
fof(f420,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f419])).
fof(f419,plain,(
  s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f415,f416])).
fof(f414,plain,(
  s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f416,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f252,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f195])).
fof(f195,plain,(
  ? [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ? [X0,X1] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X6))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X6))) = s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.bvlSem.v_|49||49||split|4','thm.bvlSem.v_|49||49||split|4')).
fof(f415,plain,(
  s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f417,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f253,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
% SZS output end Proof for thm.bvlSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 177
% Passive clauses: 1783
% Generated clauses: 20429
% Final active clauses: 125
% Final passive clauses: 1062
% Input formulas: 69
% Initial clauses: 143
% 
% Duplicate literals: 691
% Trivial inequalities: 174
% Fw subsumption resolutions: 874
% Fw demodulations: 3722
% Bw demodulations: 417
% Condensations: 904
% Global subsumptions: 16
% 
% Simple tautologies: 275
% Equational tautologies: 144
% Forward subsumptions: 10912
% Fw demodulations to eq. taut.: 848
% Bw demodulations to eq. taut.: 60
% 
% Unit resulting resolution: 2229
% Factoring: 1
% Forward superposition: 5754
% Backward superposition: 4664
% Self superposition: 144
% Equality factoring: 76
% Equality resolution: 31
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 1913
% SAT solver unit clauses: 131
% SAT solver binary clauses: 734
% 
% Sat splits: 26
% Sat splitting refutations: 60
% 
% Memory used [KB]: 16630
% Time elapsed: 0.789 s
% ------------------------------
% ------------------------------
% Success in time 2.12 s
