Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bviSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bviSem
% SZS output start Proof for thm.bviSem
fof(f14405,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f126,f190,f188,f186,f184,f182,f180,f13106,f121,f181,f122,f183,f123,f124,f125,f189,f179,f14394,f14392,f14388,f14389,f187,f14390,f185])).
fof(f185,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK8) | $spl6),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK8) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f14390,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK8) | $spl0),
  inference(forward_demodulation,[],[f14352,f157])).
fof(f157,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bviSem.state_global_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bviSem.state_global_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bviSem.state_global_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X27,X38,X44,X45,X46] : s('type.option.option'('type.num.num'),X44) = s('type.option.option'('type.num.num'),'const.bviSem.state_global_1'(s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27),s('type.num.num',X38),s('type.option.option'('type.num.num'),X44),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45),s('type.ffi.ffi_state'(X6),X46)))))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_global')).
fof(f14352,plain,(
  s('type.option.option'('type.num.num'),sK8) = s('type.option.option'('type.num.num'),'const.bviSem.state_global_1'(s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))))) | $spl0),
  inference(superposition,[],[f157,f179])).
fof(f187,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | $spl8),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f14389,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | $spl0),
  inference(forward_demodulation,[],[f14351,f158])).
fof(f158,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),'const.bviSem.state_code_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),'const.bviSem.state_code_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),'const.bviSem.state_code_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X27,X38,X44,X45,X46] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),'const.bviSem.state_code_1'(s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27),s('type.num.num',X38),s('type.option.option'('type.num.num'),X44),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45),s('type.ffi.ffi_state'(X6),X46)))))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_code')).
fof(f14351,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),'const.bviSem.state_code_1'(s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))))) | $spl0),
  inference(superposition,[],[f158,f179])).
fof(f14388,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f14387,f182])).
fof(f14387,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | $spl0),
  inference(forward_demodulation,[],[f14350,f159])).
fof(f159,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bviSem.state_refs_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bviSem.state_refs_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bviSem.state_refs_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X27,X38,X44,X45,X46] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bviSem.state_refs_1'(s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27),s('type.num.num',X38),s('type.option.option'('type.num.num'),X44),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45),s('type.ffi.ffi_state'(X6),X46)))))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_refs')).
fof(f14350,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bviSem.state_refs_1'(s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))))) | $spl0),
  inference(superposition,[],[f159,f179])).
fof(f14392,plain,(
  $false | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f14391,f190])).
fof(f14391,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) = s('type.ffi.ffi_state'(sK0),sK10) | $spl0),
  inference(forward_demodulation,[],[f14353,f156])).
fof(f156,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),'const.bviSem.state_ffi_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),'const.bviSem.state_ffi_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),'const.bviSem.state_ffi_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X27,X38,X44,X45,X46] : s('type.ffi.ffi_state'(X6),X46) = s('type.ffi.ffi_state'(X6),'const.bviSem.state_ffi_1'(s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27),s('type.num.num',X38),s('type.option.option'('type.num.num'),X44),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45),s('type.ffi.ffi_state'(X6),X46)))))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_ffi')).
fof(f14353,plain,(
  s('type.ffi.ffi_state'(sK0),sK10) = s('type.ffi.ffi_state'(sK0),'const.bviSem.state_ffi_1'(s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))))) | $spl0),
  inference(superposition,[],[f156,f179])).
fof(f14394,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f14393,f184])).
fof(f14393,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK7) | $spl0),
  inference(forward_demodulation,[],[f14354,f155])).
fof(f155,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num',X2) = s('type.num.num','const.bviSem.state_clock_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.num.num',X2) = s('type.num.num','const.bviSem.state_clock_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.num.num',X2) = s('type.num.num','const.bviSem.state_clock_1'(s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X27,X38,X44,X45,X46] : s('type.num.num',X38) = s('type.num.num','const.bviSem.state_clock_1'(s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X27),s('type.num.num',X38),s('type.option.option'('type.num.num'),X44),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X45),s('type.ffi.ffi_state'(X6),X46)))))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_clock')).
fof(f14354,plain,(
  s('type.num.num',sK7) = s('type.num.num','const.bviSem.state_clock_1'(s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))))) | $spl0),
  inference(superposition,[],[f155,f179])).
fof(f179,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) | $spl0),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f189,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) = s('type.ffi.ffi_state'(sK0),sK10) | $spl10),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) = s('type.ffi.ffi_state'(sK0),sK10) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f125,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) = s('type.ffi.ffi_state'(sK0),sK10) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  (s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) | (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) & s('type.num.num',sK2) = s('type.num.num',sK7) & s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) & s('type.ffi.ffi_state'(sK0),sK5) = s('type.ffi.ffi_state'(sK0),sK10))) & (s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | s('type.num.num',sK2) != s('type.num.num',sK7) | s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK8) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | s('type.ffi.ffi_state'(sK0),sK5) != s('type.ffi.ffi_state'(sK0),sK10))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f110])).
fof(f110,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) = s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) | (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) & s('type.num.num',X2) = s('type.num.num',X7) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) & s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),X10))) & (s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) != s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) | s('type.num.num',X2) != s('type.num.num',X7) | s('type.option.option'('type.num.num'),X3) != s('type.option.option'('type.num.num'),X8) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) | s('type.ffi.ffi_state'(X0),X5) != s('type.ffi.ffi_state'(X0),X10)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) = s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) | (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) & s('type.num.num',X2) = s('type.num.num',X7) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) & s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),X10))) & (s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) != s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) | (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) | s('type.num.num',X2) != s('type.num.num',X7) | s('type.option.option'('type.num.num'),X3) != s('type.option.option'('type.num.num'),X8) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) | s('type.ffi.ffi_state'(X0),X5) != s('type.ffi.ffi_state'(X0),X10))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) = s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) <~> (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) & s('type.num.num',X2) = s('type.num.num',X7) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) & s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),X10)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) = s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) <=> (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) & s('type.num.num',X2) = s('type.num.num',X7) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) & s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),X10)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1),s('type.num.num',X2),s('type.option.option'('type.num.num'),X3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4),s('type.ffi.ffi_state'(X0),X5))) = s('type.bviSem.state'(X0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.option.option'('type.num.num'),X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9),s('type.ffi.ffi_state'(X0),X10))) <=> (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) & s('type.num.num',X2) = s('type.num.num',X7) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X8) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X9) & s('type.ffi.ffi_state'(X0),X5) = s('type.ffi.ffi_state'(X0),X10)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.option.option'('type.num.num'),X9),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X10),s('type.ffi.ffi_state'(X6),X11))) = s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X13),s('type.option.option'('type.num.num'),X14),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X15),s('type.ffi.ffi_state'(X6),X16))) <=> (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12) & s('type.num.num',X8) = s('type.num.num',X13) & s('type.option.option'('type.num.num'),X9) = s('type.option.option'('type.num.num'),X14) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X10) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X15) & s('type.ffi.ffi_state'(X6),X11) = s('type.ffi.ffi_state'(X6),X16)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.option.option'('type.num.num'),X9),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X10),s('type.ffi.ffi_state'(X6),X11))) = s('type.bviSem.state'(X6),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X13),s('type.option.option'('type.num.num'),X14),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X15),s('type.ffi.ffi_state'(X6),X16))) <=> (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12) & s('type.num.num',X8) = s('type.num.num',X13) & s('type.option.option'('type.num.num'),X9) = s('type.option.option'('type.num.num'),X14) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X10) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),X15) & s('type.ffi.ffi_state'(X6),X11) = s('type.ffi.ffi_state'(X6),X16)))),
  file('thm.bviSem.state_|49||49|','thm.bviSem.state_|49||49|')).
fof(f124,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
fof(f123,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK8) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
fof(f183,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK7) | $spl4),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK7) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f122,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK7) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
fof(f181,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | $spl2),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f121,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) = s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
fof(f13106,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f13105])).
fof(f13105,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f13104,f181])).
fof(f13104,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) | ($spl1 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f13103,f183])).
fof(f13103,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) | ($spl1 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f13102,f185])).
fof(f13102,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f13101,f187])).
fof(f13101,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK5))) | ($spl1 | $spl10)),
  inference(forward_demodulation,[],[f180,f189])).
fof(f180,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) | $spl1),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f182,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | $spl3),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f184,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK7) | $spl5),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK7) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f186,plain,(
  s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK8) | $spl7),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK8) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f188,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | $spl9),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f190,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) != s('type.ffi.ffi_state'(sK0),sK10) | $spl11),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) != s('type.ffi.ffi_state'(sK0),sK10) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f126,plain,(
  s('type.ffi.ffi_state'(sK0),sK5) != s('type.ffi.ffi_state'(sK0),sK10) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9) | s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK8) | s('type.num.num',sK2) != s('type.num.num',sK7) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6) | s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK1),s('type.num.num',sK2),s('type.option.option'('type.num.num'),sK3),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK4),s('type.ffi.ffi_state'(sK0),sK5))) != s('type.bviSem.state'(sK0),'const.bviSem.recordtype|46|state_5'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK6),s('type.num.num',sK7),s('type.option.option'('type.num.num'),sK8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvi.exp')),sK9),s('type.ffi.ffi_state'(sK0),sK10)))),
  inference(cnf_transformation,[],[f111])).
% SZS output end Proof for thm.bviSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 143
% Passive clauses: 2332
% Generated clauses: 14111
% Final active clauses: 112
% Final passive clauses: 1999
% Input formulas: 37
% Initial clauses: 58
% 
% Duplicate literals: 430
% Trivial inequalities: 232
% Fw subsumption resolutions: 1527
% Fw demodulations: 2143
% Bw demodulations: 18
% 
% Simple tautologies: 203
% Equational tautologies: 41
% Forward subsumptions: 6891
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 483
% Factoring: 3
% Forward superposition: 5673
% Backward superposition: 3243
% Self superposition: 131
% Equality factoring: 76
% Equality resolution: 24
% 
% Split clauses: 164
% Split components: 332
% SAT solver clauses: 438
% SAT solver unit clauses: 63
% SAT solver binary clauses: 167
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 4
% 
% Sat splits: 164
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 435
% TWLsolver calls for satisfiability: 106
% 
% Memory used [KB]: 8315
% Time elapsed: 0.452 s
% ------------------------------
% ------------------------------
% Success in time 0.46 s
