Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bviSem
% (32156)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 167
% Passive clauses: 3259
% Generated clauses: 22636
% Final active clauses: 112
% Final passive clauses: 2873
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 820
% Trivial inequalities: 110
% Fw subsumption resolutions: 1867
% Fw demodulations: 3631
% Bw demodulations: 48
% 
% Simple tautologies: 143
% Equational tautologies: 49
% Forward subsumptions: 12213
% 
% Binary resolution: 496
% Factoring: 2
% Forward superposition: 8091
% Backward superposition: 7068
% Self superposition: 227
% Equality factoring: 96
% Equality resolution: 11
% 
% Split clauses: 241
% Split components: 482
% SAT solver clauses: 711
% SAT solver unit clauses: 34
% SAT solver binary clauses: 133
% 
% Sat splits: 241
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 711
% TWLsolver calls for satisfiability: 150
% 
% Memory used [KB]: 20852
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bviSem
% (32169)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 1159
% Generated clauses: 13164
% Final active clauses: 91
% Final passive clauses: 808
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 369
% Trivial inequalities: 19
% Fw subsumption resolutions: 692
% Fw demodulations: 1494
% Bw demodulations: 253
% Condensations: 843
% 
% Simple tautologies: 48
% Equational tautologies: 36
% Forward subsumptions: 6668
% 
% Binary resolution: 335
% Factoring: 1
% Forward superposition: 4387
% Backward superposition: 3403
% Self superposition: 169
% Equality factoring: 60
% Equality resolution: 1
% 
% Split clauses: 36
% Split components: 72
% SAT solver clauses: 1789
% SAT solver unit clauses: 131
% SAT solver binary clauses: 44
% 
% Sat splits: 36
% Sat splitting refutations: 31
% 
% TWLsolver clauses: 1789
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bviSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bviSem
% SZS output start Proof for thm.bviSem
fof(f15936,plain,(
  $false),
  inference(subsumption_resolution,[],[f15935,f215])).
fof(f215,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','thm.bool.TRUTH')).
fof(f15935,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f15933,f4013])).
fof(f4013,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f4012])).
fof(f4012,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f3940,f299])).
fof(f299,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f217,f218])).
fof(f218,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bviSem.state_fupdfupds_comp|split|9',aHL_BOOL_CASES)).
fof(f217,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bviSem.state_fupdfupds_comp|split|9',aHL_FALSITY)).
fof(f3940,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f3383,f297])).
fof(f297,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f218])).
fof(f3383,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3283,f1312])).
fof(f1312,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1265,f268])).
fof(f268,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X8,X17] : ! [X12,X11,X18] : s(X16,'const.combin.o_3'(s(fun(X8,X16),X12),s(fun(X17,X8),X11),s(X17,X18))) = s(X16,happ(s(fun(X8,X16),X12),s(X8,happ(s(fun(X17,X8),X11),s(X17,X18)))))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','thm.combin.o_THM')).
fof(f1265,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f219,f265])).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X44,X45,X46] : (p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X44,fun(X44,bool)),'=_0'),s(X44,X45))),s(X44,X46)))) <=> s(X44,X45) = s(X44,X46))),
  file('thm.bviSem.state_fupdfupds_comp|split|9',arity_0_HL_EQ)).
fof(f219,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X34] : ! [X36] : s(X34,X36) = s(X34,happ(s(fun(X34,X34),'_i_0_0'),s(X34,X36)))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','_i0_ax')).
fof(f3283,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,happ(s(fun(X73,X75),X76),s(X73,X74))) = s(X75,'const.combin.o_3'(s(fun(X73,X75),X76),s(fun(X73,X73),'_i_0_0'),s(X73,X74)))) )),
  inference(superposition,[],[f268,f219])).
fof(f15933,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.bviSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.bviSem.state'(sK0)),fun(fun(sK2,'type.bviSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))))),s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))))))),
  inference(backward_demodulation,[],[f15883,f4906])).
fof(f4906,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.bviSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.bviSem.state'(sK0)),fun(fun(sK2,'type.bviSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))))),s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK1),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.bviSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))))))))),
  inference(unit_resulting_resolution,[],[f214,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f214,plain,(
  s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))) != s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK1),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.bviSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4))) != s(fun(sK2,'type.bviSem.state'(sK0)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK1),'type.bviSem.state'(sK0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.bviSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bviSem.state'(sK3),'type.bviSem.state'(sK1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.bviSem.state'(sK3)),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.bviSem.state'(X3)),X4))) != s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.bviSem.state'(X1)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.bviSem.state'(X3)),X4)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.bviSem.state'(X3)),X4))) = s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.bviSem.state'(X1)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.bviSem.state'(X3)),X4)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.bviSem.state'(X3)),X4))) = s(fun(X2,'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.bviSem.state'(X1)),'const.combin.o_2'(s(fun('type.bviSem.state'(X3),'type.bviSem.state'(X1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.bviSem.state'(X3)),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.bviSem.state'(X6)),'const.combin.o_2'(s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.bviSem.state'(X9)),X10))) = s(fun(X8,'type.bviSem.state'(X6)),'const.combin.o_2'(s(fun('type.bviSem.state'(X7),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.bviSem.state'(X7)),'const.combin.o_2'(s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X7)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.bviSem.state'(X9)),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.bviSem.state'(X6)),'const.combin.o_2'(s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.bviSem.state'(X9)),X10))) = s(fun(X8,'type.bviSem.state'(X6)),'const.combin.o_2'(s(fun('type.bviSem.state'(X7),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.bviSem.state'(X7)),'const.combin.o_2'(s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X7)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.bviSem.state'(X9)),X10)))))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','thm.bviSem.state_fupdfupds_comp|split|9')).
fof(f15883,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X5,'type.bviSem.state'(X1)),'const.combin.o_2'(s(fun('type.bviSem.state'(X0),'type.bviSem.state'(X1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))))),s(fun(X5,'type.bviSem.state'(X0)),X6))) = s(fun(X5,'type.bviSem.state'(X1)),'const.combin.o_2'(s(fun('type.bviSem.state'(X2),'type.bviSem.state'(X1)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3))),s(fun(X5,'type.bviSem.state'(X2)),'const.combin.o_2'(s(fun('type.bviSem.state'(X0),'type.bviSem.state'(X2)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))),s(fun(X5,'type.bviSem.state'(X0)),X6)))))) )),
  inference(superposition,[],[f272,f267])).
fof(f267,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X2),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X2)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X2),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X2)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X0)),'const.combin.o_2'(s(fun('type.bviSem.state'(X2),'type.bviSem.state'(X0)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.bviSem.state'(X1),'type.bviSem.state'(X2)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X9,X7] : ! [X11,X12] : s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))) = s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X6)),'const.combin.o_2'(s(fun('type.bviSem.state'(X7),'type.bviSem.state'(X6)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun('type.bviSem.state'(X9),'type.bviSem.state'(X7)),'const.bviSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11)))))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','thm.bviSem.state_fupdfupds_comp|split|8')).
fof(f272,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X16,X8,X22,X17] : ! [X12,X11,X10] : s(fun(X22,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X22,X8),'const.combin.o_2'(s(fun(X17,X8),X11),s(fun(X22,X17),X10))))) = s(fun(X22,X16),'const.combin.o_2'(s(fun(X17,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X17,X8),X11))),s(fun(X22,X17),X10)))),
  file('thm.bviSem.state_fupdfupds_comp|split|9','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.bviSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 149
% Passive clauses: 1573
% Generated clauses: 17307
% Final active clauses: 83
% Final passive clauses: 906
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 491
% Trivial inequalities: 42
% Fw subsumption resolutions: 539
% Fw demodulations: 3104
% Bw demodulations: 507
% Condensations: 804
% Global subsumptions: 350
% 
% Simple tautologies: 120
% Equational tautologies: 131
% Forward subsumptions: 9706
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 20
% 
% Unit resulting resolution: 689
% Factoring: 1
% Forward superposition: 4881
% Backward superposition: 3823
% Self superposition: 177
% Equality factoring: 68
% Equality resolution: 2
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 2351
% SAT solver unit clauses: 785
% SAT solver binary clauses: 797
% 
% Sat splits: 27
% Sat splitting refutations: 372
% 
% Memory used [KB]: 24178
% Time elapsed: 0.892 s
% ------------------------------
% ------------------------------
% Success in time 2.26 s
