Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSysProps
% (85779)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 207
% Passive clauses: 3800
% Generated clauses: 31313
% Final active clauses: 182
% Final passive clauses: 3234
% Input formulas: 46
% Initial clauses: 77
% 
% Duplicate literals: 1149
% Trivial inequalities: 510
% Fw subsumption resolutions: 4377
% Fw demodulations: 4617
% Bw demodulations: 24
% 
% Simple tautologies: 560
% Equational tautologies: 80
% Forward subsumptions: 15862
% 
% Binary resolution: 193
% Factoring: 1
% Forward superposition: 10298
% Backward superposition: 9506
% Self superposition: 261
% Equality factoring: 172
% Equality resolution: 76
% 
% Split clauses: 273
% Split components: 550
% SAT solver clauses: 701
% SAT solver unit clauses: 26
% SAT solver binary clauses: 128
% 
% Sat splits: 273
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 699
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 17398
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.typeSysProps
% (85803)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 800
% Generated clauses: 17023
% Final active clauses: 167
% Final passive clauses: 603
% Input formulas: 46
% Initial clauses: 77
% 
% Duplicate literals: 461
% Trivial inequalities: 258
% Fw subsumption resolutions: 1083
% Fw demodulations: 2946
% Bw demodulations: 28
% Condensations: 738
% Global subsumptions: 5
% 
% Simple tautologies: 208
% Equational tautologies: 31
% Forward subsumptions: 8956
% 
% Binary resolution: 147
% Factoring: 1
% Forward superposition: 5155
% Backward superposition: 4738
% Self superposition: 86
% Equality factoring: 38
% Equality resolution: 49
% 
% Split clauses: 155
% Split components: 310
% SAT solver clauses: 1561
% SAT solver unit clauses: 105
% SAT solver binary clauses: 44
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 2
% 
% Sat splits: 155
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1540
% TWLsolver calls for satisfiability: 398
% 
% Memory used [KB]: 15735
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.typeSysProps
% (85828)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 557
% Passive clauses: 4921
% Generated clauses: 316334
% Final active clauses: 154
% Final passive clauses: 1017
% Input formulas: 46
% Initial clauses: 77
% Discarded non-redundant clauses: 256
% 
% Duplicate literals: 2849
% Trivial inequalities: 679
% Fw subsumption resolutions: 5000
% Fw demodulations: 185162
% Bw demodulations: 2004
% Condensations: 4229
% Global subsumptions: 299
% 
% Simple tautologies: 618
% Equational tautologies: 425
% Forward subsumptions: 64213
% Fw demodulations to eq. taut.: 46496
% Bw demodulations to eq. taut.: 119
% 
% Unit resulting resolution: 66296
% Factoring: 1
% Forward superposition: 25813
% Backward superposition: 20221
% Self superposition: 495
% Equality factoring: 140
% Equality resolution: 139
% 
% Split clauses: 191
% Split components: 389
% SAT solver clauses: 5919
% SAT solver unit clauses: 856
% SAT solver binary clauses: 1521
% 
% Sat splits: 191
% Sat splitting refutations: 315
% 
% Memory used [KB]: 58591
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.typeSysProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSysProps
% SZS output start Proof for thm.typeSysProps
fof(f15921,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f182,f205,f193,f194,f195,f196,f197,f198,f199,f200,f201,f202,f235,f239,f228,f229,f366,f391,f392,f393,f407,f598,f599,f600,f328,f327,f338,f336,f619,f644,f925,f926,f432,f1071,f1116,f1118,f1136,f1138,f641,f642,f643,f635,f636,f637,f638,f639,f640,f1267,f1300,f1301,f1419,f1420,f1421,f1507,f1508,f1676,f1677,f1875,f1913,f1975,f1991,f1992,f1993,f1994,f1995,f1996,f2000,f2035,f2036,f2151,f2152,f2495,f2497,f1999,f1998,f1997,f2584,f2585,f2586,f2587,f2590,f2591,f2592,f2708,f2709,f2887,f2889,f2593,f2589,f2588,f3020,f3083,f3084,f3392,f3393,f3437,f1107,f3482,f3483,f3694,f1108,f3732,f3741,f3743,f3986,f3988,f4081,f4082,f4083,f4084,f4087,f4088,f4089,f4202,f4203,f4205,f4090,f4086,f4085,f4219,f4218,f2897,f2896,f2935,f2934,f4215,f4214,f1967,f1966,f1466,f1465,f4301,f4302,f4303,f4304,f4305,f4306,f4310,f211,f217,f4420,f4421,f4672,f4673,f2496,f4852,f4853,f4908,f4909,f1117,f5261,f5262,f3742,f5542,f5543,f3987,f5707,f5737,f5735,f5710,f5809,f5811,f5810,f6786,f5738,f6814,f7343,f7346,f7354,f7355,f7345,f7431,f7432,f7433,f2898,f7577,f7574,f7573,f7627,f7791,f7792,f7793,f7954,f7794,f7795,f5739,f7796,f7953,f7952,f8106,f8113,f8100,f8115,f8143,f12063,f12065,f12523,f12525,f14237,f14246,f14247,f12983,f14282,f14284,f14285,f7921,f14288,f14291,f14292,f14293,f14283,f14294,f14286,f3515,f14287,f14330,f14333,f14336,f14338,f14341,f14343,f14346,f14348,f14350,f14359,f14360,f14361,f13782,f13781,f13780,f14365,f14366,f13768,f13767,f14399,f14401,f14402,f9770,f14403,f14404,f14407,f9765,f14409,f14410,f14412,f14413,f14414,f14416,f14415,f14417,f14418,f14420,f14421,f14422,f14423,f14424,f9659,f14425,f14426,f14427,f14428,f14429,f14430,f14432,f14433,f14435,f14436,f14438,f14440,f14442,f14443,f14444,f9302,f8114,f14447,f14400,f14448,f14449,f14450,f14451,f14453,f9246,f14454,f9268,f14455,f14456,f14411,f14457,f14459,f14460,f14461,f14462,f14463,f7320,f7316,f14486,f7301,f14362,f7299,f14487,f7298,f14488,f7297,f14489,f7295,f14490,f7293,f14491,f7292,f14492,f7256,f7254,f7252,f7250,f14496,f7231,f14497,f7229,f14498,f7227,f14367,f7225,f14499,f7224,f14500,f7223,f14501,f14502,f7219,f14503,f7218,f14504,f7169,f7168,f240,f7344,f14505,f7154,f7153,f14506,f14507,f14434,f7147,f14508,f7146,f14509,f14510,f7134,f7133,f7128,f14513,f14405,f14515,f7097,f7096,f14517,f7090,f14512,f7089,f7049,f7045,f14516,f7025,f14518,f7023,f14368,f14519,f7019,f14520,f7017,f14521,f14522,f14408,f6985,f6981,f14511,f6966,f14363,f6964,f14523,f6962,f14524,f6960,f14525,f6959,f14526,f14528,f14529,f14530,f14289,f6891,f14364,f14531,f14532,f14533,f14437,f6749,f14535,f14536,f14537,f14538,f14439,f6527,f14540,f14541,f14542,f14543,f14441,f6373,f14545,f6334,f14546,f14547,f14549,f14548,f6324,f14550,f2594,f203,f4309,f4308,f4307,f14358,f14357,f7555,f4060,f4059,f2563,f2562,f241,f236,f14639,f14640,f237,f14657,f14658,f208,f207,f14659,f14660,f209,f14663,f14664,f14665,f213,f215,f14817,f14818,f618,f206,f15461,f15462,f238,f2888,f15880,f15890,f15891])).
fof(f15891,plain,(
  $false | ($spl12 | $spl25)),
  inference(subsumption_resolution,[],[f15830,f212])).
fof(f212,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl12),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f15830,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl12 | $spl25)),
  inference(superposition,[],[f14698,f15646])).
fof(f15646,plain,(
  ( ! [X21] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X21))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X21)))))) ) | ($spl12 | $spl25)),
  inference(superposition,[],[f149,f14564])).
fof(f14564,plain,(
  ( ! [X4] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X4))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X4)))) ) | ($spl12 | $spl25)),
  inference(resolution,[],[f2706,f8834])).
fof(f8834,plain,(
  ( ! [X28,X29] : (p(s(bool,sK2(X28,X29))) | s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X29),s('type.typeSystem.tenv_val',X28))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X29),s('type.typeSystem.tenv_val',X28)))) )),
  inference(forward_demodulation,[],[f8743,f164])).
fof(f164,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X12) = s(X14,'const.bool.COND_3'(s(bool,'F_0'),s(X14,X11),s(X14,X12)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.COND_CLAUSES|split|1')).
fof(f8743,plain,(
  ( ! [X28,X29] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X29),s('type.typeSystem.tenv_val',X28))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,'F_0'),s('type.typeSystem.tenv_val',X28),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X29),s('type.typeSystem.tenv_val',X28))))) | p(s(bool,sK2(X28,X29)))) )),
  inference(superposition,[],[f138,f128])).
fof(f128,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f138,plain,(
  ( ! [X0,X1] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,sK2(X1,X0)),s('type.typeSystem.tenv_val',X1),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1)))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((~p(s(bool,sK2(X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK2(X1,X0)))) & s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,sK2(X1,X0)),s('type.typeSystem.tenv_val',X1),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ? [X2] : ((~p(s(bool,X2)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X2))) & s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,X2),s('type.typeSystem.tenv_val',X1),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ? [X2] : (((~p(s(bool,X2)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X2)))) & s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,X2),s('type.typeSystem.tenv_val',X1),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,X2),s('type.typeSystem.tenv_val',X1),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X29] : ? [X30] : ((p(s(bool,X30)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) & s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',X29))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,X30),s('type.typeSystem.tenv_val',X29),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',X29))))))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSystem.bind_tvar_def')).
fof(f2706,plain,(
  ( ! [X10] : (~p(s(bool,sK2(X10,sK0)))) ) | ($spl12 | $spl25)),
  inference(subsumption_resolution,[],[f2682,f922])).
fof(f922,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl25),
  inference(trivial_inequality_removal,[],[f921])).
fof(f921,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl25),
  inference(superposition,[],[f391,f135])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3',aHL_BOOL_CASES)).
fof(f2682,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X10,sK0)))) ) | $spl12),
  inference(superposition,[],[f2596,f136])).
fof(f136,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,sK2(X1,X0)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f2596,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl12),
  inference(backward_demodulation,[],[f212,f114])).
fof(f114,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f84])).
fof(f84,plain,(
  ? [X0,X1] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1)))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X1))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X7))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X7))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',X7)))))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSysProps.bind_tvar_rewrites|split|3')).
fof(f149,plain,(
  ( ! [X0,X1] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X1))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X1))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9,X16] : s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X16))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X9),s('type.typeSystem.tenv_val',X16)))))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSoundInvariants.tenv_val_ok_def|split|1')).
fof(f14698,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl12),
  inference(backward_demodulation,[],[f212,f114])).
fof(f15890,plain,(
  $false | ($spl12 | $spl14 | $spl25)),
  inference(subsumption_resolution,[],[f15889,f2597])).
fof(f2597,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl12 | $spl14)),
  inference(backward_demodulation,[],[f212,f1905])).
fof(f1905,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl14),
  inference(forward_demodulation,[],[f114,f215])).
fof(f15889,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl12 | $spl14 | $spl25)),
  inference(forward_demodulation,[],[f15829,f212])).
fof(f15829,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl12 | $spl14 | $spl25)),
  inference(superposition,[],[f215,f15646])).
fof(f15880,plain,(
  $false | ($spl12 | $spl14 | $spl25)),
  inference(subsumption_resolution,[],[f15879,f2597])).
fof(f15879,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl12 | $spl14 | $spl25)),
  inference(forward_demodulation,[],[f15818,f212])).
fof(f15818,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl12 | $spl14 | $spl25)),
  inference(superposition,[],[f15646,f215])).
fof(f2888,plain,(
  ( ! [X14,X12,X15] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X14),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X12))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X12))))) ) | $spl38),
  inference(cnf_transformation,[],[f2888_D])).
fof(f2888_D,plain,(
  ( ! [X14,X12,X15] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X14),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X12))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X12))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f238,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f15462,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,X23) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X25),s('type.typeSystem.tenv_val',X22))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22))))) ) | $spl14),
  inference(subsumption_resolution,[],[f15441,f127])).
fof(f127,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f15441,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,X23) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X25),s('type.typeSystem.tenv_val',X22))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22)))) | p(s(bool,X23))) ) | $spl14),
  inference(superposition,[],[f14624,f120])).
fof(f120,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f14624,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl14),
  inference(superposition,[],[f1905,f116])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val',X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : ((s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val',X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.typeSystem.tenv_val',X1) != s('type.typeSystem.tenv_val',X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : ((s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val',X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.typeSystem.tenv_val',X1) != s('type.typeSystem.tenv_val',X3)) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3)))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val',X3)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X21,X22,X23,X24] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X21),s('type.typeSystem.tenv_val',X22))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X24))) <=> (s('type.num.num',X21) = s('type.num.num',X23) & s('type.typeSystem.tenv_val',X22) = s('type.typeSystem.tenv_val',X24)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSystem.tenv_val_|49||49||split|0')).
fof(f15461,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X21),s('type.typeSystem.tenv_val',X18))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X18))))) ) | $spl14),
  inference(subsumption_resolution,[],[f15440,f127])).
fof(f15440,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') != s(bool,X19) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X21),s('type.typeSystem.tenv_val',X18))) | p(s(bool,X19)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X18))))) ) | $spl14),
  inference(superposition,[],[f14624,f120])).
fof(f206,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | $spl5),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f618,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl28),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f14818,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14806,f131])).
fof(f131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f14806,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(superposition,[],[f14698,f123])).
fof(f123,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f86])).
fof(f14817,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14805,f131])).
fof(f14805,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f14698,f123])).
fof(f215,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl14),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f213,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl13),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f14665,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f14664,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f134,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f14663,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f132,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f209,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X2))) ) | $spl8),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X2))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f14660,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f14659,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f207,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | $spl4),
  inference(cnf_transformation,[],[f207_D])).
fof(f207_D,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f208,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0))) ) | $spl6),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f14658,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f14657,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f237,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl18),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f14640,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(subsumption_resolution,[],[f14627,f127])).
fof(f14627,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f1905,f120])).
fof(f14639,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(subsumption_resolution,[],[f14626,f127])).
fof(f14626,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(superposition,[],[f1905,f120])).
fof(f236,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl19),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f241,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl23),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2562,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(superposition,[],[f1905,f132])).
fof(f2563,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(superposition,[],[f1905,f134])).
fof(f4059,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(superposition,[],[f1905,f132])).
fof(f4060,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(superposition,[],[f1905,f134])).
fof(f7555,plain,(
  ( ! [X19] : (p(s(bool,X19))) ) | $spl64),
  inference(cnf_transformation,[],[f7555_D])).
fof(f7555_D,plain,(
  ( ! [X19] : (p(s(bool,X19))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14357,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4033])).
fof(f4033,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(superposition,[],[f2567,f120])).
fof(f2567,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(trivial_inequality_removal,[],[f2561])).
fof(f2561,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl14),
  inference(superposition,[],[f1905,f130])).
fof(f130,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.EQ_CLAUSES|split|3')).
fof(f14358,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4034])).
fof(f4034,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f2567,f120])).
fof(f4307,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f4308,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f4309,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f203,plain,(
  ( ! [X10] : (~p(s(bool,sK2(X10,sK0)))) ) | $spl0),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  ( ! [X10] : (~p(s(bool,sK2(X10,sK0)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2594,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,sK2(X10,sK0)))) ) | $spl25),
  inference(forward_demodulation,[],[f2573,f922])).
fof(f2573,plain,(
  ( ! [X10] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X10,sK0)))) )),
  inference(superposition,[],[f114,f136])).
fof(f14550,plain,(
  ( ! [X3] : (~p(s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) | $spl194),
  inference(cnf_transformation,[],[f14550_D])).
fof(f14550_D,plain,(
  ( ! [X3] : (~p(s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f6324,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) | $spl51),
  inference(superposition,[],[f5963,f5815])).
fof(f5815,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,sK2(X4,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(resolution,[],[f5768,f128])).
fof(f5768,plain,(
  ( ! [X12] : (~p(s(bool,sK2(X12,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))) ) | $spl51),
  inference(trivial_inequality_removal,[],[f5765])).
fof(f5765,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ~p(s(bool,sK2(X12,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))) ) | $spl51),
  inference(superposition,[],[f5737,f136])).
fof(f5963,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,sK2(X18,X17)) | ~p(s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17)))))) ) | $spl51),
  inference(resolution,[],[f5822,f131])).
fof(f5822,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK2(X2,X0))) | ~p(s(bool,sK2(X1,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) ) | $spl51),
  inference(superposition,[],[f5768,f136])).
fof(f14548,plain,(
  ( ! [X24,X25] : (~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) ) | $spl192),
  inference(cnf_transformation,[],[f14548_D])).
fof(f14548_D,plain,(
  ( ! [X24,X25] : (~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f14549,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) ) | $spl51),
  inference(subsumption_resolution,[],[f6329,f5822])).
fof(f6329,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))))) | p(s(bool,sK2(X20,X21)))) ) | $spl51),
  inference(superposition,[],[f5963,f128])).
fof(f14547,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) ) | $spl51),
  inference(subsumption_resolution,[],[f6330,f5822])).
fof(f6330,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24))))) | p(s(bool,sK2(X23,X24)))) ) | $spl51),
  inference(superposition,[],[f5963,f130])).
fof(f14546,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,sK2(X37,'const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,sK2(X35,X36))) ) | $spl190),
  inference(cnf_transformation,[],[f14546_D])).
fof(f14546_D,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,sK2(X37,'const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,sK2(X35,X36))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f6334,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X37,'const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,sK2(X35,X36))) ) | $spl51),
  inference(superposition,[],[f5963,f135])).
fof(f14545,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | $spl189),
  inference(cnf_transformation,[],[f14545_D])).
fof(f14545_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f6373,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | ($spl51 | $spl54)),
  inference(superposition,[],[f5988,f5815])).
fof(f5988,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,sK2(X12,X11)) | s('type.num.num',X11) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11)))) ) | $spl54),
  inference(resolution,[],[f5739,f131])).
fof(f14441,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10)))) ) | $spl124),
  inference(cnf_transformation,[],[f14441_D])).
fof(f14441_D,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f14543,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X15) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15)))) ) | $spl54),
  inference(subsumption_resolution,[],[f6378,f5739])).
fof(f6378,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X15) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15))) | p(s(bool,sK2(X14,X15)))) ) | $spl54),
  inference(superposition,[],[f5988,f128])).
fof(f14542,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17)))) ) | $spl54),
  inference(subsumption_resolution,[],[f6379,f5739])).
fof(f6379,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17))) | p(s(bool,sK2(X16,X17)))) ) | $spl54),
  inference(superposition,[],[f5988,f130])).
fof(f14541,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X25) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))) ) | $spl54),
  inference(subsumption_resolution,[],[f6383,f5988])).
fof(f6383,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X25) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))) | s(bool,'T_0') = s(bool,sK2(X24,X25))) ) | $spl54),
  inference(superposition,[],[f5988,f135])).
fof(f14540,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | $spl187),
  inference(cnf_transformation,[],[f14540_D])).
fof(f14540_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f6527,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | $spl51),
  inference(superposition,[],[f6110,f5815])).
fof(f6110,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,sK2(X12,X11)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(resolution,[],[f6096,f131])).
fof(f6096,plain,(
  ( ! [X14,X15] : (~p(s(bool,sK2(X15,X14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(trivial_inequality_removal,[],[f6078])).
fof(f6078,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK2(X15,X14)))) ) | $spl51),
  inference(superposition,[],[f5759,f136])).
fof(f5759,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(superposition,[],[f5737,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9,X10] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))) <=> s('type.num.num',X8) = s('type.num.num',X9))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.arithmetic.EQ_MONO_ADD_EQ')).
fof(f14439,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f14439_D])).
fof(f14439_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f14538,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f6532,f6096])).
fof(f6532,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK2(X14,X15)))) ) | $spl51),
  inference(superposition,[],[f6110,f128])).
fof(f14537,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f6533,f6096])).
fof(f6533,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK2(X16,X17)))) ) | $spl51),
  inference(superposition,[],[f6110,f130])).
fof(f14536,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f6537,f6110])).
fof(f6537,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK2(X24,X25))) ) | $spl51),
  inference(superposition,[],[f6110,f135])).
fof(f14535,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | $spl185),
  inference(cnf_transformation,[],[f14535_D])).
fof(f14535_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f6749,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | $spl51),
  inference(superposition,[],[f6207,f5815])).
fof(f6207,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,sK2(X12,X11)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11)))) ) | $spl51),
  inference(resolution,[],[f6194,f131])).
fof(f6194,plain,(
  ( ! [X14,X15] : (~p(s(bool,sK2(X15,X14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X14)))) ) | $spl51),
  inference(trivial_inequality_removal,[],[f6176])).
fof(f6176,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X14))) | ~p(s(bool,sK2(X15,X14)))) ) | $spl51),
  inference(superposition,[],[f6011,f136])).
fof(f6011,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))) ) | $spl51),
  inference(resolution,[],[f5958,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (p(s(bool,sK2(X1,X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f99])).
fof(f5958,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK2(X0,'const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl51),
  inference(resolution,[],[f5822,f137])).
fof(f14437,plain,(
  ( ! [X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16)))) ) | $spl120),
  inference(cnf_transformation,[],[f14437_D])).
fof(f14437_D,plain,(
  ( ! [X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f14533,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15)))) ) | $spl51),
  inference(subsumption_resolution,[],[f6754,f6194])).
fof(f6754,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15))) | p(s(bool,sK2(X14,X15)))) ) | $spl51),
  inference(superposition,[],[f6207,f128])).
fof(f14532,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17)))) ) | $spl51),
  inference(subsumption_resolution,[],[f6755,f6194])).
fof(f6755,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17))) | p(s(bool,sK2(X16,X17)))) ) | $spl51),
  inference(superposition,[],[f6207,f130])).
fof(f14531,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))) ) | $spl51),
  inference(subsumption_resolution,[],[f6759,f6207])).
fof(f6759,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))) | s(bool,'T_0') = s(bool,sK2(X24,X25))) ) | $spl51),
  inference(superposition,[],[f6207,f135])).
fof(f14364,plain,(
  ( ! [X53] : (~p(s(bool,sK2(X53,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) | $spl96),
  inference(cnf_transformation,[],[f14364_D])).
fof(f14364_D,plain,(
  ( ! [X53] : (~p(s(bool,sK2(X53,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6891,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f5815])).
fof(f6830,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,sK2(X18,X17)) | ~p(s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X17)))))) ) | $spl58),
  inference(resolution,[],[f6814,f131])).
fof(f14289,plain,(
  ( ! [X24,X23] : (~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23)))))) ) | $spl88),
  inference(cnf_transformation,[],[f14289_D])).
fof(f14289_D,plain,(
  ( ! [X24,X23] : (~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f14530,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f6896,f6814])).
fof(f6896,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21))))) | p(s(bool,sK2(X20,X21)))) ) | $spl58),
  inference(superposition,[],[f6830,f128])).
fof(f14529,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f6897,f6814])).
fof(f6897,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) | p(s(bool,sK2(X23,X24)))) ) | $spl58),
  inference(superposition,[],[f6830,f130])).
fof(f14528,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X37,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f6901,f6830])).
fof(f6901,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X37,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,sK2(X35,X36))) ) | $spl58),
  inference(superposition,[],[f6830,f135])).
fof(f14526,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,sK2(X66,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))))))) | s(bool,'F_0') = s(bool,sK2(X67,X65))) ) | $spl182),
  inference(cnf_transformation,[],[f14526_D])).
fof(f14526_D,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,sK2(X66,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))))))) | s(bool,'F_0') = s(bool,sK2(X67,X65))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6959,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X66,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X65))))))) | s(bool,'F_0') = s(bool,sK2(X67,X65))) ) | $spl51),
  inference(superposition,[],[f5963,f6231])).
fof(f6231,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') = s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11)))) | s(bool,'F_0') = s(bool,sK2(X12,X11))) ) | $spl51),
  inference(resolution,[],[f5844,f128])).
fof(f5844,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK2(X2,X0))) | s(bool,'F_0') = s(bool,sK2(X1,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl51),
  inference(superposition,[],[f5815,f136])).
fof(f14525,plain,(
  ( ! [X70,X69] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))))) | s(bool,'F_0') = s(bool,sK2(X70,X69))) ) | $spl180),
  inference(cnf_transformation,[],[f14525_D])).
fof(f14525_D,plain,(
  ( ! [X70,X69] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))))) | s(bool,'F_0') = s(bool,sK2(X70,X69))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6960,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))))) | s(bool,'F_0') = s(bool,sK2(X70,X69))) ) | ($spl51 | $spl54)),
  inference(superposition,[],[f5988,f6231])).
fof(f14524,plain,(
  ( ! [X76,X75] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X75),s('type.num.num',X75))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X76,X75))) ) | $spl178),
  inference(cnf_transformation,[],[f14524_D])).
fof(f14524_D,plain,(
  ( ! [X76,X75] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X75),s('type.num.num',X75))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X76,X75))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6962,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X75),s('type.num.num',X75))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X76,X75))) ) | $spl51),
  inference(superposition,[],[f6110,f6231])).
fof(f14523,plain,(
  ( ! [X81,X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))))) | s(bool,'F_0') = s(bool,sK2(X82,X81))) ) | $spl176),
  inference(cnf_transformation,[],[f14523_D])).
fof(f14523_D,plain,(
  ( ! [X81,X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))))) | s(bool,'F_0') = s(bool,sK2(X82,X81))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6964,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))))) | s(bool,'F_0') = s(bool,sK2(X82,X81))) ) | $spl51),
  inference(superposition,[],[f6207,f6231])).
fof(f14363,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,sK2(X56,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X55))))))) | s(bool,'F_0') = s(bool,sK2(X57,X55))) ) | $spl94),
  inference(cnf_transformation,[],[f14363_D])).
fof(f14363_D,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,sK2(X56,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X55))))))) | s(bool,'F_0') = s(bool,sK2(X57,X55))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6966,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X89,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X88),s('type.num.num',X88))))))) | s(bool,'F_0') = s(bool,sK2(X90,X88))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6231])).
fof(f14511,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') = s(bool,sK2(X31,X30)) | p(s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X30)))))) ) | $spl164),
  inference(cnf_transformation,[],[f14511_D])).
fof(f14511_D,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') = s(bool,sK2(X31,X30)) | p(s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X30)))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6981,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X145,'const.arithmetic.+_2'(s('type.num.num',X146),s('type.num.num',X146))))) | s(bool,'F_0') = s(bool,sK2(X147,X146))) ) | $spl51),
  inference(superposition,[],[f131,f6231])).
fof(f6985,plain,(
  ( ! [X158,X159,X157] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X157,'const.arithmetic.+_2'(s('type.num.num',X158),s('type.num.num',X158))))) | s(bool,'F_0') = s(bool,sK2(X159,X158))) ) | $spl51),
  inference(superposition,[],[f133,f6231])).
fof(f133,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f14408,plain,(
  ( ! [X18,X16] : (s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) | $spl108),
  inference(cnf_transformation,[],[f14408_D])).
fof(f14408_D,plain,(
  ( ! [X18,X16] : (s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f14522,plain,(
  ( ! [X48,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X48,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7016,f8437])).
fof(f8437,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,sK2(X11,X12))) | s(bool,'F_0') = s(bool,sK2(X13,X12))) )),
  inference(resolution,[],[f8101,f128])).
fof(f8101,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,sK2(X20,X18))) | p(s(bool,sK2(X19,X18)))) )),
  inference(trivial_inequality_removal,[],[f8083])).
fof(f8083,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,sK2(X19,X18))) | ~p(s(bool,sK2(X20,X18)))) )),
  inference(superposition,[],[f137,f136])).
fof(f7016,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X47,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s(bool,'F_0') = s(bool,sK2(X48,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) ) | $spl51),
  inference(superposition,[],[f5963,f6231])).
fof(f14521,plain,(
  ( ! [X50,X51] : (s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))))) ) | $spl174),
  inference(cnf_transformation,[],[f14521_D])).
fof(f14521_D,plain,(
  ( ! [X50,X51] : (s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7017,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))))) ) | ($spl51 | $spl54)),
  inference(superposition,[],[f5988,f6231])).
fof(f14520,plain,(
  ( ! [X57,X56] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X57,'const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X56))))) ) | $spl172),
  inference(cnf_transformation,[],[f14520_D])).
fof(f14520_D,plain,(
  ( ! [X57,X56] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X57,'const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X56))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7019,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X57,'const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X56))))) ) | $spl51),
  inference(superposition,[],[f6110,f6231])).
fof(f14519,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62)))) ) | $spl51),
  inference(subsumption_resolution,[],[f7021,f8048])).
fof(f8048,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s(bool,'F_0') != s(bool,sK2(X8,X7))) )),
  inference(resolution,[],[f137,f127])).
fof(f7021,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62))) | s(bool,'F_0') = s(bool,sK2(X63,'const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62))))) ) | $spl51),
  inference(superposition,[],[f6207,f6231])).
fof(f14368,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,sK2(X2,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) | $spl100),
  inference(cnf_transformation,[],[f14368_D])).
fof(f14368_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,sK2(X2,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7023,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X70,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X69))))) | s(bool,'F_0') = s(bool,sK2(X71,'const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num',X69))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6231])).
fof(f14518,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) | $spl170),
  inference(cnf_transformation,[],[f14518_D])).
fof(f14518_D,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7025,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X76,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) | $spl51),
  inference(superposition,[],[f5817,f6231])).
fof(f5817,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(resolution,[],[f5768,f131])).
fof(f14516,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') = s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) | $spl168),
  inference(cnf_transformation,[],[f14516_D])).
fof(f14516_D,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') = s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7045,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X151,X152))) | s(bool,'F_0') = s(bool,sK2(X153,'const.arithmetic.+_2'(s('type.num.num',X152),s('type.num.num',X152))))) ) | $spl51),
  inference(superposition,[],[f131,f6231])).
fof(f7049,plain,(
  ( ! [X165,X163,X164] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X163,X164))) | s(bool,'F_0') = s(bool,sK2(X165,'const.arithmetic.+_2'(s('type.num.num',X164),s('type.num.num',X164))))) ) | $spl51),
  inference(superposition,[],[f133,f6231])).
fof(f7089,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) | $spl51),
  inference(superposition,[],[f6233,f5815])).
fof(f6233,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,sK2(X18,X17)) | s(bool,'F_0') = s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17))))) ) | $spl51),
  inference(resolution,[],[f5844,f131])).
fof(f14512,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') = s(bool,sK2(X8,X7)) | s(bool,'F_0') = s(bool,sK2(X9,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))))))) ) | $spl166),
  inference(cnf_transformation,[],[f14512_D])).
fof(f14512_D,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') = s(bool,sK2(X8,X7)) | s(bool,'F_0') = s(bool,sK2(X9,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7090,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,sK2(X7,X5))) ) | $spl51),
  inference(superposition,[],[f6233,f6231])).
fof(f14517,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7091,f9050])).
fof(f9050,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,sK2(X15,X14)) | s(bool,'F_0') = s(bool,sK2(X13,X14))) )),
  inference(resolution,[],[f8437,f127])).
fof(f7091,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,sK2(X11,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))))) ) | $spl51),
  inference(superposition,[],[f6233,f6231])).
fof(f7096,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X30,'const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X29)))) | p(s(bool,sK2(X28,X29)))) ) | $spl51),
  inference(superposition,[],[f6233,f128])).
fof(f7097,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) | $spl51),
  inference(superposition,[],[f6233,f130])).
fof(f14515,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X45,'const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7101,f6233])).
fof(f7101,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X45,'const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44)))) | s(bool,'T_0') = s(bool,sK2(X43,X44))) ) | $spl51),
  inference(superposition,[],[f6233,f135])).
fof(f14405,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') = s(bool,sK2(X26,X24))) ) | $spl106),
  inference(cnf_transformation,[],[f14405_D])).
fof(f14405_D,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') = s(bool,sK2(X26,X24))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f14513,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X4,X3))) ) | $spl51),
  inference(subsumption_resolution,[],[f7127,f9050])).
fof(f7127,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X4,X3)) | s(bool,'F_0') = s(bool,sK2(X5,X3))) ) | $spl51),
  inference(superposition,[],[f6256,f6231])).
fof(f6256,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11)))) | s(bool,'F_0') = s(bool,sK2(X12,X11))) ) | $spl51),
  inference(resolution,[],[f5930,f128])).
fof(f5930,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK2(X2,X0))) | s(bool,'T_0') != s(bool,sK2(X1,'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))) ) | $spl51),
  inference(superposition,[],[f5817,f136])).
fof(f7128,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X8,X7)) | s(bool,'F_0') = s(bool,sK2(X9,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))))))) ) | $spl51),
  inference(superposition,[],[f6256,f6231])).
fof(f7133,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X28,X27)) | p(s(bool,sK2(X26,'const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))))) ) | $spl51),
  inference(superposition,[],[f6256,f128])).
fof(f7134,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X31,X30)) | p(s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X30)))))) ) | $spl51),
  inference(superposition,[],[f6256,f130])).
fof(f14510,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X43,X42))) ) | $spl51),
  inference(subsumption_resolution,[],[f7138,f6256])).
fof(f7138,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X43,X42)) | s(bool,'T_0') = s(bool,sK2(X41,'const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X42))))) ) | $spl51),
  inference(superposition,[],[f6256,f135])).
fof(f14509,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) | $spl162),
  inference(cnf_transformation,[],[f14509_D])).
fof(f14509_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7146,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) ) | $spl51),
  inference(superposition,[],[f6258,f5815])).
fof(f6258,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,sK2(X18,X17)) | s(bool,'T_0') != s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X17))))) ) | $spl51),
  inference(resolution,[],[f5930,f131])).
fof(f14508,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,sK2(X7,X5))) ) | $spl160),
  inference(cnf_transformation,[],[f14508_D])).
fof(f14508_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,sK2(X7,X5))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f7147,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,sK2(X7,X5))) ) | $spl51),
  inference(superposition,[],[f6258,f6231])).
fof(f14434,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,sK2(X26,'const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) ) | $spl118),
  inference(cnf_transformation,[],[f14434_D])).
fof(f14434_D,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,sK2(X26,'const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f14507,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7148,f9083])).
fof(f9083,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,sK2(X13,X14)) | s(bool,'F_0') != s(bool,sK2(X15,X14))) )),
  inference(resolution,[],[f8439,f127])).
fof(f8439,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,sK2(X17,X18))) | s(bool,'T_0') != s(bool,sK2(X19,X18))) )),
  inference(resolution,[],[f8101,f131])).
fof(f7148,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,sK2(X11,'const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))))) ) | $spl51),
  inference(superposition,[],[f6258,f6231])).
fof(f14506,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) | $spl158),
  inference(cnf_transformation,[],[f14506_D])).
fof(f14506_D,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f7153,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X30,'const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X29)))) | p(s(bool,sK2(X28,X29)))) ) | $spl51),
  inference(superposition,[],[f6258,f128])).
fof(f7154,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X33,'const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))) | p(s(bool,sK2(X31,X32)))) ) | $spl51),
  inference(superposition,[],[f6258,f130])).
fof(f14505,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X45,'const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7158,f6258])).
fof(f7158,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X45,'const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44)))) | s(bool,'T_0') = s(bool,sK2(X43,X44))) ) | $spl51),
  inference(superposition,[],[f6258,f135])).
fof(f7344,plain,(
  ( ! [X21,X19,X18] : (s(bool,'F_0') = s(bool,sK2(X21,'const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X18,X19)))) ) | $spl60),
  inference(cnf_transformation,[],[f7344_D])).
fof(f7344_D,plain,(
  ( ! [X21,X19,X18] : (s(bool,'F_0') = s(bool,sK2(X21,'const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X18,X19)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl22),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7168,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X32,'const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X30,X31)))) ) | $spl51),
  inference(superposition,[],[f6476,f132])).
fof(f6476,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') = s(bool,sK2(X10,X11)) | s(bool,'F_0') = s(bool,sK2(X12,'const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(resolution,[],[f6470,f128])).
fof(f6470,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,sK2(X30,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X29,X28))) ) | $spl51),
  inference(trivial_inequality_removal,[],[f6467])).
fof(f6467,plain,(
  ( ! [X30,X28,X29] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s(bool,'F_0') = s(bool,sK2(X29,X28)) | ~p(s(bool,sK2(X30,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))))) ) | $spl51),
  inference(superposition,[],[f6108,f136])).
fof(f6108,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X8,X7))) ) | $spl51),
  inference(resolution,[],[f6096,f128])).
fof(f7169,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X35,'const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X33,X34)))) ) | $spl51),
  inference(superposition,[],[f6476,f134])).
fof(f14504,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,sK2(X47,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s(bool,'F_0') = s(bool,sK2(X48,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.num.|48|_0'))))) ) | $spl156),
  inference(cnf_transformation,[],[f14504_D])).
fof(f14504_D,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,sK2(X47,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s(bool,'F_0') = s(bool,sK2(X48,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f7218,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X47,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s(bool,'F_0') = s(bool,sK2(X48,'const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f5963,f6476])).
fof(f14503,plain,(
  ( ! [X50,X51] : (s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0'))))) ) | $spl154),
  inference(cnf_transformation,[],[f14503_D])).
fof(f14503_D,plain,(
  ( ! [X50,X51] : (s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7219,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X50) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s(bool,'F_0') = s(bool,sK2(X51,'const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.num.|48|_0'))))) ) | ($spl51 | $spl54)),
  inference(superposition,[],[f5988,f6476])).
fof(f14502,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f7221,f8048])).
fof(f7221,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X57,'const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f6110,f6476])).
fof(f14501,plain,(
  ( ! [X62,X63] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62))) | s(bool,'F_0') = s(bool,sK2(X63,'const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num','const.num.|48|_0'))))) ) | $spl152),
  inference(cnf_transformation,[],[f14501_D])).
fof(f14501_D,plain,(
  ( ! [X62,X63] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62))) | s(bool,'F_0') = s(bool,sK2(X63,'const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f7223,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X62))) | s(bool,'F_0') = s(bool,sK2(X63,'const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f6207,f6476])).
fof(f14500,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') = s(bool,sK2(X68,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num',X67)))) | s(bool,'F_0') = s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num','const.num.|48|_0'))))) ) | $spl150),
  inference(cnf_transformation,[],[f14500_D])).
fof(f14500_D,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') = s(bool,sK2(X68,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num',X67)))) | s(bool,'F_0') = s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f7224,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X68,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num',X67)))) | s(bool,'F_0') = s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f6233,f6476])).
fof(f14499,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,sK2(X72,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num',X71)))) | s(bool,'F_0') = s(bool,sK2(X73,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))))) ) | $spl148),
  inference(cnf_transformation,[],[f14499_D])).
fof(f14499_D,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,sK2(X72,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num',X71)))) | s(bool,'F_0') = s(bool,sK2(X73,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f7225,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X72,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num',X71)))) | s(bool,'F_0') = s(bool,sK2(X73,'const.arithmetic.+_2'(s('type.num.num',X71),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f6258,f6476])).
fof(f14367,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) | s(bool,'F_0') = s(bool,sK2(X7,'const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) ) | $spl98),
  inference(cnf_transformation,[],[f14367_D])).
fof(f14367_D,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) | s(bool,'F_0') = s(bool,sK2(X7,'const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7227,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X80,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X79))))) | s(bool,'F_0') = s(bool,sK2(X81,'const.arithmetic.+_2'(s('type.num.num',X79),s('type.num.num','const.num.|48|_0'))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6476])).
fof(f14498,plain,(
  ( ! [X86] : (s(bool,'F_0') = s(bool,sK2(X86,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))) ) | $spl146),
  inference(cnf_transformation,[],[f14498_D])).
fof(f14498_D,plain,(
  ( ! [X86] : (s(bool,'F_0') = s(bool,sK2(X86,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7229,plain,(
  ( ! [X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X86,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f5817,f6476])).
fof(f14497,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,sK2(X94,X93)) | s(bool,'F_0') = s(bool,sK2(X95,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num',X93))),s('type.num.num','const.num.|48|_0'))))) ) | $spl144),
  inference(cnf_transformation,[],[f14497_D])).
fof(f14497_D,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,sK2(X94,X93)) | s(bool,'F_0') = s(bool,sK2(X95,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num',X93))),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7231,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X94,X93)) | s(bool,'F_0') = s(bool,sK2(X95,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num',X93))),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f6256,f6476])).
fof(f14496,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,sK2(X179,X180))) | s(bool,'F_0') = s(bool,sK2(X181,'const.arithmetic.+_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0'))))) ) | $spl142),
  inference(cnf_transformation,[],[f14496_D])).
fof(f14496_D,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,sK2(X179,X180))) | s(bool,'F_0') = s(bool,sK2(X181,'const.arithmetic.+_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7250,plain,(
  ( ! [X167,X169,X168] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X167,X168))) | s(bool,'F_0') = s(bool,sK2(X169,'const.arithmetic.+_2'(s('type.num.num',X168),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f131,f6476])).
fof(f7252,plain,(
  ( ! [X175,X173,X174] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,sK2(X173,X174))) | s(bool,'F_0') = s(bool,sK2(X175,'const.arithmetic.+_2'(s('type.num.num',X174),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f132,f6476])).
fof(f7254,plain,(
  ( ! [X180,X181,X179] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X179,X180))) | s(bool,'F_0') = s(bool,sK2(X181,'const.arithmetic.+_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f133,f6476])).
fof(f7256,plain,(
  ( ! [X187,X185,X186] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,sK2(X185,X186))) | s(bool,'F_0') = s(bool,sK2(X187,'const.arithmetic.+_2'(s('type.num.num',X186),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f134,f6476])).
fof(f14492,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X70,X68))) ) | $spl140),
  inference(cnf_transformation,[],[f14492_D])).
fof(f14492_D,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X70,X68))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7292,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X69,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X70,X68))) ) | $spl51),
  inference(superposition,[],[f5963,f6476])).
fof(f14491,plain,(
  ( ! [X72,X73] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X73,X72))) ) | $spl138),
  inference(cnf_transformation,[],[f14491_D])).
fof(f14491_D,plain,(
  ( ! [X72,X73] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X73,X72))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7293,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X73,X72))) ) | ($spl51 | $spl54)),
  inference(superposition,[],[f5988,f6476])).
fof(f14490,plain,(
  ( ! [X78,X79] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X79,X78))) ) | $spl136),
  inference(cnf_transformation,[],[f14490_D])).
fof(f14490_D,plain,(
  ( ! [X78,X79] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X79,X78))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f7295,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X79,X78))) ) | $spl51),
  inference(superposition,[],[f6110,f6476])).
fof(f14489,plain,(
  ( ! [X85,X84] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X85,X84))) ) | $spl134),
  inference(cnf_transformation,[],[f14489_D])).
fof(f14489_D,plain,(
  ( ! [X85,X84] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X85,X84))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7297,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X85,X84))) ) | $spl51),
  inference(superposition,[],[f6207,f6476])).
fof(f14488,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') = s(bool,sK2(X90,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X91,X89))) ) | $spl132),
  inference(cnf_transformation,[],[f14488_D])).
fof(f14488_D,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') = s(bool,sK2(X90,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X91,X89))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7298,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X90,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X91,X89))) ) | $spl51),
  inference(superposition,[],[f6233,f6476])).
fof(f14487,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,sK2(X94,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X95,X93))) ) | $spl130),
  inference(cnf_transformation,[],[f14487_D])).
fof(f14487_D,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,sK2(X94,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X95,X93))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7299,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X94,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,sK2(X95,X93))) ) | $spl51),
  inference(superposition,[],[f6258,f6476])).
fof(f14362,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,sK2(X60,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X61,X59))) ) | $spl92),
  inference(cnf_transformation,[],[f14362_D])).
fof(f14362_D,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,sK2(X60,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X61,X59))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7301,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X102,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X101),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X103,X101))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6476])).
fof(f14486,plain,(
  ( ! [X171,X172,X170] : (p(s(bool,sK2(X170,'const.arithmetic.+_2'(s('type.num.num',X171),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X172,X171))) ) | $spl128),
  inference(cnf_transformation,[],[f14486_D])).
fof(f14486_D,plain,(
  ( ! [X171,X172,X170] : (p(s(bool,sK2(X170,'const.arithmetic.+_2'(s('type.num.num',X171),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X172,X171))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7316,plain,(
  ( ! [X158,X159,X160] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X158,'const.arithmetic.+_2'(s('type.num.num',X159),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X160,X159))) ) | $spl51),
  inference(superposition,[],[f131,f6476])).
fof(f7320,plain,(
  ( ! [X171,X172,X170] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X170,'const.arithmetic.+_2'(s('type.num.num',X171),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,sK2(X172,X171))) ) | $spl51),
  inference(superposition,[],[f133,f6476])).
fof(f14463,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X7) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))) ) | $spl54),
  inference(subsumption_resolution,[],[f9184,f5983])).
fof(f5983,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl54),
  inference(resolution,[],[f5739,f137])).
fof(f9184,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X7) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) | $spl54),
  inference(superposition,[],[f5988,f7670])).
fof(f7670,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,sK2(X7,X6)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) )),
  inference(resolution,[],[f136,f128])).
fof(f14462,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f9185,f5759])).
fof(f9185,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) | $spl51),
  inference(superposition,[],[f6110,f7670])).
fof(f14461,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11)))) ) | $spl51),
  inference(subsumption_resolution,[],[f9186,f6011])).
fof(f9186,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) ) | $spl51),
  inference(superposition,[],[f6207,f7670])).
fof(f14460,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9187,f6014])).
fof(f6014,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s(bool,'F_0') = s(bool,sK2(X8,'const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))))) ) | $spl51),
  inference(resolution,[],[f5958,f128])).
fof(f9187,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X16,'const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X15)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15)) ) | $spl51),
  inference(superposition,[],[f6233,f7670])).
fof(f14459,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X19,'const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X18))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9188,f6016])).
fof(f6016,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X11) | s(bool,'T_0') != s(bool,sK2(X12,'const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))))) ) | $spl51),
  inference(resolution,[],[f5958,f131])).
fof(f9188,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X19,'const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X18)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) ) | $spl51),
  inference(superposition,[],[f6258,f7670])).
fof(f14457,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X36,X35))) ) | $spl51),
  inference(subsumption_resolution,[],[f9194,f6278])).
fof(f6278,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) | s(bool,'F_0') = s(bool,sK2(X0,X1))) ) | $spl51),
  inference(resolution,[],[f5961,f137])).
fof(f5961,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))))) | s(bool,'F_0') = s(bool,sK2(X12,X11))) ) | $spl51),
  inference(resolution,[],[f5822,f128])).
fof(f9194,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X36,X35)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num',X35)))) ) | $spl51),
  inference(superposition,[],[f6256,f7670])).
fof(f14411,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) | $spl112),
  inference(cnf_transformation,[],[f14411_D])).
fof(f14411_D,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f14456,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X85)) )),
  inference(subsumption_resolution,[],[f9211,f136])).
fof(f9211,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X84,X85))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X85)) )),
  inference(superposition,[],[f131,f7670])).
fof(f14455,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X93)) )),
  inference(subsumption_resolution,[],[f9215,f136])).
fof(f9215,plain,(
  ( ! [X92,X93] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X92,X93))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X93)) )),
  inference(superposition,[],[f133,f7670])).
fof(f9268,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(duplicate_literal_removal,[],[f9245])).
fof(f9245,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f7672,f7670])).
fof(f7672,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,sK2(X11,X10)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(resolution,[],[f136,f131])).
fof(f14454,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s(bool,'F_0') = s(bool,sK2(X4,'const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))) ) | $spl126),
  inference(cnf_transformation,[],[f14454_D])).
fof(f14454_D,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s(bool,'F_0') = s(bool,sK2(X4,'const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f9246,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s(bool,'F_0') = s(bool,sK2(X4,'const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f7672,f6476])).
fof(f14453,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X7,'const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9247,f6014])).
fof(f9247,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s(bool,'F_0') = s(bool,sK2(X7,'const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) ) | $spl51),
  inference(superposition,[],[f7672,f6231])).
fof(f14451,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X12,X11))) ) | $spl51),
  inference(subsumption_resolution,[],[f9250,f6278])).
fof(f9250,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))) | s(bool,'F_0') = s(bool,sK2(X12,X11))) ) | $spl51),
  inference(superposition,[],[f7672,f6231])).
fof(f14450,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X15,X14))) ) | $spl51),
  inference(subsumption_resolution,[],[f9251,f6108])).
fof(f9251,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK2(X15,X14))) ) | $spl51),
  inference(superposition,[],[f7672,f6476])).
fof(f14449,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X31)) )),
  inference(subsumption_resolution,[],[f9258,f136])).
fof(f9258,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X31) | p(s(bool,sK2(X30,X31)))) )),
  inference(superposition,[],[f7672,f128])).
fof(f14448,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f9259,f136])).
fof(f9259,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33) | p(s(bool,sK2(X32,X33)))) )),
  inference(superposition,[],[f7672,f130])).
fof(f14400,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,sK2(X66,X67))) ) | $spl102),
  inference(cnf_transformation,[],[f14400_D])).
fof(f14400_D,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,sK2(X66,X67))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f14447,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X40,X41))) )),
  inference(subsumption_resolution,[],[f9263,f8050])).
fof(f8050,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X11) | s(bool,'T_0') = s(bool,sK2(X12,X11))) )),
  inference(resolution,[],[f137,f132])).
fof(f9263,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X41) | s(bool,'T_0') = s(bool,sK2(X40,X41))) )),
  inference(superposition,[],[f7672,f135])).
fof(f8114,plain,(
  ( ! [X39,X41] : (~p(s(bool,sK2(X41,X39)))) ) | $spl72),
  inference(cnf_transformation,[],[f8114_D])).
fof(f8114_D,plain,(
  ( ! [X39,X41] : (~p(s(bool,sK2(X41,X39)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9302,plain,(
  ( ! [X39,X41] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X41,X39)))) )),
  inference(forward_demodulation,[],[f9300,f8121])).
fof(f8121,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,sK2(X6,'const.num.|48|_0'))) )),
  inference(resolution,[],[f8098,f132])).
fof(f8098,plain,(
  ( ! [X0] : (p(s(bool,sK2(X0,'const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f137])).
fof(f9300,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,sK2(X40,'const.num.|48|_0')) | ~p(s(bool,sK2(X41,X39)))) )),
  inference(trivial_inequality_removal,[],[f9286])).
fof(f9286,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',X39) != s('type.num.num',X39) | s(bool,'F_0') != s(bool,sK2(X40,'const.num.|48|_0')) | ~p(s(bool,sK2(X41,X39)))) )),
  inference(superposition,[],[f8048,f136])).
fof(f14444,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X51,X50))) )),
  inference(subsumption_resolution,[],[f9461,f8437])).
fof(f9461,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X51,X50)) | ~p(s(bool,sK2(X49,X50)))) )),
  inference(superposition,[],[f9050,f132])).
fof(f14443,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X54,X53))) )),
  inference(subsumption_resolution,[],[f9462,f8437])).
fof(f9462,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X54,X53)) | ~p(s(bool,sK2(X52,X53)))) )),
  inference(superposition,[],[f9050,f134])).
fof(f14442,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X57,X56))) )),
  inference(subsumption_resolution,[],[f9463,f9050])).
fof(f9463,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X57,X56)) | s(bool,'F_0') = s(bool,sK2(X55,X56))) )),
  inference(superposition,[],[f9050,f135])).
fof(f14440,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10)))) ) | $spl54),
  inference(subsumption_resolution,[],[f9541,f5988])).
fof(f9541,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))) | s(bool,'T_0') = s(bool,sK2(X11,X10))) ) | $spl54),
  inference(superposition,[],[f5988,f9052])).
fof(f9052,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') = s(bool,sK2(X19,X20)) | s(bool,'T_0') = s(bool,sK2(X21,X20))) )),
  inference(resolution,[],[f8437,f132])).
fof(f14438,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) ) | $spl51),
  inference(subsumption_resolution,[],[f9542,f6110])).
fof(f9542,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK2(X14,X13))) ) | $spl51),
  inference(superposition,[],[f6110,f9052])).
fof(f14436,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16)))) ) | $spl51),
  inference(subsumption_resolution,[],[f9543,f6207])).
fof(f9543,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))) | s(bool,'T_0') = s(bool,sK2(X17,X16))) ) | $spl51),
  inference(superposition,[],[f6207,f9052])).
fof(f14435,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9544,f6233])).
fof(f9544,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))) | s(bool,'T_0') = s(bool,sK2(X23,X21))) ) | $spl51),
  inference(superposition,[],[f6233,f9052])).
fof(f14433,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X26,'const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9545,f6258])).
fof(f9545,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,sK2(X26,'const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))) | s(bool,'T_0') = s(bool,sK2(X27,X25))) ) | $spl51),
  inference(superposition,[],[f6258,f9052])).
fof(f14432,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X34,X33))) )),
  inference(subsumption_resolution,[],[f9546,f8050])).
fof(f9546,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33) | s(bool,'T_0') = s(bool,sK2(X34,X33))) )),
  inference(superposition,[],[f7672,f9052])).
fof(f14430,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X76,X75))) ) | $spl51),
  inference(subsumption_resolution,[],[f9557,f6256])).
fof(f9557,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X76,X75)) | s(bool,'T_0') = s(bool,sK2(X77,'const.arithmetic.+_2'(s('type.num.num',X75),s('type.num.num',X75))))) ) | $spl51),
  inference(superposition,[],[f6256,f9052])).
fof(f14429,plain,(
  ( ! [X142,X143] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X143,X142))) )),
  inference(subsumption_resolution,[],[f9574,f9373])).
fof(f9373,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,sK2(X20,X18))) | s(bool,'T_0') = s(bool,sK2(X19,X18))) )),
  inference(trivial_inequality_removal,[],[f9355])).
fof(f9355,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s(bool,'T_0') = s(bool,sK2(X19,X18)) | ~p(s(bool,sK2(X20,X18)))) )),
  inference(superposition,[],[f8050,f136])).
fof(f9574,plain,(
  ( ! [X142,X143,X141] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X141,X142))) | s(bool,'T_0') = s(bool,sK2(X143,X142))) )),
  inference(superposition,[],[f131,f9052])).
fof(f14428,plain,(
  ( ! [X154,X155] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X155,X154))) )),
  inference(subsumption_resolution,[],[f9578,f9373])).
fof(f9578,plain,(
  ( ! [X154,X155,X153] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,sK2(X153,X154))) | s(bool,'T_0') = s(bool,sK2(X155,X154))) )),
  inference(superposition,[],[f133,f9052])).
fof(f14427,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X64,X63))) )),
  inference(subsumption_resolution,[],[f9613,f9050])).
fof(f9613,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X64,X63)) | s(bool,'F_0') = s(bool,sK2(X65,X63))) )),
  inference(superposition,[],[f9050,f9052])).
fof(f14426,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X130,X129))) )),
  inference(subsumption_resolution,[],[f9631,f8437])).
fof(f9631,plain,(
  ( ! [X130,X128,X129] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X128,X129))) | s(bool,'F_0') = s(bool,sK2(X130,X129))) )),
  inference(superposition,[],[f127,f9052])).
fof(f14425,plain,(
  ( ! [X142,X141] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X142,X141))) )),
  inference(subsumption_resolution,[],[f9635,f8437])).
fof(f9635,plain,(
  ( ! [X142,X140,X141] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X140,X141))) | s(bool,'F_0') = s(bool,sK2(X142,X141))) )),
  inference(superposition,[],[f129,f9052])).
fof(f129,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f9659,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X0,X1))) )),
  inference(equality_factoring,[],[f9052])).
fof(f14424,plain,(
  ( ! [X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X7,X5))) )),
  inference(subsumption_resolution,[],[f9693,f9052])).
fof(f9693,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X6,X5)) | s(bool,'T_0') = s(bool,sK2(X7,X5))) )),
  inference(superposition,[],[f9083,f9052])).
fof(f14423,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f9694,f7670])).
fof(f9694,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X10,X9)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f9083,f7670])).
fof(f14422,plain,(
  ( ! [X14,X12] : (s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) | $spl116),
  inference(cnf_transformation,[],[f14422_D])).
fof(f14422_D,plain,(
  ( ! [X14,X12] : (s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f14421,plain,(
  ( ! [X14,X12] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9695,f6476])).
fof(f9695,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X13,X12)) | s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f9083,f6476])).
fof(f14420,plain,(
  ( ! [X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9696,f6231])).
fof(f9696,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X17,X16)) | s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) | $spl51),
  inference(superposition,[],[f9083,f6231])).
fof(f14418,plain,(
  ( ! [X26,X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X26,X24))) ) | $spl51),
  inference(subsumption_resolution,[],[f9699,f6231])).
fof(f9699,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))) | s(bool,'F_0') = s(bool,sK2(X26,X24))) ) | $spl51),
  inference(superposition,[],[f9083,f6231])).
fof(f14417,plain,(
  ( ! [X30,X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X30,X28))) ) | $spl51),
  inference(subsumption_resolution,[],[f9700,f6476])).
fof(f9700,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') = s(bool,sK2(X30,X28))) ) | $spl51),
  inference(superposition,[],[f9083,f6476])).
fof(f14415,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') != s(bool,sK2(X56,X55))) ) | $spl114),
  inference(cnf_transformation,[],[f14415_D])).
fof(f14415_D,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') != s(bool,sK2(X56,X55))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f14416,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X53,X52))) )),
  inference(subsumption_resolution,[],[f9707,f9301])).
fof(f9301,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,sK2(X20,X18))) | s(bool,'F_0') != s(bool,sK2(X19,X18))) )),
  inference(trivial_inequality_removal,[],[f9283])).
fof(f9283,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s(bool,'F_0') != s(bool,sK2(X19,X18)) | ~p(s(bool,sK2(X20,X18)))) )),
  inference(superposition,[],[f8048,f136])).
fof(f9707,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X53,X52)) | p(s(bool,sK2(X51,X52)))) )),
  inference(superposition,[],[f9083,f128])).
fof(f14414,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X56,X55))) )),
  inference(subsumption_resolution,[],[f9708,f9301])).
fof(f9708,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X56,X55)) | p(s(bool,sK2(X54,X55)))) )),
  inference(superposition,[],[f9083,f130])).
fof(f14413,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X66,X67))) )),
  inference(subsumption_resolution,[],[f9712,f9052])).
fof(f9712,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,sK2(X68,X67)) | s(bool,'T_0') = s(bool,sK2(X66,X67))) )),
  inference(superposition,[],[f9083,f135])).
fof(f14412,plain,(
  ( ! [X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X7,X5))) )),
  inference(subsumption_resolution,[],[f9763,f9085])).
fof(f9085,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,sK2(X19,X20)) | s(bool,'T_0') = s(bool,sK2(X21,X20))) )),
  inference(resolution,[],[f8439,f132])).
fof(f9763,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X6,X5)) | s(bool,'T_0') = s(bool,sK2(X7,X5))) )),
  inference(superposition,[],[f9085,f9052])).
fof(f14410,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f9764,f7672])).
fof(f9764,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X10,X9)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f9085,f7670])).
fof(f14409,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,sK2(X13,X12)) | s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) | $spl110),
  inference(cnf_transformation,[],[f14409_D])).
fof(f14409_D,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,sK2(X13,X12)) | s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f9765,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X13,X12)) | s(bool,'F_0') = s(bool,sK2(X14,'const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f9085,f6476])).
fof(f14407,plain,(
  ( ! [X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) | $spl51),
  inference(subsumption_resolution,[],[f9766,f6233])).
fof(f9766,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X17,X16)) | s(bool,'F_0') = s(bool,sK2(X18,'const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X16))))) ) | $spl51),
  inference(superposition,[],[f9085,f6231])).
fof(f14404,plain,(
  ( ! [X26,X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK2(X26,X24))) ) | $spl51),
  inference(subsumption_resolution,[],[f9769,f6256])).
fof(f9769,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))) | s(bool,'F_0') = s(bool,sK2(X26,X24))) ) | $spl51),
  inference(superposition,[],[f9085,f6231])).
fof(f14403,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') = s(bool,sK2(X30,X28))) ) | $spl104),
  inference(cnf_transformation,[],[f14403_D])).
fof(f14403_D,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') = s(bool,sK2(X30,X28))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f9770,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X29,'const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') = s(bool,sK2(X30,X28))) ) | $spl51),
  inference(superposition,[],[f9085,f6476])).
fof(f14402,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X53,X52))) )),
  inference(subsumption_resolution,[],[f9777,f9373])).
fof(f9777,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X53,X52)) | p(s(bool,sK2(X51,X52)))) )),
  inference(superposition,[],[f9085,f128])).
fof(f14401,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X56,X55))) )),
  inference(subsumption_resolution,[],[f9778,f9373])).
fof(f9778,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X56,X55)) | p(s(bool,sK2(X54,X55)))) )),
  inference(superposition,[],[f9085,f130])).
fof(f14399,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X66,X67))) )),
  inference(subsumption_resolution,[],[f9782,f9085])).
fof(f9782,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK2(X68,X67)) | s(bool,'T_0') = s(bool,sK2(X66,X67))) )),
  inference(superposition,[],[f9085,f135])).
fof(f13767,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X2,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,sK2(X3,'const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6231])).
fof(f13768,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X6,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) | s(bool,'F_0') = s(bool,sK2(X7,'const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6476])).
fof(f14366,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f13769,f6825])).
fof(f6825,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | ~p(s(bool,sK2(X0,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))))) ) | $spl58),
  inference(resolution,[],[f6814,f137])).
fof(f13769,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X10,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) ) | $spl58),
  inference(superposition,[],[f6830,f7670])).
fof(f14365,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X13,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f13770,f6830])).
fof(f13770,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X13,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) | s(bool,'T_0') = s(bool,sK2(X14,X12))) ) | $spl58),
  inference(superposition,[],[f6830,f9052])).
fof(f13780,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X53,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f5815])).
fof(f13781,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X56,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X55))))))) | s(bool,'F_0') = s(bool,sK2(X57,X55))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6231])).
fof(f13782,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X60,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,sK2(X61,X59))) ) | ($spl51 | $spl58)),
  inference(superposition,[],[f6830,f6476])).
fof(f14361,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X90,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X89)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f13791,f6814])).
fof(f13791,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X90,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X89))))) | p(s(bool,sK2(X88,X89)))) ) | $spl58),
  inference(superposition,[],[f6830,f128])).
fof(f14360,plain,(
  ( ! [X92,X93] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X93,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X92)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f13792,f6814])).
fof(f13792,plain,(
  ( ! [X92,X93,X91] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X93,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X92))))) | p(s(bool,sK2(X91,X92)))) ) | $spl58),
  inference(superposition,[],[f6830,f130])).
fof(f14359,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X105,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X104)))))) ) | $spl58),
  inference(subsumption_resolution,[],[f13796,f6830])).
fof(f13796,plain,(
  ( ! [X103,X105,X104] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,sK2(X105,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X104))))) | s(bool,'T_0') = s(bool,sK2(X103,X104))) ) | $spl58),
  inference(superposition,[],[f6830,f135])).
fof(f14350,plain,(
  $false | $spl3),
  inference(trivial_inequality_removal,[],[f14349])).
fof(f14349,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl3),
  inference(forward_demodulation,[],[f205,f8842])).
fof(f8842,plain,(
  ( ! [X64] : (s('type.typeSystem.tenv_val',X64) = s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X64)))) )),
  inference(forward_demodulation,[],[f8841,f165])).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X11) = s(X14,'const.bool.COND_3'(s(bool,'T_0'),s(X14,X11),s(X14,X12)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.COND_CLAUSES|split|0')).
fof(f8841,plain,(
  ( ! [X64] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X64))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,'T_0'),s('type.typeSystem.tenv_val',X64),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X64)))))) )),
  inference(forward_demodulation,[],[f8763,f8177])).
fof(f8177,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,sK2(X6,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0'))))) )),
  inference(resolution,[],[f8099,f132])).
fof(f8099,plain,(
  ( ! [X64] : (p(s(bool,sK2(X64,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))))) )),
  inference(trivial_inequality_removal,[],[f8093])).
fof(f8093,plain,(
  ( ! [X64] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,sK2(X64,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))))) )),
  inference(superposition,[],[f137,f163])).
fof(f163,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))),
  inference(cnf_transformation,[],[f43])).
fof(f43,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSystem.num_tvs_def|split|0')).
fof(f8763,plain,(
  ( ! [X64] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X64))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,sK2(X64,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))),s('type.typeSystem.tenv_val',X64),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X64)))))) )),
  inference(superposition,[],[f138,f163])).
fof(f14348,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f14347,f210])).
fof(f210,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl10),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f14347,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f970,f8842])).
fof(f970,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f205,f210])).
fof(f14346,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f14345,f175])).
fof(f175,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3',aHL_FALSITY)).
fof(f14345,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f14344,f210])).
fof(f14344,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2361,f8842])).
fof(f2361,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | ($spl3 | $spl10)),
  inference(trivial_inequality_removal,[],[f2353])).
fof(f2353,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | ($spl3 | $spl10)),
  inference(superposition,[],[f970,f128])).
fof(f14343,plain,(
  $false | ($spl3 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f14342,f1128])).
fof(f1128,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl16),
  inference(forward_demodulation,[],[f114,f217])).
fof(f14342,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2359,f8842])).
fof(f2359,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ($spl3 | $spl10)),
  inference(trivial_inequality_removal,[],[f2358])).
fof(f2358,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ($spl3 | $spl10)),
  inference(superposition,[],[f970,f135])).
fof(f14341,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f14340,f175])).
fof(f14340,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f14339,f210])).
fof(f14339,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f2360,f8842])).
fof(f2360,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | ($spl3 | $spl10)),
  inference(trivial_inequality_removal,[],[f2354])).
fof(f2354,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | ($spl3 | $spl10)),
  inference(superposition,[],[f970,f130])).
fof(f14338,plain,(
  $false | ($spl16 | $spl25)),
  inference(subsumption_resolution,[],[f14337,f1128])).
fof(f14337,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl25),
  inference(forward_demodulation,[],[f922,f8842])).
fof(f14336,plain,(
  $false | ($spl10 | $spl25)),
  inference(subsumption_resolution,[],[f14335,f175])).
fof(f14335,plain,(
  p(s(bool,'F_0')) | ($spl10 | $spl25)),
  inference(forward_demodulation,[],[f14334,f210])).
fof(f14334,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl25),
  inference(forward_demodulation,[],[f923,f8842])).
fof(f923,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | $spl25),
  inference(trivial_inequality_removal,[],[f917])).
fof(f917,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | $spl25),
  inference(superposition,[],[f391,f130])).
fof(f14333,plain,(
  $false | ($spl10 | $spl25)),
  inference(subsumption_resolution,[],[f14332,f175])).
fof(f14332,plain,(
  p(s(bool,'F_0')) | ($spl10 | $spl25)),
  inference(forward_demodulation,[],[f14331,f210])).
fof(f14331,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl25),
  inference(forward_demodulation,[],[f924,f8842])).
fof(f924,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | $spl25),
  inference(trivial_inequality_removal,[],[f916])).
fof(f916,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | $spl25),
  inference(superposition,[],[f391,f128])).
fof(f14330,plain,(
  $false | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14329,f175])).
fof(f14329,plain,(
  p(s(bool,'F_0')) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f14328,f210])).
fof(f14328,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl26),
  inference(forward_demodulation,[],[f432,f8842])).
fof(f14287,plain,(
  ( ! [X447,X446] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X446))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X447)))) ) | $spl86),
  inference(cnf_transformation,[],[f14287_D])).
fof(f14287_D,plain,(
  ( ! [X447,X446] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X446))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X447)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3515,plain,(
  ( ! [X19,X20,X18] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X18))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X19))) | ~p(s(bool,sK2(X20,sK0)))) ) | $spl30),
  inference(superposition,[],[f1107,f136])).
fof(f14286,plain,(
  ( ! [X453,X451,X452] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X451))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X452),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X453),s('type.typeSystem.tenv_val',X451)))) ) | $spl84),
  inference(cnf_transformation,[],[f14286_D])).
fof(f14286_D,plain,(
  ( ! [X453,X451,X452] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X451))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X452),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X453),s('type.typeSystem.tenv_val',X451)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f14294,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22))) | ~p(s(bool,sK2(X25,sK0)))) ) | $spl10),
  inference(forward_demodulation,[],[f5502,f8842])).
fof(f5502,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X22))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22))) | ~p(s(bool,sK2(X25,sK0)))) ) | $spl10),
  inference(superposition,[],[f4399,f136])).
fof(f4399,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X14))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl10),
  inference(superposition,[],[f4313,f116])).
fof(f4313,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl10),
  inference(backward_demodulation,[],[f210,f114])).
fof(f14283,plain,(
  ( ! [X6,X5] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',sK1)))) ) | $spl82),
  inference(cnf_transformation,[],[f14283_D])).
fof(f14283_D,plain,(
  ( ! [X6,X5] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',sK1)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f14293,plain,(
  ( ! [X19,X20,X18] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X18),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,sK2(X20,sK0)))) ) | $spl16),
  inference(forward_demodulation,[],[f5628,f8842])).
fof(f5628,plain,(
  ( ! [X19,X20,X18] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X18),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X20,sK0)))) ) | $spl16),
  inference(superposition,[],[f4666,f136])).
fof(f4666,plain,(
  ( ! [X2,X3] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X3),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) ) | $spl16),
  inference(trivial_inequality_removal,[],[f4648])).
fof(f4648,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'T_0') | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X3),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) ) | $spl16),
  inference(superposition,[],[f4271,f217])).
fof(f4271,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl16),
  inference(superposition,[],[f1128,f116])).
fof(f14292,plain,(
  ( ! [X24,X23,X22] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22)))) ) | $spl90),
  inference(cnf_transformation,[],[f14292_D])).
fof(f14292_D,plain,(
  ( ! [X24,X23,X22] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f14291,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22))) | ~p(s(bool,sK2(X25,sK0)))) )),
  inference(trivial_inequality_removal,[],[f14290])).
fof(f14290,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22))) | ~p(s(bool,sK2(X25,sK0)))) )),
  inference(forward_demodulation,[],[f6393,f8842])).
fof(f6393,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X22))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X22))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val',X22))) | ~p(s(bool,sK2(X25,sK0)))) )),
  inference(superposition,[],[f187,f136])).
fof(f187,plain,(
  ( ! [X14,X15,X16] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X14))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) )),
  inference(superposition,[],[f114,f116])).
fof(f14288,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23))))) | ~p(s(bool,sK2(X25,sK0)))) ) | $spl42),
  inference(subsumption_resolution,[],[f6779,f8101])).
fof(f6779,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23))))) | ~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23))))) | ~p(s(bool,sK2(X25,sK0)))) ) | $spl42),
  inference(superposition,[],[f5734,f136])).
fof(f5734,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,sK2(X31,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X30))))) | ~p(s(bool,sK2(X32,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X30)))))) ) | $spl42),
  inference(trivial_inequality_removal,[],[f5731])).
fof(f5731,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ~p(s(bool,sK2(X31,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X30))))) | ~p(s(bool,sK2(X32,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X30)))))) ) | $spl42),
  inference(superposition,[],[f4945,f136])).
fof(f4945,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X21))) | ~p(s(bool,sK2(X22,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X21)))))) ) | $spl42),
  inference(superposition,[],[f3732,f136])).
fof(f7921,plain,(
  ( ! [X447,X446,X448] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X446))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X447))) | ~p(s(bool,sK2(X448,sK0)))) ) | $spl30),
  inference(superposition,[],[f1107,f136])).
fof(f14285,plain,(
  ( ! [X453,X451,X452,X454] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X451))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X452),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X453),s('type.typeSystem.tenv_val',X451))) | ~p(s(bool,sK2(X454,sK0)))) ) | $spl10),
  inference(forward_demodulation,[],[f7923,f8842])).
fof(f7923,plain,(
  ( ! [X453,X451,X452,X454] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X451))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X452),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X453),s('type.typeSystem.tenv_val',X451))) | ~p(s(bool,sK2(X454,sK0)))) ) | $spl10),
  inference(superposition,[],[f4399,f136])).
fof(f14284,plain,(
  ( ! [X455,X457,X456] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X455),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X456),s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,sK2(X457,sK0)))) ) | $spl16),
  inference(forward_demodulation,[],[f7924,f8842])).
fof(f7924,plain,(
  ( ! [X455,X457,X456] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X455),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X456),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X457,sK0)))) ) | $spl16),
  inference(superposition,[],[f4666,f136])).
fof(f14282,plain,(
  s(bool,'F_0') = s(bool,sK2(sK1,sK0)) | $spl80),
  inference(cnf_transformation,[],[f14282_D])).
fof(f14282_D,plain,(
  s(bool,'F_0') = s(bool,sK2(sK1,sK0)) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f12983,plain,(
  ( ! [X6,X5] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,sK2(sK1,sK0))) ) | $spl16),
  inference(superposition,[],[f4666,f8838])).
fof(f8838,plain,(
  ( ! [X37,X36] : (s('type.typeSystem.tenv_val',X36) = s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X37),s('type.typeSystem.tenv_val',X36))) | s(bool,'F_0') = s(bool,sK2(X36,X37))) )),
  inference(forward_demodulation,[],[f8747,f165])).
fof(f8747,plain,(
  ( ! [X37,X36] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X37),s('type.typeSystem.tenv_val',X36))) = s('type.typeSystem.tenv_val','const.bool.COND_3'(s(bool,'T_0'),s('type.typeSystem.tenv_val',X36),s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X37),s('type.typeSystem.tenv_val',X36))))) | s(bool,'F_0') = s(bool,sK2(X36,X37))) )),
  inference(superposition,[],[f138,f135])).
fof(f14247,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f14236,f210])).
fof(f14236,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl0 | $spl10)),
  inference(trivial_inequality_removal,[],[f14187])).
fof(f14187,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl10)),
  inference(superposition,[],[f4392,f14010])).
fof(f14010,plain,(
  ( ! [X21] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X21))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X21)))))) ) | $spl0),
  inference(superposition,[],[f149,f13969])).
fof(f13969,plain,(
  ( ! [X51] : (s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X51))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X51)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13943,f175])).
fof(f13943,plain,(
  ( ! [X51] : (p(s(bool,'F_0')) | s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X51))) = s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X51)))) ) | $spl0),
  inference(superposition,[],[f8834,f439])).
fof(f439,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,sK2(X4,sK0))) ) | $spl0),
  inference(resolution,[],[f203,f128])).
fof(f4392,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',sK1))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))) ) | $spl10),
  inference(superposition,[],[f4313,f118])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f14246,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f14186,f1128])).
fof(f14186,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl0 | $spl16)),
  inference(superposition,[],[f217,f14010])).
fof(f14237,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f14175,f1128])).
fof(f14175,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl0 | $spl16)),
  inference(superposition,[],[f14010,f217])).
fof(f12525,plain,(
  ( ! [X47,X48] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X48)))))) ) | $spl78),
  inference(cnf_transformation,[],[f12525_D])).
fof(f12525_D,plain,(
  ( ! [X47,X48] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X48)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f12523,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X48))))) | ~p(s(bool,sK2(X49,X46)))) )),
  inference(trivial_inequality_removal,[],[f12509])).
fof(f12509,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.num.num',X46) != s('type.num.num',X46) | p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X48))))) | ~p(s(bool,sK2(X49,X46)))) )),
  inference(superposition,[],[f8096,f136])).
fof(f8096,plain,(
  ( ! [X72,X71,X73] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X71) | p(s(bool,sK2(X73,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X71),s('type.num.num',X72)))))) )),
  inference(superposition,[],[f137,f167])).
fof(f167,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X11) = s(X14,'const.bool.COND_3'(s(bool,'T_0'),s(X14,X11),s(X14,X12)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.COND_CLAUSES|split|0')).
fof(f12065,plain,(
  ( ! [X47,X48] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X48),s('type.num.num','const.num.|48|_0')))))) ) | $spl76),
  inference(cnf_transformation,[],[f12065_D])).
fof(f12065_D,plain,(
  ( ! [X47,X48] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X48),s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f12063,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X48),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK2(X49,X46)))) )),
  inference(trivial_inequality_removal,[],[f12049])).
fof(f12049,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.num.num',X46) != s('type.num.num',X46) | p(s(bool,sK2(X47,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X48),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK2(X49,X46)))) )),
  inference(superposition,[],[f8094,f136])).
fof(f8094,plain,(
  ( ! [X66,X67,X65] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X66) | p(s(bool,sK2(X67,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X65),s('type.num.num',X66)))))) )),
  inference(superposition,[],[f137,f166])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X12) = s(X14,'const.bool.COND_3'(s(bool,'F_0'),s(X14,X11),s(X14,X12)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.COND_CLAUSES|split|1')).
fof(f8143,plain,(
  $false | $spl72),
  inference(resolution,[],[f8114,f8098])).
fof(f8115,plain,(
  ( ! [X40] : (p(s(bool,sK2(X40,'const.num.|48|_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f8115_D])).
fof(f8115_D,plain,(
  ( ! [X40] : (p(s(bool,sK2(X40,'const.num.|48|_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8100,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,sK2(X40,'const.num.|48|_0'))) | ~p(s(bool,sK2(X41,X39)))) )),
  inference(trivial_inequality_removal,[],[f8086])).
fof(f8086,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',X39) != s('type.num.num',X39) | p(s(bool,sK2(X40,'const.num.|48|_0'))) | ~p(s(bool,sK2(X41,X39)))) )),
  inference(superposition,[],[f137,f136])).
fof(f8113,plain,(
  $false | $spl66),
  inference(subsumption_resolution,[],[f8112,f175])).
fof(f8112,plain,(
  p(s(bool,'F_0')) | $spl66),
  inference(forward_demodulation,[],[f8098,f7952])).
fof(f8106,plain,(
  $false | $spl66),
  inference(subsumption_resolution,[],[f8104,f175])).
fof(f8104,plain,(
  p(s(bool,'F_0')) | $spl66),
  inference(trivial_inequality_removal,[],[f8056])).
fof(f8056,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl66),
  inference(superposition,[],[f137,f7952])).
fof(f7952,plain,(
  ( ! [X335] : (s(bool,'F_0') = s(bool,sK2(X335,'const.num.|48|_0'))) ) | $spl66),
  inference(cnf_transformation,[],[f7952_D])).
fof(f7952_D,plain,(
  ( ! [X335] : (s(bool,'F_0') = s(bool,sK2(X335,'const.num.|48|_0'))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7953,plain,(
  ( ! [X349,X351,X348] : (s(bool,'F_0') = s(bool,sK2(X349,'const.arithmetic.+_2'(s('type.num.num',X348),s('type.num.num',X348)))) | ~p(s(bool,sK2(X351,X348)))) ) | $spl68),
  inference(cnf_transformation,[],[f7953_D])).
fof(f7953_D,plain,(
  ( ! [X349,X351,X348] : (s(bool,'F_0') = s(bool,sK2(X349,'const.arithmetic.+_2'(s('type.num.num',X348),s('type.num.num',X348)))) | ~p(s(bool,sK2(X351,X348)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7796,plain,(
  ( ! [X364,X362,X365,X363] : (s(bool,'F_0') = s(bool,sK2(X363,'const.arithmetic.+_2'(s('type.num.num',X362),s('type.num.num',X362)))) | s(bool,'F_0') = s(bool,sK2(X364,'const.num.|48|_0')) | ~p(s(bool,sK2(X365,X362)))) ) | $spl51),
  inference(superposition,[],[f6476,f136])).
fof(f5739,plain,(
  ( ! [X39,X41] : (~p(s(bool,sK2(X41,X39))) | s('type.num.num',X39) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X39)))) ) | $spl54),
  inference(cnf_transformation,[],[f5739_D])).
fof(f5739_D,plain,(
  ( ! [X39,X41] : (~p(s(bool,sK2(X41,X39))) | s('type.num.num',X39) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X39)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7795,plain,(
  ( ! [X360,X359,X361] : (s('type.num.num',X359) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X359),s('type.num.num',X359))) | s(bool,'F_0') = s(bool,sK2(X360,'const.num.|48|_0')) | ~p(s(bool,sK2(X361,X359)))) ) | $spl51),
  inference(superposition,[],[f6459,f136])).
fof(f6459,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s(bool,'F_0') = s(bool,sK2(X1,X0))) ) | $spl51),
  inference(superposition,[],[f6108,f124])).
fof(f124,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X8)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.arithmetic.ADD_SYM')).
fof(f7794,plain,(
  ( ! [X356,X358,X357] : (s('type.num.num',X356) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X356),s('type.num.num',X356))) | s(bool,'F_0') = s(bool,sK2(X357,'const.num.|48|_0')) | ~p(s(bool,sK2(X358,X356)))) ) | $spl51),
  inference(superposition,[],[f6278,f136])).
fof(f7954,plain,(
  ( ! [X352,X353,X355] : (s(bool,'T_0') != s(bool,sK2(X353,'const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352)))) | ~p(s(bool,sK2(X355,X352)))) ) | $spl70),
  inference(cnf_transformation,[],[f7954_D])).
fof(f7954_D,plain,(
  ( ! [X352,X353,X355] : (s(bool,'T_0') != s(bool,sK2(X353,'const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352)))) | ~p(s(bool,sK2(X355,X352)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7793,plain,(
  ( ! [X352,X354,X353,X355] : (s(bool,'T_0') != s(bool,sK2(X353,'const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352)))) | s(bool,'F_0') = s(bool,sK2(X354,'const.num.|48|_0')) | ~p(s(bool,sK2(X355,X352)))) ) | $spl51),
  inference(superposition,[],[f6256,f136])).
fof(f7792,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'F_0') = s(bool,sK2(X349,'const.arithmetic.+_2'(s('type.num.num',X348),s('type.num.num',X348)))) | s(bool,'F_0') = s(bool,sK2(X350,'const.num.|48|_0')) | ~p(s(bool,sK2(X351,X348)))) ) | $spl51),
  inference(superposition,[],[f6231,f136])).
fof(f7791,plain,(
  ( ! [X346,X345,X347] : (s('type.num.num',X345) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X345),s('type.num.num',X345))) | s(bool,'F_0') = s(bool,sK2(X346,'const.num.|48|_0')) | ~p(s(bool,sK2(X347,X345)))) ) | $spl51),
  inference(superposition,[],[f6108,f136])).
fof(f7627,plain,(
  $false | $spl64),
  inference(resolution,[],[f7555,f175])).
fof(f7573,plain,(
  $false | $spl62),
  inference(resolution,[],[f7434,f162])).
fof(f162,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3',aHL_TRUTH)).
fof(f7434,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl62),
  inference(superposition,[],[f175,f7345])).
fof(f7574,plain,(
  $false | $spl62),
  inference(resolution,[],[f7434,f161])).
fof(f161,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.TRUTH')).
fof(f7577,plain,(
  $false | $spl62),
  inference(resolution,[],[f7434,f176])).
fof(f176,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X14] : ! [X15] : (s(X14,X15) = s(X14,X15) <=> p(s(bool,'T_0')))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.REFL_CLAUSE')).
fof(f2898,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl40),
  inference(cnf_transformation,[],[f2898_D])).
fof(f2898_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7433,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,X23))) ) | $spl62),
  inference(superposition,[],[f174,f7345])).
fof(f174,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : (p(s(bool,'F_0')) => p(s(bool,X13)))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.FALSITY')).
fof(f7432,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21))) ) | $spl62),
  inference(superposition,[],[f173,f7345])).
fof(f173,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) & p(s(bool,X13))) <=> p(s(bool,'F_0')))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.AND_CLAUSES|split|2')).
fof(f7431,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19))) ) | $spl62),
  inference(superposition,[],[f169,f7345])).
fof(f169,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : ((p(s(bool,X13)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.bool.AND_CLAUSES|split|3')).
fof(f7345,plain,(
  ( ! [X20] : (s(bool,'F_0') = s(bool,X20)) ) | $spl62),
  inference(cnf_transformation,[],[f7345_D])).
fof(f7345_D,plain,(
  ( ! [X20] : (s(bool,'F_0') = s(bool,X20)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7355,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'F_0') = s(bool,X135) | ~p(s(bool,sK2(X133,X134))) | s(bool,'F_0') = s(bool,sK2(X136,'const.arithmetic.+_2'(s('type.num.num',X134),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7241,f128])).
fof(f7241,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'F_0') = s(bool,X135) | ~p(s(bool,X135)) | ~p(s(bool,sK2(X133,X134))) | s(bool,'F_0') = s(bool,sK2(X136,'const.arithmetic.+_2'(s('type.num.num',X134),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f123,f6476])).
fof(f7354,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'F_0') = s(bool,X131) | ~p(s(bool,sK2(X129,X130))) | s(bool,'F_0') = s(bool,sK2(X132,'const.arithmetic.+_2'(s('type.num.num',X130),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(subsumption_resolution,[],[f7240,f128])).
fof(f7240,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'F_0') = s(bool,X131) | ~p(s(bool,sK2(X129,X130))) | ~p(s(bool,X131)) | s(bool,'F_0') = s(bool,sK2(X132,'const.arithmetic.+_2'(s('type.num.num',X130),s('type.num.num','const.num.|48|_0'))))) ) | $spl51),
  inference(superposition,[],[f123,f6476])).
fof(f7346,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,'F_0') = s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X22,X23)))) ) | $spl51),
  inference(subsumption_resolution,[],[f7167,f128])).
fof(f7167,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,'F_0') = s(bool,sK2(X25,'const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X24)) | ~p(s(bool,sK2(X22,X23)))) ) | $spl51),
  inference(superposition,[],[f6476,f123])).
fof(f7343,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK2(X21,'const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X18,X19)))) ) | $spl51),
  inference(subsumption_resolution,[],[f7166,f128])).
fof(f7166,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK2(X21,'const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK2(X18,X19))) | ~p(s(bool,X20))) ) | $spl51),
  inference(superposition,[],[f6476,f123])).
fof(f6814,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,sK2(X25,X22))) | ~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))))) ) | $spl58),
  inference(cnf_transformation,[],[f6814_D])).
fof(f6814_D,plain,(
  ( ! [X24,X25,X22] : (~p(s(bool,sK2(X25,X22))) | ~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5738,plain,(
  ( ! [X40] : (~p(s(bool,sK2(X40,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0')))))) ) | $spl52),
  inference(cnf_transformation,[],[f5738_D])).
fof(f5738_D,plain,(
  ( ! [X40] : (~p(s(bool,sK2(X40,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6786,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,sK2(X23,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK2(X24,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22))))) | ~p(s(bool,sK2(X25,X22)))) ) | $spl42),
  inference(superposition,[],[f5734,f136])).
fof(f5810,plain,(
  ( ! [X19,X17,X20] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X17))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X17))))) ) | $spl56),
  inference(cnf_transformation,[],[f5810_D])).
fof(f5810_D,plain,(
  ( ! [X19,X17,X20] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X17))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X17))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5811,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X21))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X21))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5789,f127])).
fof(f5789,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X21))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X24),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X21)))) | p(s(bool,X22))) ) | $spl10),
  inference(superposition,[],[f4403,f120])).
fof(f4403,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) ) | $spl10),
  inference(superposition,[],[f4313,f116])).
fof(f5809,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X17))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X17))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5788,f127])).
fof(f5788,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'F_0') != s(bool,X18) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X17))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X20),s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | p(s(bool,X18)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X17))))) ) | $spl10),
  inference(superposition,[],[f4403,f120])).
fof(f5710,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',X39) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X39))) | ~p(s(bool,sK2(X40,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK2(X41,X39)))) ) | $spl42),
  inference(superposition,[],[f4945,f136])).
fof(f5735,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,sK2(X19,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X18))))) | ~p(s(bool,sK2(X20,X18)))) ) | $spl48),
  inference(cnf_transformation,[],[f5735_D])).
fof(f5735_D,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,sK2(X19,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X18))))) | ~p(s(bool,sK2(X20,X18)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5737,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl51),
  inference(cnf_transformation,[],[f5737_D])).
fof(f5737_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f5707,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK2(X19,'const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X18))))) | ~p(s(bool,sK2(X20,X18)))) ) | $spl42),
  inference(superposition,[],[f4945,f136])).
fof(f3987,plain,(
  ( ! [X4,X2,X5] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) | $spl46),
  inference(cnf_transformation,[],[f3987_D])).
fof(f3987_D,plain,(
  ( ! [X4,X2,X5] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5543,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X9),s('type.typeSystem.tenv_val',X6))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5524,f127])).
fof(f5524,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X9),s('type.typeSystem.tenv_val',X6))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6)))))) | p(s(bool,X7))) ) | $spl10),
  inference(superposition,[],[f4399,f120])).
fof(f5542,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5523,f127])).
fof(f5523,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,X3)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) | $spl10),
  inference(superposition,[],[f4399,f120])).
fof(f3742,plain,(
  ( ! [X4,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl44),
  inference(cnf_transformation,[],[f3742_D])).
fof(f3742_D,plain,(
  ( ! [X4,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5262,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5238,f127])).
fof(f5238,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X6))) ) | $spl10),
  inference(superposition,[],[f4392,f120])).
fof(f5261,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5237,f127])).
fof(f5237,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,X3)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f4392,f120])).
fof(f1117,plain,(
  ( ! [X6,X4,X7] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl34),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ( ! [X6,X4,X7] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4909,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4886,f127])).
fof(f4886,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X9))) ) | $spl10),
  inference(superposition,[],[f4390,f120])).
fof(f4390,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',sK1))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X2)))) ) | $spl10),
  inference(superposition,[],[f4313,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X0) = s('type.num.num',X2) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',X3)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f4908,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4885,f127])).
fof(f4885,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,X5)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f4390,f120])).
fof(f4853,plain,(
  $false | ($spl24 | $spl32)),
  inference(trivial_inequality_removal,[],[f4847])).
fof(f4847,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ($spl24 | $spl32)),
  inference(superposition,[],[f4801,f163])).
fof(f4801,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f4800,f137])).
fof(f4800,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK2(X1,X0)))) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3266,f1108])).
fof(f3266,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK2(X1,X0))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',sK1)))))) ) | $spl24),
  inference(superposition,[],[f390,f136])).
fof(f390,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl24),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4852,plain,(
  $false | ($spl24 | $spl32)),
  inference(equality_resolution,[],[f4801])).
fof(f2496,plain,(
  ( ! [X23,X22,X20] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20))))) ) | $spl36),
  inference(cnf_transformation,[],[f2496_D])).
fof(f2496_D,plain,(
  ( ! [X23,X22,X20] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4673,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4655,f131])).
fof(f4655,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,X25)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(superposition,[],[f4271,f123])).
fof(f4672,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4654,f131])).
fof(f4654,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20)))) | ~p(s(bool,X21))) ) | $spl16),
  inference(superposition,[],[f4271,f123])).
fof(f4421,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4406,f127])).
fof(f4406,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f4313,f120])).
fof(f4420,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4405,f127])).
fof(f4405,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f4313,f120])).
fof(f217,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl16),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f211,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl11),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f4310,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f4306,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f4305,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f4304,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f4303,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f4302,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f4301,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f1465,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f128])).
fof(f1466,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f130])).
fof(f1966,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f128])).
fof(f1967,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f130])).
fof(f4214,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f1947])).
fof(f1947,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | ~p(s(bool,X2))) ) | $spl16),
  inference(superposition,[],[f1472,f123])).
fof(f1472,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(trivial_inequality_removal,[],[f1468])).
fof(f1468,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f134])).
fof(f4215,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f1948])).
fof(f1948,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(superposition,[],[f1472,f123])).
fof(f2934,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f128])).
fof(f2935,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl16),
  inference(superposition,[],[f1128,f130])).
fof(f2896,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f1947])).
fof(f2897,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f1948])).
fof(f4218,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2915])).
fof(f2915,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | ~p(s(bool,X2))) ) | $spl16),
  inference(superposition,[],[f1472,f123])).
fof(f4219,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2916])).
fof(f2916,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl16),
  inference(superposition,[],[f1472,f123])).
fof(f4085,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f4086,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f4090,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f4205,plain,(
  $false | ($spl2 | $spl12 | $spl14 | $spl24)),
  inference(subsumption_resolution,[],[f4204,f2597])).
fof(f4204,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl12 | $spl24)),
  inference(forward_demodulation,[],[f4000,f212])).
fof(f4000,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl2 | $spl24)),
  inference(forward_demodulation,[],[f204,f390])).
fof(f204,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl2),
  inference(cnf_transformation,[],[f204_D])).
fof(f204_D,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4203,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4192,f131])).
fof(f4192,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(superposition,[],[f4092,f123])).
fof(f4092,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl12),
  inference(backward_demodulation,[],[f212,f114])).
fof(f4202,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4191,f131])).
fof(f4191,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f4092,f123])).
fof(f4089,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f4088,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f4087,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f4084,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f4083,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f4082,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f4081,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f3988,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X9),s('type.typeSystem.tenv_val',X6))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3968,f127])).
fof(f3968,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X9),s('type.typeSystem.tenv_val',X6))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6)))))) | p(s(bool,X7))) ) | $spl10),
  inference(superposition,[],[f3063,f120])).
fof(f3063,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X14))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl10),
  inference(superposition,[],[f3019,f116])).
fof(f3019,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl10),
  inference(forward_demodulation,[],[f114,f210])).
fof(f3986,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3967,f127])).
fof(f3967,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',X2))) | p(s(bool,X3)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X2))))))) ) | $spl10),
  inference(superposition,[],[f3063,f120])).
fof(f3743,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3717,f127])).
fof(f3717,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X5),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X6))) ) | $spl10),
  inference(superposition,[],[f3056,f120])).
fof(f3056,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X6),s('type.typeSystem.tenv_val',sK1))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X7))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))) ) | $spl10),
  inference(superposition,[],[f3019,f118])).
fof(f3741,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3716,f127])).
fof(f3716,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X4))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))) | p(s(bool,X3)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X2),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f3056,f120])).
fof(f3732,plain,(
  ( ! [X40] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X40))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X40)))) ) | $spl42),
  inference(cnf_transformation,[],[f3732_D])).
fof(f3732_D,plain,(
  ( ! [X40] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X40))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X40)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1108,plain,(
  ( ! [X46,X49] : (~p(s(bool,sK2(X49,X46))) | s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X46),s('type.typeSystem.tenv_val',sK1)))))) ) | $spl32),
  inference(cnf_transformation,[],[f1108_D])).
fof(f1108_D,plain,(
  ( ! [X46,X49] : (~p(s(bool,sK2(X49,X46))) | s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X46),s('type.typeSystem.tenv_val',sK1)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3694,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X39),s('type.typeSystem.tenv_val',sK1))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X40))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X40))) | ~p(s(bool,sK2(X41,X39)))) ) | $spl10),
  inference(superposition,[],[f3056,f136])).
fof(f3483,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3460,f127])).
fof(f3460,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X9))) ) | $spl10),
  inference(superposition,[],[f3054,f120])).
fof(f3054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',sK1))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X2)))) ) | $spl10),
  inference(superposition,[],[f3019,f115])).
fof(f3482,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3459,f127])).
fof(f3459,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,X5)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f3054,f120])).
fof(f1107,plain,(
  ( ! [X47,X48] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X48))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X47)))) ) | $spl30),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  ( ! [X47,X48] : (s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X48))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X47)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3437,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X46),s('type.typeSystem.tenv_val',sK1))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X48))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X47))) | ~p(s(bool,sK2(X49,X46)))) ) | $spl10),
  inference(superposition,[],[f3054,f136])).
fof(f3393,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(subsumption_resolution,[],[f3375,f131])).
fof(f3375,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,X25)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(superposition,[],[f2928,f123])).
fof(f2928,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl16),
  inference(superposition,[],[f1128,f116])).
fof(f3392,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20))))) ) | $spl16),
  inference(subsumption_resolution,[],[f3374,f131])).
fof(f3374,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20)))) | ~p(s(bool,X21))) ) | $spl16),
  inference(superposition,[],[f2928,f123])).
fof(f3084,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3070,f127])).
fof(f3070,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f3019,f120])).
fof(f3083,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3069,f127])).
fof(f3069,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f3019,f120])).
fof(f3020,plain,(
  $false | ($spl2 | $spl16 | $spl25)),
  inference(subsumption_resolution,[],[f2505,f1128])).
fof(f2505,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl2 | $spl25)),
  inference(forward_demodulation,[],[f204,f922])).
fof(f2588,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f2589,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f2593,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f2889,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X18),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X16))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X16))))) ) | $spl14),
  inference(subsumption_resolution,[],[f2869,f127])).
fof(f2869,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X18),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X19),s('type.typeSystem.tenv_val',X16))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X16)))) | p(s(bool,X17))) ) | $spl14),
  inference(superposition,[],[f2554,f120])).
fof(f2554,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl14),
  inference(superposition,[],[f1905,f116])).
fof(f2887,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X14),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X12))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X12))))) ) | $spl14),
  inference(subsumption_resolution,[],[f2868,f127])).
fof(f2868,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X14),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X12))) | p(s(bool,X13)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X12))))) ) | $spl14),
  inference(superposition,[],[f2554,f120])).
fof(f2709,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f2696,f131])).
fof(f2696,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(superposition,[],[f2596,f123])).
fof(f2708,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f2695,f131])).
fof(f2695,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f2596,f123])).
fof(f2592,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f2591,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f2590,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f2587,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f2586,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f2585,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f2584,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f1997,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f1998,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f1999,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f2497,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2477,f131])).
fof(f2477,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X26),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X27),s('type.typeSystem.tenv_val',X24))) | ~p(s(bool,X25)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X24))))) ) | $spl16),
  inference(superposition,[],[f1960,f123])).
fof(f1960,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X14))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X16),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X15),s('type.typeSystem.tenv_val',X14)))) ) | $spl16),
  inference(superposition,[],[f1128,f116])).
fof(f2495,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2476,f131])).
fof(f2476,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X22),s('type.typeSystem.tenv_val',sK1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X23),s('type.typeSystem.tenv_val',X20))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X20)))) | ~p(s(bool,X21))) ) | $spl16),
  inference(superposition,[],[f1960,f123])).
fof(f2152,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2138,f127])).
fof(f2138,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f2045,f120])).
fof(f2045,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl10),
  inference(backward_demodulation,[],[f210,f114])).
fof(f2151,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2137,f127])).
fof(f2137,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f2045,f120])).
fof(f2036,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f2028])).
fof(f2028,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl28),
  inference(superposition,[],[f618,f150])).
fof(f150,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))),
  inference(cnf_transformation,[],[f28])).
fof(f28,axiom,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Empty_0')))),
  file('thm.typeSysProps.bind_tvar_rewrites|split|3','thm.typeSoundInvariants.tenv_val_ok_def|split|0')).
fof(f2035,plain,(
  $false | $spl28),
  inference(equality_resolution,[],[f618])).
fof(f2000,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f1996,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f1995,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f1994,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f1993,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f1992,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f1991,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f1975,plain,(
  $false | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f237,f1472])).
fof(f1913,plain,(
  $false | ($spl4 | $spl14)),
  inference(subsumption_resolution,[],[f1912,f175])).
fof(f1912,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl14)),
  inference(forward_demodulation,[],[f207,f215])).
fof(f1875,plain,(
  $false | ($spl24 | $spl26)),
  inference(subsumption_resolution,[],[f1811,f175])).
fof(f1811,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl26)),
  inference(backward_demodulation,[],[f390,f432])).
fof(f1677,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1665,f131])).
fof(f1665,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(superposition,[],[f1644,f123])).
fof(f1644,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl12),
  inference(forward_demodulation,[],[f114,f212])).
fof(f1676,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1664,f131])).
fof(f1664,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f1644,f123])).
fof(f1508,plain,(
  $false | ($spl0 | $spl20)),
  inference(trivial_inequality_removal,[],[f1502])).
fof(f1502,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl20)),
  inference(superposition,[],[f238,f439])).
fof(f1507,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f238])).
fof(f1421,plain,(
  $false | ($spl8 | $spl12)),
  inference(trivial_inequality_removal,[],[f1411])).
fof(f1411,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl8 | $spl12)),
  inference(superposition,[],[f1384,f150])).
fof(f1384,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f1383,f133])).
fof(f1383,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl8 | $spl12)),
  inference(forward_demodulation,[],[f209,f212])).
fof(f1420,plain,(
  $false | ($spl8 | $spl12)),
  inference(trivial_inequality_removal,[],[f1413])).
fof(f1413,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl8 | $spl12)),
  inference(superposition,[],[f1384,f212])).
fof(f1419,plain,(
  $false | ($spl8 | $spl12)),
  inference(equality_resolution,[],[f1384])).
fof(f1301,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1289,f131])).
fof(f1289,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(superposition,[],[f1266,f123])).
fof(f1266,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl12),
  inference(forward_demodulation,[],[f114,f212])).
fof(f1300,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1288,f131])).
fof(f1288,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f1266,f123])).
fof(f1267,plain,(
  $false | ($spl3 | $spl12 | $spl25)),
  inference(subsumption_resolution,[],[f1127,f212])).
fof(f1127,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ($spl3 | $spl25)),
  inference(forward_demodulation,[],[f205,f922])).
fof(f640,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f639,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f638,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f637,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f636,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f635,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f643,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f642,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f641,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f1138,plain,(
  $false | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f1137,f162])).
fof(f1137,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl16)),
  inference(forward_demodulation,[],[f206,f217])).
fof(f1136,plain,(
  $false | ($spl14 | $spl16 | $spl23)),
  inference(subsumption_resolution,[],[f1135,f241])).
fof(f1135,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f215,f217])).
fof(f1118,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1094,f127])).
fof(f1094,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X10))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',X11))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X8),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X9))) ) | $spl10),
  inference(superposition,[],[f769,f120])).
fof(f769,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',sK1))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X1))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X0),s('type.typeSystem.tenv_val',X2)))) ) | $spl10),
  inference(superposition,[],[f693,f115])).
fof(f693,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl10),
  inference(backward_demodulation,[],[f210,f114])).
fof(f1116,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1093,f127])).
fof(f1093,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X6))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',X7))) | p(s(bool,X5)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X4),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f769,f120])).
fof(f1071,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',X46),s('type.typeSystem.tenv_val',sK1))))) | s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X48))) != s('type.typeSystem.tenv_val','const.typeSystem.Bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',X47))) | ~p(s(bool,sK2(X49,X46)))) ) | $spl10),
  inference(superposition,[],[f769,f136])).
fof(f432,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | $spl26),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f926,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl25),
  inference(subsumption_resolution,[],[f913,f127])).
fof(f913,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) ) | $spl25),
  inference(superposition,[],[f391,f120])).
fof(f925,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl25),
  inference(subsumption_resolution,[],[f912,f127])).
fof(f912,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl25),
  inference(superposition,[],[f391,f120])).
fof(f644,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f619,plain,(
  $false | ($spl11 | $spl13)),
  inference(subsumption_resolution,[],[f232,f213])).
fof(f232,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl11),
  inference(trivial_inequality_removal,[],[f231])).
fof(f231,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | $spl11),
  inference(superposition,[],[f211,f135])).
fof(f336,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl13),
  inference(subsumption_resolution,[],[f325,f131])).
fof(f325,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | ~p(s(bool,X2))) ) | $spl13),
  inference(superposition,[],[f213,f123])).
fof(f338,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl13),
  inference(subsumption_resolution,[],[f326,f131])).
fof(f326,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl13),
  inference(superposition,[],[f213,f123])).
fof(f327,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl13),
  inference(superposition,[],[f213,f128])).
fof(f328,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl13),
  inference(superposition,[],[f213,f130])).
fof(f600,plain,(
  $false | ($spl6 | $spl10)),
  inference(trivial_inequality_removal,[],[f592])).
fof(f592,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl10)),
  inference(superposition,[],[f563,f210])).
fof(f563,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f562,f129])).
fof(f562,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f208,f210])).
fof(f599,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(trivial_inequality_removal,[],[f593])).
fof(f593,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl6 | $spl10)),
  inference(superposition,[],[f563,f439])).
fof(f598,plain,(
  $false | ($spl6 | $spl10)),
  inference(equality_resolution,[],[f563])).
fof(f407,plain,(
  ( ! [X10] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,sK2(X10,sK0)))) ) | $spl4),
  inference(superposition,[],[f207,f136])).
fof(f393,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f378,f127])).
fof(f378,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f243,f120])).
fof(f243,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))) | $spl10),
  inference(backward_demodulation,[],[f210,f114])).
fof(f392,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f377,f127])).
fof(f377,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) ) | $spl10),
  inference(superposition,[],[f243,f120])).
fof(f391,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl25),
  inference(cnf_transformation,[],[f391_D])).
fof(f391_D,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f366,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X10,sK0)))) ) | $spl10),
  inference(superposition,[],[f243,f136])).
fof(f229,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl11),
  inference(superposition,[],[f211,f134])).
fof(f228,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | $spl11),
  inference(superposition,[],[f211,f132])).
fof(f239,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl11),
  inference(subsumption_resolution,[],[f223,f127])).
fof(f223,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1)))) | p(s(bool,X1))) ) | $spl11),
  inference(superposition,[],[f211,f120])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl11),
  inference(subsumption_resolution,[],[f222,f127])).
fof(f222,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))))) ) | $spl11),
  inference(superposition,[],[f211,f120])).
fof(f202,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f201,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))),
  inference(superposition,[],[f114,f135])).
fof(f200,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f134])).
fof(f199,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f132])).
fof(f198,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f130])).
fof(f197,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))),
  inference(superposition,[],[f114,f128])).
fof(f196,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,X3)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f123])).
fof(f195,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f114,f123])).
fof(f194,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f114,f120])).
fof(f193,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) | p(s(bool,X0)) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK1))))))) )),
  inference(superposition,[],[f114,f120])).
fof(f205,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | $spl3),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f182,plain,(
  ( ! [X10] : (s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK1))) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_tvar_2'(s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK1))))) | ~p(s(bool,sK2(X10,sK0)))) )),
  inference(superposition,[],[f114,f136])).
% SZS output end Proof for thm.typeSysProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 314
% Passive clauses: 4719
% Generated clauses: 16717
% Final active clauses: 148
% Final passive clauses: 2282
% Input formulas: 46
% Initial clauses: 62
% 
% Selected by SInE selection: 33
% SInE iterations: 1
% 
% Duplicate literals: 230
% Trivial inequalities: 575
% Fw subsumption resolutions: 1182
% Fw demodulations: 1708
% Bw demodulations: 82
% 
% Simple tautologies: 479
% Equational tautologies: 5
% Forward subsumptions: 8051
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 1226
% Factoring: 1
% Forward superposition: 7363
% Backward superposition: 3175
% Self superposition: 27
% Equality factoring: 3
% Equality resolution: 28
% 
% Split clauses: 320
% Split components: 640
% SAT solver clauses: 364
% SAT solver unit clauses: 11
% SAT solver binary clauses: 123
% 
% Sat splits: 320
% Sat splitting refutations: 43
% 
% Memory used [KB]: 6396
% Time elapsed: 0.204 s
% ------------------------------
% ------------------------------
% Success in time 11.94 s
