Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (15737)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 160
% Passive clauses: 3691
% Generated clauses: 19773
% Final active clauses: 140
% Final passive clauses: 3278
% Input formulas: 68
% Initial clauses: 101
% 
% Duplicate literals: 569
% Trivial inequalities: 333
% Fw subsumption resolutions: 1496
% Fw demodulations: 2722
% Bw demodulations: 18
% 
% Simple tautologies: 257
% Equational tautologies: 44
% Forward subsumptions: 10174
% 
% Binary resolution: 413
% Factoring: 1
% Forward superposition: 6677
% Backward superposition: 7089
% Self superposition: 175
% Equality factoring: 72
% Equality resolution: 31
% 
% Split clauses: 223
% Split components: 448
% SAT solver clauses: 609
% SAT solver unit clauses: 9
% SAT solver binary clauses: 117
% 
% Sat splits: 223
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 608
% TWLsolver calls for satisfiability: 168
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (15757)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 161
% Passive clauses: 1594
% Generated clauses: 15652
% Final active clauses: 115
% Final passive clauses: 887
% Input formulas: 68
% Initial clauses: 101
% 
% Duplicate literals: 406
% Trivial inequalities: 136
% Fw subsumption resolutions: 742
% Fw demodulations: 1731
% Bw demodulations: 477
% Condensations: 876
% Global subsumptions: 8
% 
% Simple tautologies: 171
% Equational tautologies: 40
% Forward subsumptions: 7741
% Bw demodulations to eq. taut.: 101
% 
% Binary resolution: 371
% Factoring: 1
% Forward superposition: 5003
% Backward superposition: 3934
% Self superposition: 126
% Equality factoring: 62
% Equality resolution: 20
% 
% Split clauses: 93
% Split components: 186
% SAT solver clauses: 2179
% SAT solver unit clauses: 151
% SAT solver binary clauses: 52
% SAT solver learnt clauses: 16
% 
% Sat splits: 93
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 2147
% TWLsolver calls for satisfiability: 401
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f103933,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f228,f103929])).
fof(f103929,plain,(
  ( ! [X653] : (~p(s(bool,X653))) )),
  inference(forward_demodulation,[],[f103928,f311])).
fof(f311,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X18] : ! [X22,X35,X36,X37,X38] : s(X18,X37) = s(X18,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X18)),X22),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X18)),X35),s(X18,X36),s(X18,X37),s(X18,X38)))),
  file('thm.wordSem.result_distinct|split|2','thm.wordSem.result_case_def|split|3')).
fof(f103928,plain,(
  ( ! [X652,X654,X650,X653,X651] : (~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X650),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X651),s(bool,X652),s(bool,X653),s(bool,X654))))) )),
  inference(subsumption_resolution,[],[f103769,f103689])).
fof(f103689,plain,(
  ( ! [X6,X10,X7] : (s(X6,X10) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X6)),X7),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f20856,f311])).
fof(f20856,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(X0,X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f314,f227])).
fof(f227,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f183])).
fof(f183,plain,(
  ? [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.wordSem.result_distinct|split|2','thm.wordSem.result_distinct|split|2')).
fof(f314,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X18,X6] : ! [X8,X7,X22,X35,X36,X37,X38] : s(X18,happ(s(fun('type.wordSem.word_loc'(X6),X18),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X18)),X22),s('type.wordSem.word_loc'(X6),X8))),s('type.wordSem.word_loc'(X6),X7))) = s(X18,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X18)),X22),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X18)),X35),s(X18,X36),s(X18,X37),s(X18,X38)))),
  file('thm.wordSem.result_distinct|split|2','thm.wordSem.result_case_def|split|0')).
fof(f103769,plain,(
  ( ! [X652,X654,X650,X653,X651] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X650),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X650),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X651),s(bool,X652),s(bool,X653),s(bool,X654))))) )),
  inference(superposition,[],[f242,f20856])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X27] : (s(bool,'F_0') = s(bool,X27) <=> ~p(s(bool,X27)))),
  file('thm.wordSem.result_distinct|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f228,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.wordSem.result_distinct|split|2',aHL_TRUTH)).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 348
% Passive clauses: 8003
% Generated clauses: 108619
% Final active clauses: 194
% Final passive clauses: 1345
% Input formulas: 68
% Initial clauses: 101
% Discarded non-redundant clauses: 3683
% 
% Duplicate literals: 1842
% Trivial inequalities: 561
% Fw subsumption resolutions: 6748
% Fw demodulations: 24587
% Bw demodulations: 4528
% Condensations: 4826
% Global subsumptions: 208
% 
% Simple tautologies: 611
% Equational tautologies: 520
% Forward subsumptions: 56553
% Fw demodulations to eq. taut.: 2125
% Bw demodulations to eq. taut.: 199
% 
% Unit resulting resolution: 12052
% Factoring: 1
% Forward superposition: 25098
% Backward superposition: 22116
% Self superposition: 434
% Equality factoring: 187
% Equality resolution: 55
% 
% Split clauses: 52
% Split components: 104
% SAT solver clauses: 9266
% SAT solver unit clauses: 1920
% SAT solver binary clauses: 3790
% 
% Sat splits: 52
% Sat splitting refutations: 949
% 
% Memory used [KB]: 89550
% Time elapsed: 6.585 s
% ------------------------------
% ------------------------------
% Success in time 7.92 s
