Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (61940)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 169
% Passive clauses: 2299
% Generated clauses: 16992
% Final active clauses: 160
% Final passive clauses: 1960
% Input formulas: 69
% Initial clauses: 341
% 
% Duplicate literals: 527
% Trivial inequalities: 174
% Fw subsumption resolutions: 2633
% Fw demodulations: 2354
% Bw demodulations: 5
% 
% Simple tautologies: 229
% Equational tautologies: 73
% Forward subsumptions: 7988
% 
% Binary resolution: 249
% Factoring: 2
% Forward superposition: 6163
% Backward superposition: 5724
% Self superposition: 350
% Equality factoring: 179
% Equality resolution: 41
% 
% Split clauses: 471
% Split components: 943
% SAT solver clauses: 879
% SAT solver unit clauses: 5
% SAT solver binary clauses: 78
% 
% Sat splits: 471
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 879
% TWLsolver calls for satisfiability: 157
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (61968)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 102
% Passive clauses: 766
% Generated clauses: 13306
% Final active clauses: 99
% Final passive clauses: 526
% Input formulas: 69
% Initial clauses: 341
% 
% Duplicate literals: 367
% Trivial inequalities: 90
% Fw subsumption resolutions: 1492
% Fw demodulations: 1265
% Bw demodulations: 139
% Condensations: 2038
% Global subsumptions: 10
% 
% Simple tautologies: 199
% Equational tautologies: 36
% Forward subsumptions: 6001
% 
% Binary resolution: 36
% Factoring: 1
% Forward superposition: 3237
% Backward superposition: 3403
% Self superposition: 219
% Equality factoring: 78
% Equality resolution: 37
% 
% Split clauses: 85
% Split components: 172
% SAT solver clauses: 1317
% SAT solver unit clauses: 352
% SAT solver binary clauses: 146
% SAT solver learnt clauses: 20
% 
% Sat splits: 85
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1277
% TWLsolver calls for satisfiability: 4674
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% (61998)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 541
% Passive clauses: 10524
% Generated clauses: 141748
% Final active clauses: 205
% Final passive clauses: 1309
% Input formulas: 69
% Initial clauses: 341
% Discarded non-redundant clauses: 2371
% 
% Duplicate literals: 4668
% Trivial inequalities: 983
% Fw subsumption resolutions: 13195
% Fw demodulations: 14749
% Bw demodulations: 6246
% Condensations: 17061
% Global subsumptions: 44
% 
% Simple tautologies: 1313
% Equational tautologies: 572
% Forward subsumptions: 77644
% Fw demodulations to eq. taut.: 45
% Bw demodulations to eq. taut.: 730
% 
% Unit resulting resolution: 3613
% Factoring: 1
% Forward superposition: 36596
% Backward superposition: 34008
% Self superposition: 692
% Equality factoring: 326
% Equality resolution: 119
% 
% Split clauses: 463
% Split components: 967
% SAT solver clauses: 11667
% SAT solver unit clauses: 111
% SAT solver binary clauses: 1011
% 
% Sat splits: 463
% Sat splitting refutations: 171
% 
% Memory used [KB]: 74455
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f1422,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f346,f672,f670,f345,f671,f669,f907,f952,f953,f974,f975,f963,f999,f1009,f1010,f1012,f1014,f965,f1024,f1035,f1037,f1040,f1043,f967,f1052,f1062,f1074,f1077,f1081,f1085,f971,f1098,f961,f1102,f969,f973,f1421])).
fof(f1421,plain,(
  $false | ($spl3 | $spl4 | $spl16)),
  inference(trivial_inequality_removal,[],[f1420])).
fof(f1420,plain,(
  s('type.conSem.v',sK26) != s('type.conSem.v',sK26) | ($spl3 | $spl4 | $spl16)),
  inference(superposition,[],[f1356,f961])).
fof(f1356,plain,(
  ( ! [X2] : (s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2)))) ) | ($spl3 | $spl4 | $spl16)),
  inference(subsumption_resolution,[],[f1355,f908])).
fof(f908,plain,(
  ( ! [X0] : (s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0)))) ) | $spl3),
  inference(superposition,[],[f672,f368])).
fof(f368,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) = s('type.ast.lit',X1) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : ((s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X0))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X57,X58] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X58))) <=> s('type.ast.lit',X57) = s('type.ast.lit',X58))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_|49||49||split|0')).
fof(f1355,plain,(
  ( ! [X2] : (s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2)))) ) | ($spl4 | $spl16)),
  inference(forward_demodulation,[],[f1266,f961])).
fof(f1266,plain,(
  ( ! [X2] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28))))) ) | $spl16),
  inference(superposition,[],[f973,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) = s('type.ast.lit',X1) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : ((s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X57,X58] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X57))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X58))) <=> s('type.ast.lit',X57) = s('type.ast.lit',X58))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.conSem.v_|49||49||split|0')).
fof(f973,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl16),
  inference(cnf_transformation,[],[f973_D])).
fof(f973_D,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f969,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl12),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1102,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1101,f924])).
fof(f924,plain,(
  sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl0),
  inference(resolution,[],[f669,f581])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X2,X1) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | sP25(X0,X2,X1)) & (~sP25(X0,X2,X1) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> sP25(X0,X2,X1))),
  inference(definition_folding,[],[f109,f197,f196,f195,f194,f193])).
fof(f193,plain,(
  ! [X0,X2,X1] : (sP21(X0,X2,X1) <=> ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP21])])).
fof(f194,plain,(
  ! [X0,X2,X1] : (sP22(X0,X2,X1) <=> ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f195,plain,(
  ! [X0,X2,X1] : (sP23(X0,X2,X1) <=> ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f196,plain,(
  ! [X0,X2,X1] : (sP24(X0,X2,X1) <=> ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP24])])).
fof(f197,plain,(
  ! [X0,X2,X1] : (sP25(X0,X2,X1) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP23(X0,X2,X1) | sP21(X0,X2,X1) | sP22(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP24(X0,X2,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP25])])).
fof(f109,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X34,X35,X36] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34),s('type.conSem.v',X35),s('type.exhSem.v',X36)))) <=> (? [X7] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X17,X18,X19] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.conSem.v'),X18))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X17))),s('type.list.list'('type.exhSem.v'),X19))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34),s('type.list.list'('type.conSem.v'),X18),s('type.list.list'('type.exhSem.v'),X19))))) | ? [X20,X21,X22,X23,X24] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.string.char'),X21),s('type.conLang.exp',X22))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X21),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.conLang.exp',X22))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X23,X20,X21,X25,X24] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25),s('type.list.list'('type.string.char'),X21))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25))),s('type.list.list'('type.string.char'),X21))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X7] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7)))) | ? [X18,X19] : (s('type.conSem.v',X35) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X18))) & s('type.exhSem.v',X36) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X19))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X34),s('type.list.list'('type.conSem.v'),X18),s('type.list.list'('type.exhSem.v'),X19)))))))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.dec_to_exhProof.v_rel_cases|split|0')).
fof(f1101,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1100,f962])).
fof(f962,plain,(
  ~sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(cnf_transformation,[],[f962_D])).
fof(f962_D,plain,(
  ~sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1100,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1099,f964])).
fof(f964,plain,(
  ~sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(cnf_transformation,[],[f964_D])).
fof(f964_D,plain,(
  ~sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1099,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1088,f966])).
fof(f966,plain,(
  ~sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(cnf_transformation,[],[f966_D])).
fof(f966_D,plain,(
  ~sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1088,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl15),
  inference(resolution,[],[f970,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0)))) | sP22(X0,X1,X2) | sP21(X0,X1,X2) | sP23(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0)))) | ~sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0))))) | sP23(X0,X1,X2) | sP21(X0,X1,X2) | sP22(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0))))) | sP24(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP23(X0,X1,X2) & ~sP21(X0,X1,X2) & ~sP22(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP24(X0,X1,X2)) | sP25(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK143,sK144])],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | ? [X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X6))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X6)))) | sP23(X0,X1,X2) | sP21(X0,X1,X2) | sP22(X0,X1,X2) | ? [X5] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP24(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP23(X0,X1,X2) & ~sP21(X0,X1,X2) & ~sP22(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP24(X0,X1,X2)) | sP25(X0,X1,X2)))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X0,X2,X1] : ((~sP25(X0,X2,X1) | ? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP23(X0,X2,X1) | sP21(X0,X2,X1) | sP22(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP24(X0,X2,X1)) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP23(X0,X2,X1) & ~sP21(X0,X2,X1) & ~sP22(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP24(X0,X2,X1)) | sP25(X0,X2,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X2,X1] : ((~sP25(X0,X2,X1) | (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP23(X0,X2,X1) | sP21(X0,X2,X1) | sP22(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP24(X0,X2,X1))) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP23(X0,X2,X1) & ~sP21(X0,X2,X1) & ~sP22(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP24(X0,X2,X1)) | sP25(X0,X2,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f970,plain,(
  ~sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl15),
  inference(cnf_transformation,[],[f970_D])).
fof(f970_D,plain,(
  ~sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f961,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl4),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1098,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1097,f924])).
fof(f1097,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1096,f962])).
fof(f1096,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1095,f964])).
fof(f1095,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl11 | $spl15)),
  inference(subsumption_resolution,[],[f1087,f966])).
fof(f1087,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl15),
  inference(resolution,[],[f970,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0)))) | sP22(X0,X1,X2) | sP21(X0,X1,X2) | sP23(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0)))) | ~sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f292])).
fof(f971,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl14),
  inference(cnf_transformation,[],[f971_D])).
fof(f971_D,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1085,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1084,f924])).
fof(f1084,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1083,f962])).
fof(f1083,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1082,f964])).
fof(f1082,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(subsumption_resolution,[],[f1067,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X58,X57] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X58)))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_distinct|split|3')).
fof(f1067,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(resolution,[],[f966,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0)))) | sP22(X0,X1,X2) | sP21(X0,X1,X2) | sP23(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0)))) | ~sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f292])).
fof(f1081,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1080,f924])).
fof(f1080,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1079,f962])).
fof(f1079,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1078,f964])).
fof(f1078,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(subsumption_resolution,[],[f1066,f363])).
fof(f1066,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(resolution,[],[f966,f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(X2,X1,X0)))) | sP22(X0,X1,X2) | sP21(X0,X1,X2) | sP23(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(X2,X1,X0)))) | ~sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f292])).
fof(f1077,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1076,f924])).
fof(f1076,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1075,f962])).
fof(f1075,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1065,f964])).
fof(f1065,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(resolution,[],[f966,f554])).
fof(f1074,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1073,f924])).
fof(f1073,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1072,f962])).
fof(f1072,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f1064,f964])).
fof(f1064,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl11),
  inference(resolution,[],[f966,f553])).
fof(f1062,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1054,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X36,X35,X34,X57] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X34),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),X35),s('type.list.list'('type.string.char'),X36)))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_distinct|split|2')).
fof(f1054,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK150(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK154(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK153(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))),s('type.list.list'('type.string.char'),sK152(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl10),
  inference(resolution,[],[f967,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK150(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK154(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK153(X2,X1,X0)))),s('type.list.list'('type.string.char'),sK152(X2,X1,X0)))) | ~sP22(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK151(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK153(X2,X1,X0)),s('type.list.list'('type.string.char'),sK152(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK150(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK154(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK153(X2,X1,X0)))),s('type.list.list'('type.string.char'),sK152(X2,X1,X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK154(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK151(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK150(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP22(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK150,sK151,sK152,sK153,sK154])],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11),s('type.list.list'('type.string.char'),X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11))),s('type.list.list'('type.string.char'),X10))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP22(X0,X1,X2)))),
  inference(rectify,[],[f299])).
fof(f299,plain,(
  ! [X0,X2,X1] : ((~sP22(X0,X2,X1) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6)))))) & (! [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP22(X0,X2,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f1052,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1045,f361])).
fof(f361,plain,(
  ( ! [X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X58,X57] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X58)))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_distinct|split|4')).
fof(f1045,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK146(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl14),
  inference(resolution,[],[f971,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK146(X2,X1,X0)))) | ~sP24(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK145(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK146(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK145(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK146(X2,X1,X0))))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP24(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK145,sK146])],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | ? [X5,X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP24(X0,X1,X2)))),
  inference(rectify,[],[f293])).
fof(f293,plain,(
  ! [X0,X2,X1] : ((~sP24(X0,X2,X1) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))) & (! [X3,X4] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP24(X0,X2,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f967,plain,(
  sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl10),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1043,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1042,f924])).
fof(f1042,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1041,f962])).
fof(f1041,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(subsumption_resolution,[],[f1029,f363])).
fof(f1029,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(resolution,[],[f964,f556])).
fof(f1040,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1039,f924])).
fof(f1039,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1038,f962])).
fof(f1038,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(subsumption_resolution,[],[f1028,f363])).
fof(f1028,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(resolution,[],[f964,f555])).
fof(f1037,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1036,f924])).
fof(f1036,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1027,f962])).
fof(f1027,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(resolution,[],[f964,f554])).
fof(f1035,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1034,f924])).
fof(f1034,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1026,f962])).
fof(f1026,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl9),
  inference(resolution,[],[f964,f553])).
fof(f1024,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1016,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X3))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.exhLang.exp',X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X36,X35,X34,X57] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X34),s('type.list.list'('type.string.char'),X35),s('type.exhLang.exp',X36)))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_distinct|split|1')).
fof(f1016,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK158(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)),s('type.list.list'('type.string.char'),sK156(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK159(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)),s('type.conLang.exp',sK157(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))))) | $spl8),
  inference(resolution,[],[f965,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK158(X2,X1,X0)),s('type.list.list'('type.string.char'),sK156(X2,X1,X0)),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK159(X2,X1,X0)),s('type.conLang.exp',sK157(X2,X1,X0)))))) | ~sP21(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK155(X2,X1,X0)),s('type.list.list'('type.string.char'),sK156(X2,X1,X0)),s('type.conLang.exp',sK157(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK158(X2,X1,X0)),s('type.list.list'('type.string.char'),sK156(X2,X1,X0)),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK159(X2,X1,X0)),s('type.conLang.exp',sK157(X2,X1,X0)))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK159(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK155(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK158(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP21(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK155,sK156,sK157,sK158,sK159])],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.string.char'),X9),s('type.conLang.exp',X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11),s('type.list.list'('type.string.char'),X9),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.conLang.exp',X10))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP21(X0,X1,X2)))),
  inference(rectify,[],[f302])).
fof(f302,plain,(
  ! [X0,X2,X1] : ((~sP21(X0,X2,X1) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14)))))) & (! [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | sP21(X0,X2,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f965,plain,(
  sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl8),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1014,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f1013,f924])).
fof(f1013,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(subsumption_resolution,[],[f1004,f363])).
fof(f1004,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(resolution,[],[f962,f556])).
fof(f1012,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f1011,f924])).
fof(f1011,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(subsumption_resolution,[],[f1003,f363])).
fof(f1003,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(resolution,[],[f962,f555])).
fof(f1010,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f1002,f924])).
fof(f1002,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(resolution,[],[f962,f554])).
fof(f1009,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f1001,f924])).
fof(f1001,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl7),
  inference(resolution,[],[f962,f553])).
fof(f999,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f992,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.exhSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.exhSem.v'),X0)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.exhSem.v'),X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X35,X34,X57] : ~s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X57))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num',X34),s('type.list.list'('type.exhSem.v'),X35)))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.exhSem.v_distinct|split|0')).
fof(f992,plain,(
  s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK147(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))),s('type.list.list'('type.exhSem.v'),sK149(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl6),
  inference(resolution,[],[f963,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK147(X2,X1,X0)))),s('type.list.list'('type.exhSem.v'),sK149(X2,X1,X0)))) | ~sP23(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK147(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK148(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK147(X2,X1,X0)))),s('type.list.list'('type.exhSem.v'),sK149(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK148(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK149(X2,X1,X0))))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP23(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK147,sK148,sK149])],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | ? [X6,X7,X8] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6),s('type.list.list'('type.conSem.v'),X7))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6))),s('type.list.list'('type.exhSem.v'),X8))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X8)))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP23(X0,X1,X2)))),
  inference(rectify,[],[f296])).
fof(f296,plain,(
  ! [X0,X2,X1] : ((~sP23(X0,X2,X1) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18)))))) & (! [X16,X17,X18] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | sP23(X0,X2,X1)))),
  inference(nnf_transformation,[],[f195])).
fof(f963,plain,(
  sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl6),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f975,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(subsumption_resolution,[],[f955,f363])).
fof(f955,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(resolution,[],[f924,f556])).
fof(f974,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(subsumption_resolution,[],[f954,f363])).
fof(f954,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(resolution,[],[f924,f555])).
fof(f953,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(resolution,[],[f924,f554])).
fof(f952,plain,(
  sP24(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK144(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | sP22(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP21(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | sP23(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK143(sK26,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK28)))) | $spl0),
  inference(resolution,[],[f924,f553])).
fof(f907,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f906])).
fof(f906,plain,(
  s('type.conSem.v',sK26) != s('type.conSem.v',sK26) | s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))) | ($spl1 | $spl2)),
  inference(resolution,[],[f688,f785])).
fof(f785,plain,(
  ~sP25(sK28,'const.exhSem.Litv_1'(s('type.ast.lit',sK27)),sK26) | $spl1),
  inference(resolution,[],[f670,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ~sP25(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f305])).
fof(f688,plain,(
  ( ! [X33,X31,X32] : (sP25(X32,X33,X31) | s('type.conSem.v',sK26) != s('type.conSem.v',X31) | s('type.exhSem.v',X33) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))) ) | $spl2),
  inference(superposition,[],[f557,f671])).
fof(f557,plain,(
  ( ! [X4,X2,X0,X1] : (sP25(X0,X1,X2) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4))) | s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4)))) )),
  inference(cnf_transformation,[],[f292])).
fof(f669,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) | $spl0),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f671,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) | $spl2),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f345,plain,(
  s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))))))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) | s('type.conSem.v',sK26) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) | s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27,sK28])],[f199])).
fof(f199,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v',X0),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.conSem.v',X0) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v',X0),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.conSem.v',X0) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v',X0),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))) <~> s('type.conSem.v',X0) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v',X0),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))) <=> s('type.conSem.v',X0) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.conSem.v',X6),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.conSem.v',X6) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.conSem.v',X6),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.conSem.v',X6) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|1','thm.dec_to_exhProof.v_rel_lit_loc|split|1')).
fof(f670,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) | $spl1),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f672,plain,(
  s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) | $spl3),
  inference(cnf_transformation,[],[f672_D])).
fof(f672_D,plain,(
  s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f346,plain,(
  s('type.conSem.v',sK26) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK27))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK28),s('type.conSem.v',sK26),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK27))))))),
  inference(cnf_transformation,[],[f200])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 302
% Passive clauses: 316
% Generated clauses: 750
% Final active clauses: 288
% Final passive clauses: 181
% Input formulas: 69
% Initial clauses: 322
% 
% Selected by SInE selection: 58
% SInE iterations: 1
% 
% Duplicate literals: 6
% Trivial inequalities: 5
% Fw subsumption resolutions: 56
% Fw demodulations: 143
% 
% Simple tautologies: 53
% Forward subsumptions: 165
% 
% Binary resolution: 103
% Forward superposition: 194
% Backward superposition: 228
% Equality resolution: 2
% 
% Split clauses: 20
% Split components: 72
% SAT solver clauses: 27
% SAT solver unit clauses: 4
% SAT solver binary clauses: 3
% 
% Sat splits: 20
% Sat splitting refutations: 6
% 
% Memory used [KB]: 3582
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
% Success in time 11.785 s
