Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modSem
% (139263)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 3254
% Generated clauses: 22358
% Final active clauses: 109
% Final passive clauses: 2869
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 808
% Trivial inequalities: 107
% Fw subsumption resolutions: 1836
% Fw demodulations: 3577
% Bw demodulations: 48
% 
% Simple tautologies: 131
% Equational tautologies: 49
% Forward subsumptions: 12060
% 
% Binary resolution: 496
% Factoring: 2
% Forward superposition: 7815
% Backward superposition: 7118
% Self superposition: 237
% Equality factoring: 96
% Equality resolution: 9
% 
% Split clauses: 241
% Split components: 482
% SAT solver clauses: 704
% SAT solver unit clauses: 34
% SAT solver binary clauses: 133
% 
% Sat splits: 241
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 704
% TWLsolver calls for satisfiability: 148
% 
% Memory used [KB]: 20724
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modSem
% (139299)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 1287
% Generated clauses: 14311
% Final active clauses: 93
% Final passive clauses: 931
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 389
% Trivial inequalities: 23
% Fw subsumption resolutions: 704
% Fw demodulations: 1561
% Bw demodulations: 256
% Condensations: 889
% 
% Simple tautologies: 56
% Equational tautologies: 41
% Forward subsumptions: 7150
% 
% Binary resolution: 353
% Factoring: 1
% Forward superposition: 4853
% Backward superposition: 4321
% Self superposition: 198
% Equality factoring: 80
% Equality resolution: 1
% 
% Split clauses: 38
% Split components: 76
% SAT solver clauses: 1948
% SAT solver unit clauses: 131
% SAT solver binary clauses: 45
% 
% Sat splits: 38
% Sat splitting refutations: 31
% 
% TWLsolver clauses: 1948
% TWLsolver calls for satisfiability: 275
% 
% Memory used [KB]: 24178
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.modSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modSem
% SZS output start Proof for thm.modSem
fof(f14689,plain,(
  $false),
  inference(subsumption_resolution,[],[f14688,f215])).
fof(f215,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f30])).
fof(f30,axiom,(
  p(s(bool,'T_0'))),
  file('thm.modSem.state_fupdfupds_comp|split|5','thm.bool.TRUTH')).
fof(f14688,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f14686,f4139])).
fof(f4139,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f4138])).
fof(f4138,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f4066,f299])).
fof(f299,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f217,f218])).
fof(f218,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.modSem.state_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f217,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.modSem.state_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f4066,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f3383,f297])).
fof(f297,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f218])).
fof(f3383,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3283,f1312])).
fof(f1312,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1265,f264])).
fof(f264,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X8,X17] : ! [X12,X11,X18] : s(X16,'const.combin.o_3'(s(fun(X8,X16),X12),s(fun(X17,X8),X11),s(X17,X18))) = s(X16,happ(s(fun(X8,X16),X12),s(X8,happ(s(fun(X17,X8),X11),s(X17,X18)))))),
  file('thm.modSem.state_fupdfupds_comp|split|5','thm.combin.o_THM')).
fof(f1265,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f219,f261])).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X47,X48,X49] : (p(s(bool,happ(s(fun(X47,bool),happ(s(fun(X47,fun(X47,bool)),'=_0'),s(X47,X48))),s(X47,X49)))) <=> s(X47,X48) = s(X47,X49))),
  file('thm.modSem.state_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f219,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X36] : ! [X38] : s(X36,X38) = s(X36,happ(s(fun(X36,X36),'_i_0_0'),s(X36,X38)))),
  file('thm.modSem.state_fupdfupds_comp|split|5','_i0_ax')).
fof(f3283,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,happ(s(fun(X73,X75),X76),s(X73,X74))) = s(X75,'const.combin.o_3'(s(fun(X73,X75),X76),s(fun(X73,X73),'_i_0_0'),s(X73,X74)))) )),
  inference(superposition,[],[f264,f219])).
fof(f14686,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.modSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.modSem.state'(sK0)),fun(fun(sK2,'type.modSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))))),s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))))))),
  inference(backward_demodulation,[],[f14640,f4905])).
fof(f4905,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.modSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.modSem.state'(sK0)),fun(fun(sK2,'type.modSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))))),s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))))))))),
  inference(unit_resulting_resolution,[],[f214,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f214,plain,(
  s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))) != s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.modSem.state'(sK3)),sK4)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.modSem.state'(sK3)),sK4))) != s(fun(sK2,'type.modSem.state'(sK0)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK3),'type.modSem.state'(sK1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.modSem.state'(sK3)),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.modSem.state'(X3)),X4))) != s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.modSem.state'(X3)),X4)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.modSem.state'(X3)),X4))) = s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.modSem.state'(X3)),X4)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.modSem.state'(X3)),X4))) = s(fun(X2,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X3),'type.modSem.state'(X1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.modSem.state'(X3)),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.modSem.state'(X6)),'const.combin.o_2'(s(fun('type.modSem.state'(X9),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.modSem.state'(X9)),X10))) = s(fun(X8,'type.modSem.state'(X6)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X9),'type.modSem.state'(X7)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.modSem.state'(X9)),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.modSem.state'(X6)),'const.combin.o_2'(s(fun('type.modSem.state'(X9),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.modSem.state'(X9)),X10))) = s(fun(X8,'type.modSem.state'(X6)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X9),'type.modSem.state'(X7)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.modSem.state'(X9)),X10)))))),
  file('thm.modSem.state_fupdfupds_comp|split|5','thm.modSem.state_fupdfupds_comp|split|5')).
fof(f14640,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X5,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))))),s(fun(X5,'type.modSem.state'(X0)),X6))) = s(fun(X5,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X2),'type.modSem.state'(X1)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3))),s(fun(X5,'type.modSem.state'(X2)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X2)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))),s(fun(X5,'type.modSem.state'(X0)),X6)))))) )),
  inference(superposition,[],[f268,f263])).
fof(f263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X2),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.modSem.state'(X1),'type.modSem.state'(X2)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X2),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.modSem.state'(X1),'type.modSem.state'(X2)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.modSem.state'(X1),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X2),'type.modSem.state'(X0)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.modSem.state'(X1),'type.modSem.state'(X2)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X9,X7] : ! [X11,X12] : s(fun('type.modSem.state'(X9),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))) = s(fun('type.modSem.state'(X9),'type.modSem.state'(X6)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X6)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun('type.modSem.state'(X9),'type.modSem.state'(X7)),'const.modSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11)))))),
  file('thm.modSem.state_fupdfupds_comp|split|5','thm.modSem.state_fupdfupds_comp|split|4')).
fof(f268,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X16,X8,X22,X17] : ! [X12,X11,X10] : s(fun(X22,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X22,X8),'const.combin.o_2'(s(fun(X17,X8),X11),s(fun(X22,X17),X10))))) = s(fun(X22,X16),'const.combin.o_2'(s(fun(X17,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X17,X8),X11))),s(fun(X22,X17),X10)))),
  file('thm.modSem.state_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.modSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 137
% Passive clauses: 1577
% Generated clauses: 16223
% Final active clauses: 72
% Final passive clauses: 869
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 463
% Trivial inequalities: 37
% Fw subsumption resolutions: 493
% Fw demodulations: 3031
% Bw demodulations: 556
% Condensations: 748
% Global subsumptions: 373
% 
% Simple tautologies: 152
% Equational tautologies: 131
% Forward subsumptions: 8752
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 21
% 
% Unit resulting resolution: 646
% Factoring: 1
% Forward superposition: 4372
% Backward superposition: 3280
% Self superposition: 165
% Equality factoring: 68
% Equality resolution: 2
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 2401
% SAT solver unit clauses: 775
% SAT solver binary clauses: 831
% 
% Sat splits: 27
% Sat splitting refutations: 395
% 
% Memory used [KB]: 20212
% Time elapsed: 0.767 s
% ------------------------------
% ------------------------------
% Success in time 2.11 s
