Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (6373)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 3526
% Generated clauses: 19161
% Final active clauses: 154
% Final passive clauses: 3171
% Input formulas: 69
% Initial clauses: 101
% 
% Duplicate literals: 538
% Trivial inequalities: 303
% Fw subsumption resolutions: 1463
% Fw demodulations: 2630
% Bw demodulations: 18
% 
% Simple tautologies: 263
% Equational tautologies: 45
% Forward subsumptions: 9955
% 
% Binary resolution: 499
% Factoring: 1
% Forward superposition: 6626
% Backward superposition: 6668
% Self superposition: 178
% Equality factoring: 70
% Equality resolution: 30
% 
% Split clauses: 225
% Split components: 452
% SAT solver clauses: 621
% SAT solver unit clauses: 9
% SAT solver binary clauses: 132
% 
% Sat splits: 225
% 
% TWLsolver clauses: 621
% TWLsolver calls for satisfiability: 167
% 
% Memory used [KB]: 20340
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (6400)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1082
% Generated clauses: 15768
% Final active clauses: 138
% Final passive clauses: 854
% Input formulas: 69
% Initial clauses: 101
% 
% Duplicate literals: 264
% Trivial inequalities: 263
% Fw subsumption resolutions: 776
% Fw demodulations: 3398
% Bw demodulations: 82
% Condensations: 750
% Global subsumptions: 8
% 
% Simple tautologies: 110
% Equational tautologies: 41
% Forward subsumptions: 7040
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 129
% Factoring: 1
% Forward superposition: 4231
% Backward superposition: 4454
% Self superposition: 156
% Equality factoring: 72
% Equality resolution: 22
% 
% Split clauses: 67
% Split components: 134
% SAT solver clauses: 2165
% SAT solver unit clauses: 151
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 16
% 
% Sat splits: 67
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2133
% TWLsolver calls for satisfiability: 364
% 
% Memory used [KB]: 23411
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f56533,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f230,f56490])).
fof(f56490,plain,(
  ( ! [X434] : (~p(s(bool,X434))) )),
  inference(forward_demodulation,[],[f56489,f317])).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X16] : ! [X17,X31,X32,X33,X34] : s(X16,X34) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),X16),X17),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|8','thm.stackSem.result_case_def|split|4')).
fof(f56489,plain,(
  ( ! [X430,X432,X434,X431,X433] : (~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X430),s(fun('type.wordSem.word_loc'(sK0),bool),X431),s(fun('type.wordSem.word_loc'(sK0),bool),X432),s(bool,X433),s(bool,X434))))) )),
  inference(subsumption_resolution,[],[f56251,f56165])).
fof(f56165,plain,(
  ( ! [X6,X11,X9] : (s(X6,X11) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),X9),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f10837,f317])).
fof(f10837,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f320,f229])).
fof(f229,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Error_0') = s('type.stackSem.result'(sK0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Error_0') = s('type.stackSem.result'(sK0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f186])).
fof(f186,plain,(
  ? [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.Error_0') = s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.Error_0') != s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : ~s('type.stackSem.result'(X0),'const.stackSem.Error_0') = s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.Error_0') = s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.Error_0') = s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.stackSem.result_distinct|split|8','thm.stackSem.result_distinct|split|8')).
fof(f320,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16,X6] : ! [X7,X17,X31,X32,X33,X34] : s(X16,happ(s(fun('type.wordSem.word_loc'(X6),X16),X32),s('type.wordSem.word_loc'(X6),X7))) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),X16),X17),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|8','thm.stackSem.result_case_def|split|2')).
fof(f56251,plain,(
  ( ! [X430,X432,X434,X431,X433] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X432),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X430),s(fun('type.wordSem.word_loc'(sK0),bool),X431),s(fun('type.wordSem.word_loc'(sK0),bool),X432),s(bool,X433),s(bool,X434))))) )),
  inference(superposition,[],[f244,f10837])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X23] : (s(bool,'F_0') = s(bool,X23) <=> ~p(s(bool,X23)))),
  file('thm.stackSem.result_distinct|split|8','thm.bool.EQ_CLAUSES|split|2')).
fof(f230,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.result_distinct|split|8',aHL_TRUTH)).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 240
% Passive clauses: 3603
% Generated clauses: 57623
% Final active clauses: 168
% Final passive clauses: 1771
% Input formulas: 69
% Initial clauses: 101
% 
% Duplicate literals: 1262
% Trivial inequalities: 404
% Fw subsumption resolutions: 4516
% Fw demodulations: 10922
% Bw demodulations: 1561
% Condensations: 4750
% Global subsumptions: 69
% 
% Simple tautologies: 408
% Equational tautologies: 224
% Forward subsumptions: 31070
% Fw demodulations to eq. taut.: 88
% Bw demodulations to eq. taut.: 94
% 
% Unit resulting resolution: 2250
% Factoring: 1
% Forward superposition: 16230
% Backward superposition: 13544
% Self superposition: 380
% Equality factoring: 184
% Equality resolution: 43
% 
% Split clauses: 51
% Split components: 102
% SAT solver clauses: 3944
% SAT solver unit clauses: 770
% SAT solver binary clauses: 1603
% 
% Sat splits: 51
% Sat splitting refutations: 168
% 
% Memory used [KB]: 46054
% Time elapsed: 2.959 s
% ------------------------------
% ------------------------------
% Success in time 4.3 s
