Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f4777,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4776,f4673,f4702,f4700,f4698,f4696,f4747,f4715,f4767])).
fof(f4767,plain,(
  $false | $spl168),
  inference(subsumption_resolution,[],[f4761,f98])).
fof(f98,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(cnf_transformation,[],[f19])).
fof(f19,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.merge_def|split|4')).
fof(f4761,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl168),
  inference(backward_demodulation,[],[f4696,f94])).
fof(f94,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK0)))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f69])).
fof(f69,plain,(
  ? [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ~! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X6)))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_knownProof.merge_Other|split|0')).
fof(f4715,plain,(
  $false | $spl170),
  inference(subsumption_resolution,[],[f4704,f94])).
fof(f4704,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK0))) | $spl170),
  inference(superposition,[],[f104,f4698])).
fof(f104,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X42] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X42)))))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.merge_def|split|18')).
fof(f4747,plain,(
  $false | $spl172),
  inference(subsumption_resolution,[],[f4731,f94])).
fof(f4731,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK0))) | $spl172),
  inference(superposition,[],[f101,f4700])).
fof(f101,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X43] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X43)))))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.merge_def|split|19')).
fof(f4696,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | $spl168),
  inference(cnf_transformation,[],[f4696_D])).
fof(f4696_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f4698,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | $spl170),
  inference(cnf_transformation,[],[f4698_D])).
fof(f4698_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f4700,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | $spl172),
  inference(cnf_transformation,[],[f4700_D])).
fof(f4700_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f4702,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | $spl174),
  inference(cnf_transformation,[],[f4702_D])).
fof(f4702_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f4673,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0)),
  inference(trivial_inequality_removal,[],[f4668])).
fof(f4668,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.Other_0') | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0)),
  inference(superposition,[],[f94,f2004])).
fof(f2004,plain,(
  ( ! [X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X1) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X1)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X1)) )),
  inference(superposition,[],[f127,f108])).
fof(f108,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f46])).
fof(f46,plain,(
  ! [X0] : (? [X3,X4] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X4))) | ? [X2] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))) | ? [X1] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : (? [X11,X12] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X11),s('type.num.num',X12))) | ? [X13] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X13))) | ? [X14] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X10) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X10))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.val_approx_nchotomy')).
fof(f127,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X40,X41] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X41),s('type.num.num',X40)))))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.merge_def|split|17')).
fof(f4776,plain,(
  $false | $spl174),
  inference(subsumption_resolution,[],[f4775,f97])).
fof(f97,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(cnf_transformation,[],[f35])).
fof(f35,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|0','thm.clos_known.merge_def|split|20')).
fof(f4775,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl174),
  inference(forward_demodulation,[],[f94,f4702])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 115
% Passive clauses: 424
% Generated clauses: 4545
% Final active clauses: 102
% Final passive clauses: 245
% Input formulas: 35
% Initial clauses: 53
% 
% Duplicate literals: 618
% Trivial inequalities: 39
% Fw subsumption resolutions: 401
% Fw demodulations: 72
% Bw demodulations: 7
% 
% Simple tautologies: 74
% Equational tautologies: 27
% Forward subsumptions: 2444
% 
% Binary resolution: 28
% Forward superposition: 1617
% Backward superposition: 1451
% Self superposition: 121
% Equality factoring: 52
% Equality resolution: 9
% 
% Split clauses: 389
% Split components: 802
% SAT solver clauses: 496
% SAT solver unit clauses: 8
% SAT solver binary clauses: 19
% SAT solver learnt clauses: 1
% 
% Sat splits: 389
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 485
% TWLsolver calls for satisfiability: 55
% 
% Memory used [KB]: 3454
% Time elapsed: 0.119 s
% ------------------------------
% ------------------------------
% Success in time 0.14 s
