Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferSound
% (103831)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 747
% Passive clauses: 9520
% Generated clauses: 20451
% Final active clauses: 682
% Final passive clauses: 8274
% Input formulas: 39
% Initial clauses: 162
% 
% Duplicate literals: 527
% Trivial inequalities: 128
% Fw subsumption resolutions: 2012
% Fw demodulations: 314
% Bw demodulations: 61
% 
% Simple tautologies: 214
% Equational tautologies: 60
% Forward subsumptions: 7262
% 
% Binary resolution: 10809
% Factoring: 1349
% Forward superposition: 3019
% Backward superposition: 2171
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 41
% 
% Split clauses: 516
% Split components: 1049
% SAT solver clauses: 872
% SAT solver unit clauses: 15
% SAT solver binary clauses: 143
% 
% Sat splits: 516
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 870
% TWLsolver calls for satisfiability: 465
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferSound
% (103862)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 411
% Passive clauses: 2554
% Generated clauses: 11437
% Final active clauses: 402
% Final passive clauses: 2039
% Input formulas: 39
% Initial clauses: 162
% 
% Duplicate literals: 202
% Trivial inequalities: 51
% Fw subsumption resolutions: 934
% Fw demodulations: 387
% Bw demodulations: 111
% Condensations: 532
% Global subsumptions: 13
% 
% Simple tautologies: 232
% Equational tautologies: 16
% Forward subsumptions: 3604
% 
% Binary resolution: 3606
% Factoring: 274
% Forward superposition: 1513
% Backward superposition: 1030
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 13
% 
% Split clauses: 109
% Split components: 223
% SAT solver clauses: 3022
% SAT solver unit clauses: 168
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 109
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2970
% TWLsolver calls for satisfiability: 1555
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferSound
% (103890)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 2282
% Passive clauses: 8332
% Generated clauses: 128077
% Final active clauses: 900
% Final passive clauses: 1294
% Input formulas: 39
% Initial clauses: 162
% 
% Duplicate literals: 8636
% Trivial inequalities: 246
% Fw subsumption resolutions: 6939
% Fw demodulations: 10284
% Bw demodulations: 3203
% Condensations: 16726
% Global subsumptions: 364
% 
% Simple tautologies: 303
% Equational tautologies: 1612
% Forward subsumptions: 73867
% Bw demodulations to eq. taut.: 283
% 
% Unit resulting resolution: 19981
% Factoring: 23
% Forward superposition: 34887
% Backward superposition: 20970
% Self superposition: 288
% Equality factoring: 216
% Equality resolution: 16
% 
% Split clauses: 309
% Split components: 648
% SAT solver clauses: 9974
% SAT solver unit clauses: 776
% SAT solver binary clauses: 1613
% 
% Sat splits: 309
% Sat splitting refutations: 658
% 
% Memory used [KB]: 46950
% Time elapsed: 10.489 s
% ------------------------------
% ------------------------------
% remaining time: 31 next slice time: 31
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.inferSound
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferSound
% SZS output start Proof for thm.inferSound
fof(f44450,plain,(
  $false),
  inference(subsumption_resolution,[],[f44315,f967])).
fof(f967,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK14),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11),s(sK10,sK15(sK11,sK14,sK13,sK10)))))))),
  inference(resolution,[],[f217,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,sK15(X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X5)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK15(X3,X2,X1,X0)),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,sK15(X3,X2,X1,X0)))))))) | p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X5)))))))) & (? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4))))))) | p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3))))))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4)))))))) & (? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4))))))) | p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) <=> ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4))))))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) <=> ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4))))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) <=> ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X4))))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X8),s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))) <=> ! [X21] : (p(s(bool,'const.bool.IN_2'(s(X6,X21),s(fun(X6,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X8),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7),s(X6,X21))))))))),
  file('thm.inferSound.inferSound_6_prove_ltb','thm.infer.check_s_def')).
fof(f217,plain,(
  ~p(s(bool,'const.infer.check_s_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK14),s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11))))),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  p(s(bool,'const.infer.check_s_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK12),s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),sK12),s(fun('type.num.num',bool),sK14)))) & ~p(s(bool,'const.infer.check_s_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK14),s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f97])).
fof(f97,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),X4)))) & ~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X4),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),X4))))) & ~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X4),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),X4))))) => p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X4),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),X4))))) => p(s(bool,'const.infer.check_s_3'(s('type.num.num',X3),s(fun('type.num.num',bool),X4),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,'const.infer.check_s_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X8),s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X8),s(fun('type.num.num',bool),X10))))) => p(s(bool,'const.infer.check_s_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X10),s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,'const.infer.check_s_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X8),s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X8),s(fun('type.num.num',bool),X10))))) => p(s(bool,'const.infer.check_s_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X10),s('type.finite_map.fmap'(X6,'type.infer_t.infer_t'),X7)))))),
  file('thm.inferSound.inferSound_6_prove_ltb','thm.inferSound.inferSound_6_prove_ltb')).
fof(f44315,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK14),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11),s(sK10,sK15(sK11,sK14,sK13,sK10)))))))),
  inference(resolution,[],[f14236,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),sK12),s('type.infer_t.infer_t',X1)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),sK14),s('type.infer_t.infer_t',X1))))) )),
  inference(resolution,[],[f216,f219])).
fof(f219,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X3),s('type.infer_t.infer_t',X0)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X1),s(fun('type.num.num',bool),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X1),s('type.infer_t.infer_t',X0))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X1),s('type.infer_t.infer_t',X0)))) | ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X1),s(fun('type.num.num',bool),X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X3),s('type.infer_t.infer_t',X0))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X1),s('type.infer_t.infer_t',X0)))) => ! [X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X1),s(fun('type.num.num',bool),X3)))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),X3),s('type.infer_t.infer_t',X0))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11,X8,X9] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X8),s('type.infer_t.infer_t',X11)))) => ! [X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),X8),s(fun('type.num.num',bool),X10)))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X9),s(fun('type.num.num',bool),X10),s('type.infer_t.infer_t',X11))))))),
  file('thm.inferSound.inferSound_6_prove_ltb','thm.inferProps.check_t_more5|split|0')).
fof(f216,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),sK12),s(fun('type.num.num',bool),sK14))))),
  inference(cnf_transformation,[],[f138])).
fof(f14236,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK12),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11),s(sK10,sK15(sK11,sK14,sK13,sK10)))))))),
  inference(resolution,[],[f480,f966])).
fof(f966,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK15(sK11,sK14,sK13,sK10)),s(fun(sK10,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11))))))),
  inference(resolution,[],[f217,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK15(X3,X2,X1,X0)),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3))))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f480,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11)))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK12),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11),s(sK10,X0))))))) )),
  inference(resolution,[],[f215,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t','const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3),s(X0,X5)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3)))))) | ~p(s(bool,'const.infer.check_s_3'(s('type.num.num',X1),s(fun('type.num.num',bool),X2),s('type.finite_map.fmap'(X0,'type.infer_t.infer_t'),X3))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f215,plain,(
  p(s(bool,'const.infer.check_s_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),sK12),s('type.finite_map.fmap'(sK10,'type.infer_t.infer_t'),sK11))))),
  inference(cnf_transformation,[],[f138])).
% SZS output end Proof for thm.inferSound
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 920
% Passive clauses: 9583
% Generated clauses: 44003
% Final active clauses: 901
% Final passive clauses: 8102
% Input formulas: 39
% Initial clauses: 161
% 
% Selected by SInE selection: 38
% SInE iterations: 1
% 
% Duplicate literals: 1563
% Trivial inequalities: 96
% Fw subsumption resolutions: 6078
% Fw demodulations: 1662
% Bw demodulations: 7
% 
% Simple tautologies: 1783
% Equational tautologies: 54
% Forward subsumptions: 21208
% 
% Binary resolution: 31157
% Factoring: 131
% Forward superposition: 2032
% Backward superposition: 872
% Equality resolution: 17
% 
% Split clauses: 2447
% Split components: 5328
% SAT solver clauses: 2585
% SAT solver unit clauses: 29
% SAT solver binary clauses: 110
% 
% Sat splits: 2447
% Sat splitting refutations: 31
% 
% Memory used [KB]: 12920
% Time elapsed: 0.775 s
% ------------------------------
% ------------------------------
% Success in time 12.601 s
