Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closSem
% (84316)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 169
% Passive clauses: 3605
% Generated clauses: 19460
% Final active clauses: 161
% Final passive clauses: 3271
% Input formulas: 69
% Initial clauses: 142
% 
% Duplicate literals: 671
% Trivial inequalities: 475
% Fw subsumption resolutions: 2003
% Fw demodulations: 2931
% Bw demodulations: 6
% 
% Simple tautologies: 434
% Equational tautologies: 53
% Forward subsumptions: 8902
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 384
% Factoring: 1
% Forward superposition: 6516
% Backward superposition: 5968
% Self superposition: 152
% Equality factoring: 88
% Equality resolution: 39
% 
% Split clauses: 211
% Split components: 434
% SAT solver clauses: 491
% SAT solver unit clauses: 14
% SAT solver binary clauses: 104
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 211
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 490
% TWLsolver calls for satisfiability: 153
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closSem
% (84346)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 140
% Passive clauses: 1308
% Generated clauses: 14564
% Final active clauses: 132
% Final passive clauses: 1130
% Input formulas: 69
% Initial clauses: 142
% 
% Duplicate literals: 432
% Trivial inequalities: 396
% Fw subsumption resolutions: 972
% Fw demodulations: 2004
% Bw demodulations: 38
% Condensations: 743
% Global subsumptions: 8
% 
% Simple tautologies: 261
% Equational tautologies: 44
% Forward subsumptions: 6315
% 
% Binary resolution: 172
% Factoring: 1
% Forward superposition: 4438
% Backward superposition: 3629
% Self superposition: 118
% Equality factoring: 60
% Equality resolution: 34
% 
% Split clauses: 45
% Split components: 90
% SAT solver clauses: 2153
% SAT solver unit clauses: 152
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 45
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2119
% TWLsolver calls for satisfiability: 591
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closSem
% SZS output start Proof for thm.closSem
fof(f22618,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f252,f415,f413,f251,f414,f412,f417,f1001,f2167,f2173,f2170,f2518,f2521,f2519,f2737,f2739,f2742,f2744,f2746,f2750,f2752,f2754,f2756,f2760,f2762,f2764,f2766,f2785,f2787,f2789,f2791,f2793,f2795,f2797,f2799,f2829,f2831,f2837,f2974,f2976,f2983,f2984,f2985,f2986,f2987,f2988,f2989,f2990,f5870,f5877,f5878,f5880,f5882,f5886,f5888,f5890,f5892,f5896,f5898,f5900,f5902,f5905,f5907,f5909,f5911,f5914,f5916,f5918,f5920,f5922,f5924,f5926,f5948,f5950,f5980,f5984,f5986,f6082,f6084,f6086,f6088,f6186,f6187,f6188,f6189,f6190,f6683,f6684,f6687,f6688,f10033,f10152,f10154,f10156,f10157,f10153,f10161,f10162,f10155,f6626,f6625,f2520,f2738,f6567,f6587,f5533,f5553,f2168,f10176,f3301,f1011,f10186,f1010,f10177,f10477,f10479,f10488,f10491,f10492,f10493,f10548,f10550,f10552,f10554,f10558,f10560,f10562,f10564,f10566,f10568,f10570,f10572,f12221,f12222,f2174,f13462,f13479,f13477,f17019,f17094,f17439,f10100,f18235,f18312,f10099,f22508,f22523,f22573])).
fof(f22573,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f22572,f1968])).
fof(f1968,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',sK0),s('type.list.list'('type.closSem.v'),X0))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.closSem.v'),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f415,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : ((s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.closSem.v'),X1) != s('type.list.list'('type.closSem.v'),X3) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3)))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : ((s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.closSem.v'),X1) != s('type.list.list'('type.closSem.v'),X3)) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3)))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X17))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X18),s('type.list.list'('type.closSem.v'),X19))) <=> (s('type.num.num',X16) = s('type.num.num',X18) & s('type.list.list'('type.closSem.v'),X17) = s('type.list.list'('type.closSem.v'),X19)))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_|49||49||split|2')).
fof(f22572,plain,(
  ( ! [X193] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',sK0),s('type.list.list'('type.closSem.v'),X193))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.closSem.v'),X193)))) ) | $spl0),
  inference(forward_demodulation,[],[f22267,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v',happ(s(fun('type.list.list'('type.closSem.v'),'type.closSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),'type.closSem.v')),'const.closSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.closSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v',happ(s(fun('type.list.list'('type.closSem.v'),'type.closSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),'type.closSem.v')),'const.closSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.closSem.v'),X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X20] : s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X8),s('type.list.list'('type.closSem.v'),X20))) = s('type.closSem.v',happ(s(fun('type.list.list'('type.closSem.v'),'type.closSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),'type.closSem.v')),'const.closSem.Block_0'),s('type.num.num',X8))),s('type.list.list'('type.closSem.v'),X20)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_2_const.closSem.Block')).
fof(f22267,plain,(
  ( ! [X193] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',sK1),s('type.list.list'('type.closSem.v'),X193))) = s('type.closSem.v',happ(s(fun('type.list.list'('type.closSem.v'),'type.closSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),'type.closSem.v')),'const.closSem.Block_0'),s('type.num.num',sK0))),s('type.list.list'('type.closSem.v'),X193)))) ) | $spl0),
  inference(superposition,[],[f321,f21912])).
fof(f21912,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f21490,f320])).
fof(f320,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.closSem.v_CASE_7'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp',X0))))),X6),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num',X0))))),X7)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.closSem.v_CASE_7'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp',X0))))),X6),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num',X0))))),X7)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.closSem.v_CASE_7'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp',X0))))),X6),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num',X0))))),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X13),s('type.num.num',X6))) = s(X9,'const.closSem.v_CASE_7'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp',X9))))),X14),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num',X9))))),X15)))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_case_def|split|3')).
fof(f21490,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.closSem.v_CASE_7'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.closSem.v'),X0)),X4),s(fun('type.num.num',X0),X1),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp',X0))))),X5),s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num',X0))))),X6)))) ) | $spl0),
  inference(superposition,[],[f320,f412])).
fof(f22523,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f22522,f415])).
fof(f22522,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f22183,f302])).
fof(f302,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.closSem.v_|49||49||split|3','_i0_ax')).
fof(f22183,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f302,f21912])).
fof(f22508,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f22507,f415])).
fof(f22507,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f22160,f302])).
fof(f22160,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f21912,f302])).
fof(f10099,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl26),
  inference(cnf_transformation,[],[f10099_D])).
fof(f10099_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f18312,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) | $spl44),
  inference(cnf_transformation,[],[f18312_D])).
fof(f18312_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f18235,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) | ($spl10 | $spl20)),
  inference(duplicate_literal_removal,[],[f18062])).
fof(f18062,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl10 | $spl20)),
  inference(superposition,[],[f7065,f16571])).
fof(f16571,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f16556])).
fof(f16556,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f16473,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f263,f263])).
fof(f263,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.closSem.v_|49||49||split|3',aHL_BOOL_CASES)).
fof(f16473,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl10),
  inference(forward_demodulation,[],[f16391,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X31,X9] : ! [X32] : s(X9,X32) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X31),'const.sum.INL_1'(s(X9,X32)))))),
  file('thm.closSem.v_|49||49||split|3','thm.sum.OUTL')).
fof(f16391,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f10160,f9041])).
fof(f9041,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f8961,f255])).
fof(f255,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.closSem.v_|49||49||split|3',aHL_FALSITY)).
fof(f8961,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(superposition,[],[f264,f3506])).
fof(f3506,plain,(
  ( ! [X20] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X20)) ) | $spl10),
  inference(forward_demodulation,[],[f3394,f2198])).
fof(f2198,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2170,f265])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f3394,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20)))) ) | $spl10),
  inference(superposition,[],[f2454,f302])).
fof(f2454,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2431,f253])).
fof(f253,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  p(s(bool,'T_0'))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.TRUTH')).
fof(f2431,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f2101,f2198])).
fof(f2101,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2100,f302])).
fof(f2100,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2068,f253])).
fof(f2068,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f326,f458])).
fof(f458,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f302,f263])).
fof(f326,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X36,X37] : ((p(s(bool,X36)) => p(s(bool,X37))) => ((p(s(bool,X37)) => p(s(bool,X36))) => s(bool,X36) = s(bool,X37)))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f264,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f10160,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2170,f2198,f2179,f2279,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2454,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f3506,f3403,f3541,f3555,f4002,f4003,f4005,f4006,f4008,f4009,f4010,f4044,f4072,f4074,f4075,f4077,f4078,f4079,f4182,f4250,f4253,f4690,f4701,f4343,f4734,f4766,f4767,f4769,f4770,f4772,f4773,f4774,f4800,f4801,f4802,f4825,f4808,f4809,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5418,f5256,f5257,f5310,f5312,f5313,f5420,f5328,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5710,f5787,f5870,f5665,f5534,f5576,f5581,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5845,f5842,f5638,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f4698,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5558,f5555,f5419,f3637,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f4828,f4820,f4821,f5874,f5871,f6191,f5535,f5156,f4983,f4768,f4073,f4004,f3508,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f4771,f4076,f4007,f3511,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6673,f6680,f6681,f6682,f6683,f6685,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7532,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8710,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f8938,f8941,f8945,f9041,f9049,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9560,f9172,f9173,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9324,f9325,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9422,f9423,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9665,f9670,f9672,f9775,f10152,f10154,f10156,f10157,f10159])).
fof(f10159,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f10158])).
fof(f10158,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f8648,f4253])).
fof(f8648,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f7444,f3403])).
fof(f9775,plain,(
  ( ! [X2,X3,X1] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),happ(s(fun(X1,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),X2),s(X1,X3))))) = s('type.closSem.v',happ(s(fun(X1,'type.closSem.v'),'_dst_x0x1_2'(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.closSem.v'),'const.closSem.Word64_0'),s(fun(X1,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),X2))),s(X1,X3)))) )),
  inference(superposition,[],[f262,f360])).
fof(f9672,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f9636,f253])).
fof(f9636,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(superposition,[],[f5576,f6685])).
fof(f9670,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f9613,f253])).
fof(f9613,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f433,f6685])).
fof(f9665,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f9595,f418])).
fof(f9595,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f6685,f3403])).
fof(f9466,plain,(
  ( ! [X600,X602,X598,X601,X603,X597,X599,X604] : (s('type.list.list'('type.closSem.v'),X602) != s('type.list.list'('type.closSem.v'),happ(s(fun(X598,'type.list.list'('type.closSem.v')),X599),s(X598,happ(s(fun(X597,X598),X600),s(X597,X601))))) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X604),s('type.list.list'('type.closSem.v'),X602))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X603),s('type.list.list'('type.closSem.v'),happ(s(fun(X597,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X598,'type.list.list'('type.closSem.v')),X599),s(fun(X597,X598),X600))),s(X597,X601))))) | s('type.num.num',X603) != s('type.num.num',X604)) )),
  inference(superposition,[],[f387,f360])).
fof(f9464,plain,(
  ( ! [X588,X584,X586,X582,X585,X587,X581,X583] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X587),s('type.list.list'('type.closSem.v'),X588))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X586),s('type.list.list'('type.closSem.v'),happ(s(fun(X582,'type.list.list'('type.closSem.v')),X583),s(X582,happ(s(fun(X581,X582),X584),s(X581,X585))))))) | s('type.list.list'('type.closSem.v'),X588) = s('type.list.list'('type.closSem.v'),happ(s(fun(X581,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X582,'type.list.list'('type.closSem.v')),X583),s(fun(X581,X582),X584))),s(X581,X585)))) )),
  inference(superposition,[],[f386,f360])).
fof(f9456,plain,(
  ( ! [X532,X534,X530,X533,X535,X531] : (s('type.integer.int',X535) != s('type.integer.int',happ(s(fun(X531,'type.integer.int'),X532),s(X531,happ(s(fun(X530,X531),X533),s(X530,X534))))) | s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X535))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X530,'type.integer.int'),'_dst_x0x1_2'(s(fun(X531,'type.integer.int'),X532),s(fun(X530,X531),X533))),s(X530,X534)))))) )),
  inference(superposition,[],[f328,f360])).
fof(f9454,plain,(
  ( ! [X520,X522,X518,X521,X523,X519] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X523))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X519,'type.integer.int'),X520),s(X519,happ(s(fun(X518,X519),X521),s(X518,X522))))))) | s('type.integer.int',X523) = s('type.integer.int',happ(s(fun(X518,'type.integer.int'),'_dst_x0x1_2'(s(fun(X519,'type.integer.int'),X520),s(fun(X518,X519),X521))),s(X518,X522)))) )),
  inference(superposition,[],[f327,f360])).
fof(f9448,plain,(
  ( ! [X476,X478,X480,X482,X477,X479,X481,X483] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X482),s('type.list.list'('type.closSem.v'),X483))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X477,'type.num.num'),X478),s(X477,happ(s(fun(X476,X477),X479),s(X476,X480))))),s('type.list.list'('type.closSem.v'),X481))) | s('type.num.num',X482) = s('type.num.num',happ(s(fun(X476,'type.num.num'),'_dst_x0x1_2'(s(fun(X477,'type.num.num'),X478),s(fun(X476,X477),X479))),s(X476,X480)))) )),
  inference(superposition,[],[f385,f360])).
fof(f9441,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X426,X425),X432),s(fun(X425,X426),happ(s(fun(X428,fun(X425,X426)),X429),s(X428,happ(s(fun(X427,X428),X430),s(X427,X431)))))))) | s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) )),
  inference(superposition,[],[f349,f360])).
fof(f9440,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X416,X417),happ(s(fun(X419,fun(X416,X417)),X420),s(X419,happ(s(fun(X418,X419),X421),s(X418,X422))))),s(fun(X417,X416),X423)))) | s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) )),
  inference(superposition,[],[f349,f360])).
fof(f9439,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X408,X407),X414),s(fun(X407,X408),happ(s(fun(X410,fun(X407,X408)),X411),s(X410,happ(s(fun(X409,X410),X412),s(X409,X413)))))))) | s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) )),
  inference(superposition,[],[f348,f360])).
fof(f9438,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X398,X399),happ(s(fun(X401,fun(X398,X399)),X402),s(X401,happ(s(fun(X400,X401),X403),s(X400,X404))))),s(fun(X399,X398),X405)))) | s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) )),
  inference(superposition,[],[f348,f360])).
fof(f9436,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (s(fun(X380,X381),X387) != s(fun(X380,X381),happ(s(fun(X383,fun(X380,X381)),X384),s(X383,happ(s(fun(X382,X383),X385),s(X382,X386))))) | s(X381,happ(s(fun(X380,X381),X387),s(X380,X388))) = s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))) )),
  inference(superposition,[],[f346,f360])).
fof(f9435,plain,(
  ( ! [X377,X379,X374,X376,X378,X375] : (p(s(bool,happ(s(fun('type.list.list'(X374),bool),happ(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(X376,happ(s(fun(X375,X376),X378),s(X375,X379))))),s('type.list.list'(X374),sK2(happ(s(fun(X375,fun('type.list.list'(X374),bool)),'_dst_x0x1_2'(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(fun(X375,X376),X378))),s(X375,X379)),X374))))) | ~p(s(bool,happ(s(fun('type.list.list'(X374),bool),happ(s(fun(X375,fun('type.list.list'(X374),bool)),'_dst_x0x1_2'(s(fun(X376,fun('type.list.list'(X374),bool)),X377),s(fun(X375,X376),X378))),s(X375,X379))),s('type.list.list'(X374),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f360])).
fof(f9434,plain,(
  ( ! [X372,X368,X370,X367,X373,X369,X371] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X367,bool),happ(s(fun(X369,fun(X367,bool)),X370),s(X369,happ(s(fun(X368,X369),X371),s(X368,X372))))))) | s(X367,X373) = s(X367,sK4(X373,happ(s(fun(X368,fun(X367,bool)),'_dst_x0x1_2'(s(fun(X369,fun(X367,bool)),X370),s(fun(X368,X369),X371))),s(X368,X372)),X367))) )),
  inference(superposition,[],[f847,f360])).
fof(f9433,plain,(
  ( ! [X364,X366,X360,X362,X365,X361,X363] : (p(s(bool,'|dollar|exists_1'(s(fun(X360,bool),happ(s(fun(X362,fun(X360,bool)),X363),s(X362,happ(s(fun(X361,X362),X364),s(X361,X365)))))))) | s(X360,X366) = s(X360,sK4(X366,happ(s(fun(X361,fun(X360,bool)),'_dst_x0x1_2'(s(fun(X362,fun(X360,bool)),X363),s(fun(X361,X362),X364))),s(X361,X365)),X360))) )),
  inference(superposition,[],[f311,f360])).
fof(f9432,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (~p(s(bool,happ(s(fun(X353,bool),happ(s(fun(X355,fun(X353,bool)),X356),s(X355,happ(s(fun(X354,X355),X357),s(X354,X358))))),s(X353,X359)))) | p(s(bool,happ(s(fun(X353,bool),happ(s(fun(X354,fun(X353,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X353,bool)),X356),s(fun(X354,X355),X357))),s(X354,X358))),s(X353,'const.min.@_1'(s(fun(X353,bool),happ(s(fun(X354,fun(X353,bool)),'_dst_x0x1_2'(s(fun(X355,fun(X353,bool)),X356),s(fun(X354,X355),X357))),s(X354,X358))))))))) )),
  inference(superposition,[],[f310,f360])).
fof(f9431,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (~p(s(bool,happ(s(fun(X346,bool),happ(s(fun(X348,fun(X346,bool)),X349),s(X348,happ(s(fun(X347,X348),X350),s(X347,X351))))),s(X346,X352)))) | p(s(bool,'|dollar|exists_1'(s(fun(X346,bool),happ(s(fun(X347,fun(X346,bool)),'_dst_x0x1_2'(s(fun(X348,fun(X346,bool)),X349),s(fun(X347,X348),X350))),s(X347,X351))))))) )),
  inference(superposition,[],[f5419,f360])).
fof(f9428,plain,(
  ( ! [X325,X327,X329,X326,X328,X330] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X330),'const.sum.INL_1'(s(bool,happ(s(fun(X326,bool),X327),s(X326,happ(s(fun(X325,X326),X328),s(X325,X329))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X325,bool),'_dst_x0x1_2'(s(fun(X326,bool),X327),s(fun(X325,X326),X328))),s(X325,X329)))) )),
  inference(superposition,[],[f7444,f360])).
fof(f9427,plain,(
  ( ! [X321,X323,X324,X320,X322,X319] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X324,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X320,bool),X321),s(X320,happ(s(fun(X319,X320),X322),s(X319,X323))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X319,bool),'_dst_x0x1_2'(s(fun(X320,bool),X321),s(fun(X319,X320),X322))),s(X319,X323)))) )),
  inference(superposition,[],[f7245,f360])).
fof(f9424,plain,(
  ( ! [X300,X302,X304,X306,X301,X303,X305] : (~p(s(bool,happ(s(fun(X301,bool),X302),s(X301,happ(s(fun(X300,X301),X303),s(X300,X304)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X305),bool),sK16(X306,happ(s(fun(X300,bool),'_dst_x0x1_2'(s(fun(X301,bool),X302),s(fun(X300,X301),X303))),s(X300,X304)),X305,bool))))) )),
  inference(superposition,[],[f5576,f360])).
fof(f9423,plain,(
  ( ! [X294,X296,X298,X295,X297,X299] : (s(bool,'T_0') != s(bool,happ(s(fun(X295,bool),X296),s(X295,happ(s(fun(X294,X295),X297),s(X294,X298))))) | s(bool,X299) = s(bool,happ(s(fun(X294,bool),'_dst_x0x1_2'(s(fun(X295,bool),X296),s(fun(X294,X295),X297))),s(X294,X298))) | s(bool,'F_0') = s(bool,X299)) ) | $spl10),
  inference(superposition,[],[f4253,f360])).
fof(f9422,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X290,bool),X291),s(X290,happ(s(fun(X289,X290),X292),s(X289,X293))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X289,bool),'_dst_x0x1_2'(s(fun(X290,bool),X291),s(fun(X289,X290),X292))),s(X289,X293)))) ) | $spl10),
  inference(superposition,[],[f3506,f360])).
fof(f9411,plain,(
  ( ! [X231,X229,X227,X230,X228,X226] : (p(s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X226,X227),X229),s(X226,X230)))))) | s(bool,X231) != s(bool,happ(s(fun(X226,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X226,X227),X229))),s(X226,X230))) | s(bool,'F_0') = s(bool,X231)) )),
  inference(superposition,[],[f2103,f360])).
fof(f9410,plain,(
  ( ! [X222,X220,X225,X223,X221,X224] : (~p(s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X220,X221),X223),s(X220,X224)))))) | s(bool,X225) = s(bool,happ(s(fun(X220,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X220,X221),X223))),s(X220,X224))) | s(bool,'F_0') = s(bool,X225)) )),
  inference(superposition,[],[f2101,f360])).
fof(f9409,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X219),'const.sum.INL_1'(s(bool,happ(s(fun(X215,bool),X216),s(X215,happ(s(fun(X214,X215),X217),s(X214,X218))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X214,bool),'_dst_x0x1_2'(s(fun(X215,bool),X216),s(fun(X214,X215),X217))),s(X214,X218)))) )),
  inference(superposition,[],[f2099,f360])).
fof(f9408,plain,(
  ( ! [X212,X210,X208,X213,X211,X209] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X213,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X209,bool),X210),s(X209,happ(s(fun(X208,X209),X211),s(X208,X212))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X208,bool),'_dst_x0x1_2'(s(fun(X209,bool),X210),s(fun(X208,X209),X211))),s(X208,X212)))) )),
  inference(superposition,[],[f2095,f360])).
fof(f9405,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X196),'const.sum.INL_1'(s(bool,happ(s(fun(X192,bool),X193),s(X192,happ(s(fun(X191,X192),X194),s(X191,X195))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X191,bool),'_dst_x0x1_2'(s(fun(X192,bool),X193),s(fun(X191,X192),X194))),s(X191,X195)))) )),
  inference(superposition,[],[f526,f360])).
fof(f9404,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X190,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X186,bool),X187),s(X186,happ(s(fun(X185,X186),X188),s(X185,X189))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X185,bool),'_dst_x0x1_2'(s(fun(X186,bool),X187),s(fun(X185,X186),X188))),s(X185,X189)))) )),
  inference(superposition,[],[f499,f360])).
fof(f9403,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (p(s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183)))))) | s(bool,X184) = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f480,f360])).
fof(f9402,plain,(
  ( ! [X177,X175,X178,X176,X174] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X175,bool),X176),s(X175,happ(s(fun(X174,X175),X177),s(X174,X178))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X174,bool),'_dst_x0x1_2'(s(fun(X175,bool),X176),s(fun(X174,X175),X177))),s(X174,X178)))) )),
  inference(superposition,[],[f458,f360])).
fof(f9401,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (~p(s(bool,happ(s(fun(X169,bool),X170),s(X169,happ(s(fun(X168,X169),X171),s(X168,X172)))))) | s(bool,X173) != s(bool,happ(s(fun(X168,bool),'_dst_x0x1_2'(s(fun(X169,bool),X170),s(fun(X168,X169),X171))),s(X168,X172))) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f433,f360])).
fof(f9396,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (~p(s(bool,happ(s(fun(X139,bool),X140),s(X139,happ(s(fun(X138,X139),X141),s(X138,X142)))))) | s(bool,X143) != s(bool,happ(s(fun(X138,bool),'_dst_x0x1_2'(s(fun(X139,bool),X140),s(fun(X138,X139),X141))),s(X138,X142))) | p(s(bool,X143))) )),
  inference(superposition,[],[f404,f360])).
fof(f9395,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(X133,bool),X134),s(X133,happ(s(fun(X132,X133),X135),s(X132,X136)))))) | ~p(s(bool,X137)) | s(bool,X137) = s(bool,happ(s(fun(X132,bool),'_dst_x0x1_2'(s(fun(X133,bool),X134),s(fun(X132,X133),X135))),s(X132,X136)))) )),
  inference(superposition,[],[f326,f360])).
fof(f9393,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | p(s(bool,X125)) | s(bool,X125) = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f323,f360])).
fof(f9392,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119)))) )),
  inference(superposition,[],[f269,f360])).
fof(f9391,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'F_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114))))) | ~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) )),
  inference(superposition,[],[f268,f360])).
fof(f9390,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f265,f360])).
fof(f9389,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,happ(s(fun(X101,bool),X102),s(X101,happ(s(fun(X100,X101),X103),s(X100,X104)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X100,bool),'_dst_x0x1_2'(s(fun(X101,bool),X102),s(fun(X100,X101),X103))),s(X100,X104)))) )),
  inference(superposition,[],[f264,f360])).
fof(f9387,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun('type.list.list'(X91),X92),X94),s('type.list.list'(X91),sK2('_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun('type.list.list'(X91),X92),X94)),X91))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X91),bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun('type.list.list'(X91),X92),X94))),s('type.list.list'(X91),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f360])).
fof(f9386,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,happ(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))),s(X86,'const.min.@_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))))) )),
  inference(superposition,[],[f310,f360])).
fof(f9385,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,X85)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f5419,f360])).
fof(f9369,plain,(
  ( ! [X633,X635,X629,X631,X632,X634,X628,X630] : (s('type.list.list'('type.closSem.v'),X633) != s('type.list.list'('type.closSem.v'),happ(s(fun(X630,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X628,'type.list.list'('type.closSem.v')),X629),s(fun(X630,X628),X631))),s(X630,X632))) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X635),s('type.list.list'('type.closSem.v'),X633))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X634),s('type.list.list'('type.closSem.v'),happ(s(fun(X628,'type.list.list'('type.closSem.v')),X629),s(X628,happ(s(fun(X630,X628),X631),s(X630,X632))))))) | s('type.num.num',X634) != s('type.num.num',X635)) )),
  inference(superposition,[],[f387,f360])).
fof(f9367,plain,(
  ( ! [X617,X619,X613,X615,X616,X618,X612,X614] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X618),s('type.list.list'('type.closSem.v'),X619))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X617),s('type.list.list'('type.closSem.v'),happ(s(fun(X614,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X612,'type.list.list'('type.closSem.v')),X613),s(fun(X614,X612),X615))),s(X614,X616))))) | s('type.list.list'('type.closSem.v'),X619) = s('type.list.list'('type.closSem.v'),happ(s(fun(X612,'type.list.list'('type.closSem.v')),X613),s(X612,happ(s(fun(X614,X612),X615),s(X614,X616)))))) )),
  inference(superposition,[],[f386,f360])).
fof(f9359,plain,(
  ( ! [X565,X561,X563,X564,X566,X562] : (s('type.integer.int',X566) != s('type.integer.int',happ(s(fun(X563,'type.integer.int'),'_dst_x0x1_2'(s(fun(X561,'type.integer.int'),X562),s(fun(X563,X561),X564))),s(X563,X565))) | s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X566))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X561,'type.integer.int'),X562),s(X561,happ(s(fun(X563,X561),X564),s(X563,X565)))))))) )),
  inference(superposition,[],[f328,f360])).
fof(f9357,plain,(
  ( ! [X553,X549,X551,X552,X554,X550] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X554))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X551,'type.integer.int'),'_dst_x0x1_2'(s(fun(X549,'type.integer.int'),X550),s(fun(X551,X549),X552))),s(X551,X553))))) | s('type.integer.int',X554) = s('type.integer.int',happ(s(fun(X549,'type.integer.int'),X550),s(X549,happ(s(fun(X551,X549),X552),s(X551,X553)))))) )),
  inference(superposition,[],[f327,f360])).
fof(f9351,plain,(
  ( ! [X509,X511,X507,X512,X514,X508,X510,X513] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X513),s('type.list.list'('type.closSem.v'),X514))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X509,'type.num.num'),'_dst_x0x1_2'(s(fun(X507,'type.num.num'),X508),s(fun(X509,X507),X510))),s(X509,X511))),s('type.list.list'('type.closSem.v'),X512))) | s('type.num.num',X513) = s('type.num.num',happ(s(fun(X507,'type.num.num'),X508),s(X507,happ(s(fun(X509,X507),X510),s(X509,X511)))))) )),
  inference(superposition,[],[f385,f360])).
fof(f9344,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X447,X446),X453),s(fun(X446,X447),happ(s(fun(X450,fun(X446,X447)),'_dst_x0x1_2'(s(fun(X448,fun(X446,X447)),X449),s(fun(X450,X448),X451))),s(X450,X452)))))) | s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) )),
  inference(superposition,[],[f349,f360])).
fof(f9343,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X437,X438),happ(s(fun(X441,fun(X437,X438)),'_dst_x0x1_2'(s(fun(X439,fun(X437,X438)),X440),s(fun(X441,X439),X442))),s(X441,X443))),s(fun(X438,X437),X444)))) | s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) )),
  inference(superposition,[],[f349,f360])).
fof(f9342,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X429,X428),X435),s(fun(X428,X429),happ(s(fun(X432,fun(X428,X429)),'_dst_x0x1_2'(s(fun(X430,fun(X428,X429)),X431),s(fun(X432,X430),X433))),s(X432,X434)))))) | s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) )),
  inference(superposition,[],[f348,f360])).
fof(f9341,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X419,X420),happ(s(fun(X423,fun(X419,X420)),'_dst_x0x1_2'(s(fun(X421,fun(X419,X420)),X422),s(fun(X423,X421),X424))),s(X423,X425))),s(fun(X420,X419),X426)))) | s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) )),
  inference(superposition,[],[f348,f360])).
fof(f9339,plain,(
  ( ! [X408,X405,X407,X401,X403,X409,X404,X406,X402] : (s(fun(X401,X402),X408) != s(fun(X401,X402),happ(s(fun(X405,fun(X401,X402)),'_dst_x0x1_2'(s(fun(X403,fun(X401,X402)),X404),s(fun(X405,X403),X406))),s(X405,X407))) | s(X402,happ(s(fun(X401,X402),X408),s(X401,X409))) = s(X402,happ(s(fun(X401,X402),happ(s(fun(X403,fun(X401,X402)),X404),s(X403,happ(s(fun(X405,X403),X406),s(X405,X407))))),s(X401,X409)))) )),
  inference(superposition,[],[f346,f360])).
fof(f9338,plain,(
  ( ! [X397,X399,X395,X396,X398,X400] : (p(s(bool,happ(s(fun('type.list.list'(X395),bool),happ(s(fun(X398,fun('type.list.list'(X395),bool)),'_dst_x0x1_2'(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(fun(X398,X396),X399))),s(X398,X400))),s('type.list.list'(X395),sK2(happ(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400)))),X395))))) | ~p(s(bool,happ(s(fun('type.list.list'(X395),bool),happ(s(fun(X396,fun('type.list.list'(X395),bool)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400))))),s('type.list.list'(X395),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f360])).
fof(f9337,plain,(
  ( ! [X389,X391,X393,X388,X390,X392,X394] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X388,bool),happ(s(fun(X391,fun(X388,bool)),'_dst_x0x1_2'(s(fun(X389,fun(X388,bool)),X390),s(fun(X391,X389),X392))),s(X391,X393))))) | s(X388,X394) = s(X388,sK4(X394,happ(s(fun(X389,fun(X388,bool)),X390),s(X389,happ(s(fun(X391,X389),X392),s(X391,X393)))),X388))) )),
  inference(superposition,[],[f847,f360])).
fof(f9336,plain,(
  ( ! [X383,X385,X387,X382,X384,X386,X381] : (p(s(bool,'|dollar|exists_1'(s(fun(X381,bool),happ(s(fun(X384,fun(X381,bool)),'_dst_x0x1_2'(s(fun(X382,fun(X381,bool)),X383),s(fun(X384,X382),X385))),s(X384,X386)))))) | s(X381,X387) = s(X381,sK4(X387,happ(s(fun(X382,fun(X381,bool)),X383),s(X382,happ(s(fun(X384,X382),X385),s(X384,X386)))),X381))) )),
  inference(superposition,[],[f311,f360])).
fof(f9335,plain,(
  ( ! [X377,X379,X374,X380,X376,X378,X375] : (~p(s(bool,happ(s(fun(X374,bool),happ(s(fun(X377,fun(X374,bool)),'_dst_x0x1_2'(s(fun(X375,fun(X374,bool)),X376),s(fun(X377,X375),X378))),s(X377,X379))),s(X374,X380)))) | p(s(bool,happ(s(fun(X374,bool),happ(s(fun(X375,fun(X374,bool)),X376),s(X375,happ(s(fun(X377,X375),X378),s(X377,X379))))),s(X374,'const.min.@_1'(s(fun(X374,bool),happ(s(fun(X375,fun(X374,bool)),X376),s(X375,happ(s(fun(X377,X375),X378),s(X377,X379))))))))))) )),
  inference(superposition,[],[f310,f360])).
fof(f9334,plain,(
  ( ! [X372,X368,X370,X367,X373,X369,X371] : (~p(s(bool,happ(s(fun(X367,bool),happ(s(fun(X370,fun(X367,bool)),'_dst_x0x1_2'(s(fun(X368,fun(X367,bool)),X369),s(fun(X370,X368),X371))),s(X370,X372))),s(X367,X373)))) | p(s(bool,'|dollar|exists_1'(s(fun(X367,bool),happ(s(fun(X368,fun(X367,bool)),X369),s(X368,happ(s(fun(X370,X368),X371),s(X370,X372))))))))) )),
  inference(superposition,[],[f5419,f360])).
fof(f9330,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X341),'const.sum.INL_1'(s(bool,happ(s(fun(X338,bool),'_dst_x0x1_2'(s(fun(X336,bool),X337),s(fun(X338,X336),X339))),s(X338,X340))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X336,bool),X337),s(X336,happ(s(fun(X338,X336),X339),s(X338,X340)))))) )),
  inference(superposition,[],[f7444,f360])).
fof(f9329,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X335,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X332,bool),'_dst_x0x1_2'(s(fun(X330,bool),X331),s(fun(X332,X330),X333))),s(X332,X334))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X330,bool),X331),s(X330,happ(s(fun(X332,X330),X333),s(X332,X334)))))) )),
  inference(superposition,[],[f7245,f360])).
fof(f9326,plain,(
  ( ! [X313,X315,X316,X312,X314,X311,X317] : (~p(s(bool,happ(s(fun(X313,bool),'_dst_x0x1_2'(s(fun(X311,bool),X312),s(fun(X313,X311),X314))),s(X313,X315)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X316),bool),sK16(X317,happ(s(fun(X311,bool),X312),s(X311,happ(s(fun(X313,X311),X314),s(X313,X315)))),X316,bool))))) )),
  inference(superposition,[],[f5576,f360])).
fof(f9325,plain,(
  ( ! [X308,X310,X306,X309,X305,X307] : (s(bool,'T_0') != s(bool,happ(s(fun(X307,bool),'_dst_x0x1_2'(s(fun(X305,bool),X306),s(fun(X307,X305),X308))),s(X307,X309))) | s(bool,X310) = s(bool,happ(s(fun(X305,bool),X306),s(X305,happ(s(fun(X307,X305),X308),s(X307,X309))))) | s(bool,'F_0') = s(bool,X310)) ) | $spl10),
  inference(superposition,[],[f4253,f360])).
fof(f9324,plain,(
  ( ! [X300,X302,X304,X301,X303] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X302,bool),'_dst_x0x1_2'(s(fun(X300,bool),X301),s(fun(X302,X300),X303))),s(X302,X304))))) | s(bool,'T_0') = s(bool,happ(s(fun(X300,bool),X301),s(X300,happ(s(fun(X302,X300),X303),s(X302,X304)))))) ) | $spl10),
  inference(superposition,[],[f3506,f360])).
fof(f9313,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (p(s(bool,happ(s(fun(X239,bool),'_dst_x0x1_2'(s(fun(X237,bool),X238),s(fun(X239,X237),X240))),s(X239,X241)))) | s(bool,X242) != s(bool,happ(s(fun(X237,bool),X238),s(X237,happ(s(fun(X239,X237),X240),s(X239,X241))))) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f2103,f360])).
fof(f9312,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X233,bool),'_dst_x0x1_2'(s(fun(X231,bool),X232),s(fun(X233,X231),X234))),s(X233,X235)))) | s(bool,X236) = s(bool,happ(s(fun(X231,bool),X232),s(X231,happ(s(fun(X233,X231),X234),s(X233,X235))))) | s(bool,'F_0') = s(bool,X236)) )),
  inference(superposition,[],[f2101,f360])).
fof(f9311,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X230),'const.sum.INL_1'(s(bool,happ(s(fun(X227,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X227,X225),X228))),s(X227,X229))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X227,X225),X228),s(X227,X229)))))) )),
  inference(superposition,[],[f2099,f360])).
fof(f9310,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X224,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X221,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X221,X219),X222))),s(X221,X223))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X221,X219),X222),s(X221,X223)))))) )),
  inference(superposition,[],[f2095,f360])).
fof(f9307,plain,(
  ( ! [X206,X204,X202,X207,X205,X203] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X207),'const.sum.INL_1'(s(bool,happ(s(fun(X204,bool),'_dst_x0x1_2'(s(fun(X202,bool),X203),s(fun(X204,X202),X205))),s(X204,X206))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X202,bool),X203),s(X202,happ(s(fun(X204,X202),X205),s(X204,X206)))))) )),
  inference(superposition,[],[f526,f360])).
fof(f9306,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X198,bool),'_dst_x0x1_2'(s(fun(X196,bool),X197),s(fun(X198,X196),X199))),s(X198,X200))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),X197),s(X196,happ(s(fun(X198,X196),X199),s(X198,X200)))))) )),
  inference(superposition,[],[f499,f360])).
fof(f9305,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (p(s(bool,happ(s(fun(X192,bool),'_dst_x0x1_2'(s(fun(X190,bool),X191),s(fun(X192,X190),X193))),s(X192,X194)))) | s(bool,X195) = s(bool,happ(s(fun(X190,bool),X191),s(X190,happ(s(fun(X192,X190),X193),s(X192,X194))))) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f480,f360])).
fof(f9304,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X187,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X187,X185),X188))),s(X187,X189))))) | s(bool,'F_0') = s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X187,X185),X188),s(X187,X189)))))) )),
  inference(superposition,[],[f458,f360])).
fof(f9303,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun(X181,bool),'_dst_x0x1_2'(s(fun(X179,bool),X180),s(fun(X181,X179),X182))),s(X181,X183)))) | s(bool,X184) != s(bool,happ(s(fun(X179,bool),X180),s(X179,happ(s(fun(X181,X179),X182),s(X181,X183))))) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f433,f360])).
fof(f9298,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,happ(s(fun(X151,bool),'_dst_x0x1_2'(s(fun(X149,bool),X150),s(fun(X151,X149),X152))),s(X151,X153)))) | s(bool,X154) != s(bool,happ(s(fun(X149,bool),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))) | p(s(bool,X154))) )),
  inference(superposition,[],[f404,f360])).
fof(f9297,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X145,bool),'_dst_x0x1_2'(s(fun(X143,bool),X144),s(fun(X145,X143),X146))),s(X145,X147)))) | ~p(s(bool,X148)) | s(bool,X148) = s(bool,happ(s(fun(X143,bool),X144),s(X143,happ(s(fun(X145,X143),X146),s(X145,X147)))))) )),
  inference(superposition,[],[f326,f360])).
fof(f9295,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (p(s(bool,happ(s(fun(X133,bool),'_dst_x0x1_2'(s(fun(X131,bool),X132),s(fun(X133,X131),X134))),s(X133,X135)))) | p(s(bool,X136)) | s(bool,X136) = s(bool,happ(s(fun(X131,bool),X132),s(X131,happ(s(fun(X133,X131),X134),s(X133,X135)))))) )),
  inference(superposition,[],[f323,f360])).
fof(f9294,plain,(
  ( ! [X127,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X128,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X128,X126),X129))),s(X128,X130)))) | s(bool,'F_0') = s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130)))))) )),
  inference(superposition,[],[f269,f360])).
fof(f9293,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'F_0') != s(bool,happ(s(fun(X123,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X123,X121),X124))),s(X123,X125))) | ~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) )),
  inference(superposition,[],[f268,f360])).
fof(f9292,plain,(
  ( ! [X118,X116,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X118,X116),X119))),s(X118,X120)))) | s(bool,'T_0') = s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X118,X116),X119),s(X118,X120)))))) )),
  inference(superposition,[],[f265,f360])).
fof(f9291,plain,(
  ( ! [X111,X114,X112,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X113,X111),X114))),s(X113,X115)))) | s(bool,'T_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X113,X111),X114),s(X113,X115)))))) )),
  inference(superposition,[],[f264,f360])).
fof(f9289,plain,(
  ( ! [X103,X101,X105,X102,X104] : (~p(s(bool,happ(s(fun(X103,bool),'_dst_x0x1_2'(s(fun(X101,bool),X102),s(fun(X103,X101),X104))),s(X103,X105)))) | p(s(bool,happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102))))))) )),
  inference(superposition,[],[f310,f360])).
fof(f9288,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,happ(s(fun(X98,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X98,X96),X99))),s(X98,X100)))) | p(s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97))))) )),
  inference(superposition,[],[f5419,f360])).
fof(f9177,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s(fun(X104,X105),'_dst_x0x1_2'(s(fun(X106,X105),happ(s(fun(X108,fun(X106,X105)),X109),s(X108,X110))),s(fun(X104,X106),X107))) = s(fun(X104,X105),happ(s(fun(X108,fun(X104,X105)),'_dst_x0x1_2'(s(fun(fun(X106,X105),fun(X104,X105)),'_dstsel_0_x01_1'(s(fun(X104,X106),X107))),s(fun(X108,fun(X106,X105)),X109))),s(X108,X110)))) )),
  inference(superposition,[],[f360,f359])).
fof(f9173,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X85) = s(bool,happ(s(fun(X82,bool),'_dst_x0x1_2'(s(fun(X80,bool),X81),s(fun(X82,X80),X83))),s(X82,X84))) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,happ(s(fun(X82,X80),X83),s(X82,X84)))))) ) | $spl10),
  inference(superposition,[],[f360,f3403])).
fof(f9172,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(X76,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X76,X74),X77))),s(X76,X78))) | s(bool,X79) = s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X76,X74),X77),s(X76,X78))))) | s(bool,'F_0') = s(bool,X79)) ) | $spl10),
  inference(superposition,[],[f360,f3403])).
fof(f9560,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun(X66,X64),X67))),s(X66,X68))) | s(bool,'F_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,happ(s(fun(X66,X64),X67),s(X66,X68)))))) ) | $spl10),
  inference(forward_demodulation,[],[f9170,f2198])).
fof(f9170,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun(X66,X64),X67))),s(X66,X68))) | s(bool,'F_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,happ(s(fun(X66,X64),X67),s(X66,X68)))))) ) | $spl10),
  inference(superposition,[],[f360,f2454])).
fof(f9166,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,X45) = s(bool,happ(s(fun(X42,bool),'_dst_x0x1_2'(s(fun(X40,bool),X41),s(fun(X42,X40),X43))),s(X42,X44))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,happ(s(fun(X42,X40),X43),s(X42,X44)))))) )),
  inference(superposition,[],[f360,f431])).
fof(f9165,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),'_dst_x0x1_2'(s(fun(X34,bool),X35),s(fun(X36,X34),X37))),s(X36,X38))) | s(bool,X39) = s(bool,happ(s(fun(X34,bool),X35),s(X34,happ(s(fun(X36,X34),X37),s(X36,X38))))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f360,f431])).
fof(f9164,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X29,bool),X30),s(fun(X31,X29),X32))),s(X31,X33))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f360,f263])).
fof(f9163,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f360,f333])).
fof(f9162,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f360,f360])).
fof(f9161,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X113] : (s(X113,happ(s(fun(fun(X108,X109),X113),X114),s(fun(X108,X109),'_dst_x0x1_2'(s(fun(X110,X109),X112),s(fun(X108,X110),X111))))) = s(X113,happ(s(fun(fun(X110,X109),X113),'_dst_x0x1_2'(s(fun(fun(X108,X109),X113),X114),s(fun(fun(X110,X109),fun(X108,X109)),'_dstsel_0_x01_1'(s(fun(X108,X110),X111))))),s(fun(X110,X109),X112)))) )),
  inference(superposition,[],[f360,f359])).
fof(f9160,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK16(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f360,f351])).
fof(f9159,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f360,f360])).
fof(f9158,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f360,f360])).
fof(f9106,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f360,f346])).
fof(f9503,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),X5))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),X5)))) )),
  inference(forward_demodulation,[],[f9123,f332])).
fof(f9123,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),X5))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.closSem.v'),X6),'const.sum.INL_1'(s('type.list.list'('type.closSem.v'),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f332,f360,f387])).
fof(f9502,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X5),X6))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X5),X6)))))) )),
  inference(forward_demodulation,[],[f9124,f333])).
fof(f9124,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X5),X6))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.closSem.v'),X5),'type.list.list'('type.closSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.closSem.v'),X5),X6)))))) )),
  inference(unit_resulting_resolution,[],[f333,f360,f387])).
fof(f9501,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X5),s('type.num.num',X6))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X5),s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f9125,f317])).
fof(f9125,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X5),s('type.num.num',X6))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.closSem.v'),X7),s(fun('type.num.num','type.list.list'('type.closSem.v')),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(unit_resulting_resolution,[],[f317,f360,f387])).
fof(f9127,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X0,'type.integer.int'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun(X2,'type.integer.int'),'_dst_x0x1_2'(s(fun(X0,'type.integer.int'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f360,f328])).
fof(f9128,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X3,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f259,f360,f387])).
fof(f9141,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X0,X1),X3))),s(X0,X4))),s('type.list.list'('type.closSem.v'),happ(s(fun(X5,'type.list.list'('type.closSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X7,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.closSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f360,f360,f387])).
fof(f9142,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X0,'type.list.list'('type.closSem.v')),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X5))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X2,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X0,'type.list.list'('type.closSem.v')),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f259,f360,f387])).
fof(f9143,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X6))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X3,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f815,f360,f387])).
fof(f9481,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),happ(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),happ(s(fun(X3,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(forward_demodulation,[],[f9152,f332])).
fof(f9152,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),happ(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X6),'const.sum.INL_1'(s('type.num.num',X0))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X3,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.closSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f332,f360,f387])).
fof(f9480,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X4,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f9153,f333])).
fof(f9153,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X0),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X0),X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X4,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f333,f360,f387])).
fof(f9479,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X4,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f9154,f317])).
fof(f9154,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),happ(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X7),s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X4,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.closSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f317,f360,f387])).
fof(f9155,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.closSem.v'),happ(s(fun(X5,'type.list.list'('type.closSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.closSem.v'),happ(s(fun(X7,'type.list.list'('type.closSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.closSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f360,f360,f387])).
fof(f9049,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8983,f255])).
fof(f8983,plain,(
  ( ! [X43,X42] : (p(s(bool,'F_0')) | s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f2103,f3506])).
fof(f8945,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,sK4(X75,X76,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76)))) ) | $spl10),
  inference(superposition,[],[f3506,f847])).
fof(f8941,plain,(
  ( ! [X66,X64,X65] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X64),s(fun('type.num.num',bool),X65),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))))) ) | $spl10),
  inference(superposition,[],[f3506,f317])).
fof(f8938,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) ) | $spl10),
  inference(superposition,[],[f3506,f333])).
fof(f8766,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),sK2(happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34)),X31)))))) )),
  inference(forward_demodulation,[],[f8746,f359])).
fof(f8746,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun('type.list.list'(X31),bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun('type.list.list'(X31),X32),X33))),s('type.list.list'(X31),sK2(happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34)),X31))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(fun(X32,bool),fun('type.list.list'(X31),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X31),X32),X33))),s(fun(X32,bool),X34))),s('type.list.list'(X31),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f359])).
fof(f8745,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0x1_2'(s(fun(X27,bool),X29),s(fun(X26,X27),X28))))) | s(X26,X30) = s(X26,sK4(X30,happ(s(fun(fun(X27,bool),fun(X26,bool)),'_dstsel_0_x01_1'(s(fun(X26,X27),X28))),s(fun(X27,bool),X29)),X26))) )),
  inference(superposition,[],[f847,f359])).
fof(f8744,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0x1_2'(s(fun(X22,bool),X24),s(fun(X21,X22),X23)))))) | s(X21,X25) = s(X21,sK4(X25,happ(s(fun(fun(X22,bool),fun(X21,bool)),'_dstsel_0_x01_1'(s(fun(X21,X22),X23))),s(fun(X22,bool),X24)),X21))) )),
  inference(superposition,[],[f311,f359])).
fof(f8680,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f7444,f847])).
fof(f8679,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f7444,f332])).
fof(f8678,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f7444,f336])).
fof(f8677,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f7444,f335])).
fof(f8676,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f7444,f317])).
fof(f8675,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7444,f309])).
fof(f8674,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f7444,f331])).
fof(f8673,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f7444,f333])).
fof(f8672,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f7444,f351])).
fof(f8710,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8652,f418])).
fof(f8652,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f7444,f3403])).
fof(f8028,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X1))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(equality_resolution,[],[f387])).
fof(f8017,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.list.list'('type.closSem.v'),X48) != s('type.list.list'('type.closSem.v'),X50) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X52),s('type.list.list'('type.closSem.v'),X50))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X51),s('type.list.list'('type.closSem.v'),sK4(X48,X49,'type.list.list'('type.closSem.v'))))) | s('type.num.num',X51) != s('type.num.num',X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.closSem.v'),bool),X49)))) )),
  inference(superposition,[],[f387,f847])).
fof(f8093,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))),s('type.list.list'('type.closSem.v'),X0)))) )),
  inference(forward_demodulation,[],[f7941,f332])).
fof(f7941,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))),s('type.list.list'('type.closSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.closSem.v'),X2),'const.sum.INL_1'(s('type.list.list'('type.closSem.v'),X0)))))))) )),
  inference(unit_resulting_resolution,[],[f259,f332,f387])).
fof(f8092,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),X1)))) )),
  inference(forward_demodulation,[],[f7942,f332])).
fof(f7942,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.closSem.v'),X3),'const.sum.INL_1'(s('type.list.list'('type.closSem.v'),X1)))))))) )),
  inference(unit_resulting_resolution,[],[f815,f332,f387])).
fof(f8067,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X0),X1))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X0),X1)))))) )),
  inference(forward_demodulation,[],[f7967,f333])).
fof(f7967,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X0),X1))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.closSem.v'),X0),'type.list.list'('type.closSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.closSem.v'),X0),X1)))))) )),
  inference(unit_resulting_resolution,[],[f259,f333,f387])).
fof(f8066,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X1),X2))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X1),X2)))))) )),
  inference(forward_demodulation,[],[f7968,f333])).
fof(f7968,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X1),X2))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.closSem.v'),X1),'type.list.list'('type.closSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.closSem.v'),X1),X2)))))) )),
  inference(unit_resulting_resolution,[],[f815,f333,f387])).
fof(f8041,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X0),s('type.num.num',X1))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X0),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f7993,f317])).
fof(f7993,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X0),s('type.num.num',X1))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))),s('type.list.list'('type.closSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.closSem.v'),X3),s(fun('type.num.num','type.list.list'('type.closSem.v')),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f259,f317,f387])).
fof(f8040,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X1),s('type.num.num',X2))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f7994,f317])).
fof(f7994,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X1),s('type.num.num',X2))))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X3))))),s('type.list.list'('type.closSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.closSem.v'),X4),s(fun('type.num.num','type.list.list'('type.closSem.v')),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))))) )),
  inference(unit_resulting_resolution,[],[f815,f317,f387])).
fof(f7503,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f7245,f847])).
fof(f7502,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f7245,f332])).
fof(f7501,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f7245,f336])).
fof(f7500,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f7245,f335])).
fof(f7499,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f7245,f317])).
fof(f7498,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7245,f309])).
fof(f7497,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f7245,f331])).
fof(f7496,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f7245,f333])).
fof(f7495,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f7245,f351])).
fof(f7532,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f7475,f418])).
fof(f7475,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f7245,f3403])).
fof(f7452,plain,(
  ( ! [X61,X62,X63] : (s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f7395,f255])).
fof(f7395,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'F_0')) | s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f2103,f2099])).
fof(f7444,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7374,f255])).
fof(f7374,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f264,f2099])).
fof(f7360,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f2099,f847])).
fof(f7358,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f2099,f336])).
fof(f7357,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f2099,f335])).
fof(f7356,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f2099,f317])).
fof(f7355,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2099,f309])).
fof(f7353,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f2099,f333])).
fof(f7352,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2099,f351])).
fof(f7253,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f7196,f255])).
fof(f7196,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f2103,f2095])).
fof(f7245,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7176,f255])).
fof(f7176,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f264,f2095])).
fof(f7162,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f2095,f847])).
fof(f7160,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X108),s(bool,X109)))))) )),
  inference(superposition,[],[f2095,f336])).
fof(f7159,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X106,bool),'const.sum.INR_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f2095,f335])).
fof(f7158,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f2095,f317])).
fof(f7157,plain,(
  ( ! [X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2095,f309])).
fof(f7155,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f2095,f333])).
fof(f7154,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X86),bool),sK16(X87,X88,X86,bool)),s('type.list.list'(X86),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2095,f351])).
fof(f6669,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK16(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f5576,f847])).
fof(f6668,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK16(X145,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X142),'const.sum.INL_1'(s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f5576,f332])).
fof(f6667,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,X139)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK16(X141,'const.pair.SND_1'(s('type.pair.prod'(X137,bool),'const.pair.,_2'(s(X137,X138),s(bool,X139)))),X140,bool))))) )),
  inference(superposition,[],[f5576,f336])).
fof(f6666,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK16(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f5576,f335])).
fof(f6665,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,happ(s(fun('type.num.num',bool),X128),s('type.num.num',X129)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK16(X131,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129)))),X130,bool))))) )),
  inference(superposition,[],[f5576,f317])).
fof(f6664,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,X123)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X125),bool),sK16(X126,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.|48|_0')),X125,bool))))) )),
  inference(superposition,[],[f5576,f309])).
fof(f6663,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK16(X122,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X120)))),X121,bool))))) )),
  inference(superposition,[],[f5576,f331])).
fof(f6662,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK16(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f5576,f333])).
fof(f6661,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK16(X114,happ(s(fun('type.list.list'(X110),bool),sK16(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f5576,f351])).
fof(f6660,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK16(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f5576,f302])).
fof(f6693,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f6657,f253])).
fof(f6657,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f5576,f526])).
fof(f6690,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f6655,f253])).
fof(f6655,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f5576,f5664])).
fof(f6686,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f6652,f253])).
fof(f6652,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f5576,f847])).
fof(f6685,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6651,f253])).
fof(f6651,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f5576,f2198])).
fof(f6682,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(subsumption_resolution,[],[f6649,f253])).
fof(f6649,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(superposition,[],[f5576,f426])).
fof(f6681,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f6648,f253])).
fof(f6648,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f5576,f499])).
fof(f6680,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6647,f253])).
fof(f6647,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f5576,f458])).
fof(f6673,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f6672,f253])).
fof(f6672,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f6636,f2198])).
fof(f6636,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f5576,f2454])).
fof(f6670,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6631,f253])).
fof(f6631,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5576,f431])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : ! [X33,X24] : (! [X32] : (s(X9,X24) = s(X9,X32) => p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X33)))))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.JRH_INDUCT_UTIL')).
fof(f5422,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f5370,f253])).
fof(f5370,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f326,f847])).
fof(f5423,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f5371,f253])).
fof(f5371,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f404,f847])).
fof(f1111,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X73,bool),X74),X75)))))) | s(X73,X76) = s(X73,sK4(X76,happ(s(fun('type.pair.prod'(fun(X73,bool),X74),fun(X73,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X73,bool),X74),X75)),X73))) )),
  inference(superposition,[],[f311,f333])).
fof(f965,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK16(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f311,f351])).
fof(f840,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f311,f302])).
fof(f841,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X6) = s(X3,sK4(X6,'const.sum.OUTR_1'(s('type.sum.sum'(X4,fun(X3,bool)),'const.sum.INR_1'(s(fun(X3,bool),X5)))),X3))) )),
  inference(superposition,[],[f311,f331])).
fof(f2375,plain,(
  ( ! [X132,X130,X128,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X128,bool),happ(s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num',X131)))))) | s(X128,X132) = s(X128,sK4(X132,'const.ind_type.FCONS_3'(s(fun(X128,bool),X129),s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num','const.num.SUC_1'(s('type.num.num',X131)))),X128))) )),
  inference(superposition,[],[f311,f317])).
fof(f842,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8)))) | s(X7,X10) = s(X7,sK4(X10,'const.ind_type.FCONS_3'(s(fun(X7,bool),X8),s(fun('type.num.num',fun(X7,bool)),X9),s('type.num.num','const.num.|48|_0')),X7))) )),
  inference(superposition,[],[f311,f309])).
fof(f843,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f311,f335])).
fof(f844,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X19)))) | s(X16,X20) = s(X16,sK4(X20,'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))) )),
  inference(superposition,[],[f311,f336])).
fof(f845,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f311,f332])).
fof(f5333,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f311,f847])).
fof(f5297,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f264,f847])).
fof(f5298,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f265,f847])).
fof(f5299,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f268,f847])).
fof(f5300,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f269,f847])).
fof(f5301,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f323,f847])).
fof(f5303,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f326,f847])).
fof(f5304,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f404,f847])).
fof(f5309,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f433,f847])).
fof(f5311,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f480,f847])).
fof(f5316,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f2101,f847])).
fof(f5317,plain,(
  ( ! [X111,X109,X110] : (p(s(bool,X109)) | s(bool,X111) != s(bool,sK4(X109,X110,bool)) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f2103,f847])).
fof(f5471,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f5419,f847])).
fof(f5331,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f310,f847])).
fof(f5334,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f406,f847])).
fof(f5808,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f349,f847])).
fof(f5798,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f349,f847])).
fof(f5601,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f348,f847])).
fof(f5591,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f348,f847])).
fof(f5335,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (s(fun(X166,X167),X168) != s(fun(X166,X167),X170) | s(X167,happ(s(fun(X166,X167),X170),s(X166,X171))) = s(X167,happ(s(fun(X166,X167),sK4(X168,X169,fun(X166,X167))),s(X166,X171))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X166,X167),bool),X169)))) )),
  inference(superposition,[],[f346,f847])).
fof(f5348,plain,(
  ( ! [X218,X219,X217] : (s('type.integer.int',X217) != s('type.integer.int',X219) | s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X219))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',sK4(X217,X218,'type.integer.int')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X218)))) )),
  inference(superposition,[],[f328,f847])).
fof(f2339,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,X7) = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f317,f431])).
fof(f2341,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) = s(bool,happ(s(fun('type.num.num',bool),X13),s('type.num.num',X14))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X12),s(fun('type.num.num',bool),X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f317,f431])).
fof(f2368,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f458,f317])).
fof(f3511,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X31),s('type.num.num',X32))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) ) | $spl10),
  inference(forward_demodulation,[],[f3399,f2198])).
fof(f3399,plain,(
  ( ! [X30,X31,X32] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X31),s('type.num.num',X32))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32)))))) ) | $spl10),
  inference(superposition,[],[f2454,f317])).
fof(f4007,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X76) = s(bool,happ(s(fun('type.num.num',bool),X74),s('type.num.num',X75))) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X73),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75)))))) ) | $spl10),
  inference(superposition,[],[f3403,f317])).
fof(f4076,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X61),s('type.num.num',X62))) | s(bool,X63) = s(bool,'const.ind_type.FCONS_3'(s(bool,X60),s(fun('type.num.num',bool),X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3403,f317])).
fof(f4771,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) | s(bool,X78) = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77))))) | s(bool,'F_0') = s(bool,X78)) ) | $spl10),
  inference(superposition,[],[f4253,f317])).
fof(f4986,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f499,f317])).
fof(f5159,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X84),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X82),s('type.num.num',X83))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X81),s(fun('type.num.num',bool),X82),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83)))))) )),
  inference(superposition,[],[f526,f317])).
fof(f1088,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4)))) )),
  inference(superposition,[],[f333,f263])).
fof(f1479,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X37),X38)))) )),
  inference(superposition,[],[f431,f333])).
fof(f1521,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X29),X30))) | s(bool,X31) = s(bool,happ(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X29),X30))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f431,f333])).
fof(f2049,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f458,f333])).
fof(f2175,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29))))) )),
  inference(forward_demodulation,[],[f2153,f333])).
fof(f2153,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f310,f333])).
fof(f3508,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),X25))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X24),X25)))) ) | $spl10),
  inference(forward_demodulation,[],[f3396,f2198])).
fof(f3396,plain,(
  ( ! [X24,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),X25))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X24),X25)))) ) | $spl10),
  inference(superposition,[],[f2454,f333])).
fof(f4004,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65)))) ) | $spl10),
  inference(superposition,[],[f3403,f333])).
fof(f4073,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X51),X52))) | s(bool,X53) = s(bool,happ(s(fun('type.pair.prod'(bool,X51),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X51),X52))) | s(bool,'F_0') = s(bool,X53)) ) | $spl10),
  inference(superposition,[],[f3403,f333])).
fof(f4768,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,X68) = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f4253,f333])).
fof(f4983,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f499,f333])).
fof(f5156,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f526,f333])).
fof(f5535,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f5419,f333])).
fof(f6191,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) )),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f2375,f1111,f965,f845,f844,f843,f842,f841,f840,f311,f1097,f2175,f1521,f1479,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f2339,f2341,f1692,f1590,f2453,f2452,f2368,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f2049,f1088,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f5068,f5155,f5156,f5157,f5158,f5159,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5297,f5298,f5299,f5300,f5301,f5303,f5304,f5309,f5310,f5311,f5312,f5313,f5316,f5317,f5331,f5332,f5333,f5334,f5335,f5340,f5346,f5348,f5353,f5422,f5423,f5424,f5471,f5535,f5562,f5591,f5601,f5664,f5658,f5787,f5798,f5808,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871])).
fof(f5871,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5859,f253])).
fof(f5859,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5535,f431])).
fof(f5874,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5873,f253])).
fof(f5873,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f5864,f2198])).
fof(f5864,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f5535,f2454])).
fof(f4821,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,f(s(X33,sK12(X34,X35,bool,X33)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(X33,sK12(X34,X35,bool,X33))))) ) | $spl10),
  inference(superposition,[],[f345,f3403])).
fof(f4820,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') != s(bool,f(s(X29,sK12(X30,X31,bool,X29)))) | s(fun(X29,bool),X30) = s(fun(X29,bool),X31) | s(bool,X32) = s(bool,g(s(X29,sK12(X30,X31,bool,X29)))) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f345,f3403])).
fof(f4828,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK12(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK12(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f4818,f2198])).
fof(f4818,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK12(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK12(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f345,f2454])).
fof(f4814,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK12(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK12(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f345,f431])).
fof(f4813,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK12(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK12(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f345,f431])).
fof(f4812,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f345,f263])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ! [X33,X32] : (p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32)))) => p(s(bool,happ(s(fun(X9,bool),X33),s(X9,'const.min.@_1'(s(fun(X9,bool),X33)))))))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.SELECT_AX')).
fof(f313,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X33] : (! [X32] : (p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X32))))) => (? [X32] : p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32)))) => ? [X32] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X32))))))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.MONO_EXISTS')).
fof(f314,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f340,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X31] : ! [X54,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X54)))) => (! [X33] : (! [X32] : p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32)))) <=> ! [X32] : p(s(bool,happ(s(fun(X9,bool),X33),s(X9,happ(s(fun(X31,X9),X54),s(X31,X32))))))) & ! [X33] : (? [X32] : p(s(bool,happ(s(fun(X9,bool),X33),s(X9,X32)))) <=> ? [X32] : p(s(bool,happ(s(fun(X9,bool),X33),s(X9,happ(s(fun(X31,X9),X54),s(X31,X32))))))) & ! [X6,X56] : (s(X9,X6) = s(X9,happ(s(fun(X31,X9),X54),s(X31,X56))) <=> s(X31,X56) = s(X31,happ(s(fun(X9,X31),X10),s(X9,X6))))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.ISO_USAGE')).
fof(f341,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f2171,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2)))))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2145,f253])).
fof(f2145,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2)))))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f310,f431])).
fof(f3637,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) ) | $spl10),
  inference(subsumption_resolution,[],[f3636,f253])).
fof(f3636,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) ) | $spl10),
  inference(forward_demodulation,[],[f3618,f2198])).
fof(f3618,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) ) | $spl10),
  inference(superposition,[],[f310,f2454])).
fof(f5419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f5281,f264])).
fof(f5281,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f312,f847])).
fof(f5555,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5522,f253])).
fof(f5522,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f5419,f431])).
fof(f5558,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5557,f253])).
fof(f5557,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f5527,f2198])).
fof(f5527,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f5419,f2454])).
fof(f339,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f342,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : ! [X33] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X33),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X24] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X33),s('type.list.list'(X9),X24)))) => ! [X34] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X33),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X34),s('type.list.list'(X9),X24)))))))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X33),s('type.list.list'(X9),X35)))))),
  file('thm.closSem.v_|49||49||split|3','thm.list.list_INDUCT')).
fof(f406,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f2728,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2713,f302])).
fof(f2713,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f302])).
fof(f2729,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f2714,f351])).
fof(f2714,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f351])).
fof(f2731,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9)))))) )),
  inference(forward_demodulation,[],[f2716,f331])).
fof(f2716,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11))))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f331])).
fof(f2732,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f2717,f309])).
fof(f2717,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f309])).
fof(f2734,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f2719,f335])).
fof(f2719,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f335])).
fof(f2735,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23)))))) )),
  inference(forward_demodulation,[],[f2720,f336])).
fof(f2720,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f336])).
fof(f2736,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27)))))) )),
  inference(forward_demodulation,[],[f2721,f332])).
fof(f2721,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29))))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f332])).
fof(f4698,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,happ(s(fun('type.list.list'(X79),bool),X80),s('type.list.list'(X79),'const.list.NIL_0')))) | s(bool,X81) = s(bool,happ(s(fun('type.list.list'(X79),bool),X80),s('type.list.list'(X79),sK2(X80,X79)))) | s(bool,'F_0') = s(bool,X81)) ) | $spl10),
  inference(subsumption_resolution,[],[f4312,f255])).
fof(f4312,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X79),bool),X80),s('type.list.list'(X79),'const.list.NIL_0')))) | s(bool,X81) = s(bool,happ(s(fun('type.list.list'(X79),bool),X80),s('type.list.list'(X79),sK2(X80,X79)))) | s(bool,'F_0') = s(bool,X81)) ) | $spl10),
  inference(superposition,[],[f406,f3403])).
fof(f2733,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15)))))) )),
  inference(forward_demodulation,[],[f2718,f317])).
fof(f2718,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15))))) | ~p(s(bool,happ(s(fun('type.list.list'(X15),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.list.list'(X15),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f317])).
fof(f2730,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f2715,f333])).
fof(f2715,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f406,f333])).
fof(f2152,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.min.@_1'(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)))))))) )),
  inference(superposition,[],[f310,f351])).
fof(f389,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X31,X51,X9,X52] : ! [X53,X54,X55,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X52),X10),s(fun(X52,X9),X55)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X31,X51),X54),s(fun(X51,X31),X53))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X31),fun(X52,X51)),'_dst_x00x11_2'(s(fun(X31,X51),X54),s(fun(fun(X9,X31),fun(X52,X31)),'_dstsel_0_x01_1'(s(fun(X52,X9),X55))))),s(fun(fun(X52,X51),fun(X9,X31)),'_dst_x00x11_2'(s(fun(X51,X31),X53),s(fun(fun(X52,X51),fun(X9,X51)),'_dstsel_0_x01_1'(s(fun(X9,X52),X10)))))))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.ISO_FUN')).
fof(f343,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f235])).
fof(f344,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f235])).
fof(f348,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0)))))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X31,X9] : ! [X10,X54] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X54)))) <=> (! [X32] : s(X31,X32) = s(X31,happ(s(fun(X9,X31),X10),s(X9,happ(s(fun(X31,X9),X54),s(X31,X32))))) & ! [X39] : s(X9,X39) = s(X9,happ(s(fun(X31,X9),X54),s(X31,happ(s(fun(X9,X31),X10),s(X9,X39)))))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.ISO')).
fof(f349,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f350,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f5635,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5604,f253])).
fof(f5604,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f348,f431])).
fof(f5638,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5637,f253])).
fof(f5637,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f5609,f2198])).
fof(f5609,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f348,f2454])).
fof(f5842,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5811,f253])).
fof(f5811,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f349,f431])).
fof(f5845,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5844,f253])).
fof(f5844,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f5816,f2198])).
fof(f5816,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f349,f2454])).
fof(f5260,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f847,f302])).
fof(f5261,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f847,f351])).
fof(f5263,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))) | s(X12,X15) = s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))) )),
  inference(superposition,[],[f847,f331])).
fof(f5264,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),X17))) | s(X16,X19) = s(X16,sK4(X19,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0')),X16))) )),
  inference(superposition,[],[f847,f309])).
fof(f5266,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f847,f335])).
fof(f5267,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X33))) | s(X30,X34) = s(X30,sK4(X34,'const.pair.SND_1'(s('type.pair.prod'(X31,fun(X30,bool)),'const.pair.,_2'(s(X31,X32),s(fun(X30,bool),X33)))),X30))) )),
  inference(superposition,[],[f847,f336])).
fof(f5268,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X37))) | s(X35,X38) = s(X35,sK4(X38,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X35,bool),X36),'const.sum.INL_1'(s(fun(X35,bool),X37)))),X35))) )),
  inference(superposition,[],[f847,f332])).
fof(f5269,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f847,f847])).
fof(f5265,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),happ(s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num',X23))))) | s(X20,X24) = s(X20,sK4(X24,'const.ind_type.FCONS_3'(s(fun(X20,bool),X21),s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X20))) )),
  inference(superposition,[],[f847,f317])).
fof(f5262,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f847,f333])).
fof(f5581,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f5572,f255])).
fof(f5572,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f5534,f3403])).
fof(f5576,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f5564,f255])).
fof(f5564,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f5534,f263])).
fof(f5534,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f5419,f351])).
fof(f5665,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f5562,f5534])).
fof(f5787,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5740,f253])).
fof(f5740,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f433,f5664])).
fof(f5710,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f5675,f255])).
fof(f5675,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f5562,f3403])).
fof(f5658,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f351,f5562,f433])).
fof(f5664,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f418,f5562,f2101])).
fof(f5562,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f253,f5534])).
fof(f5424,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f5374,f253])).
fof(f5374,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f433,f847])).
fof(f5353,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X233),s('type.list.list'('type.closSem.v'),X234))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X231))) | s('type.list.list'('type.closSem.v'),X234) = s('type.list.list'('type.closSem.v'),sK4(X231,X232,'type.list.list'('type.closSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.closSem.v'),bool),X232)))) )),
  inference(superposition,[],[f386,f847])).
fof(f5346,plain,(
  ( ! [X212,X213,X211] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X211))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X213))) | s('type.integer.int',X213) = s('type.integer.int',sK4(X211,X212,'type.integer.int')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X212)))) )),
  inference(superposition,[],[f327,f847])).
fof(f5340,plain,(
  ( ! [X191,X189,X187,X190,X188] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X187),s('type.list.list'('type.closSem.v'),X189))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X190),s('type.list.list'('type.closSem.v'),X191))) | s('type.num.num',X190) = s('type.num.num',sK4(X187,X188,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X188)))) )),
  inference(superposition,[],[f385,f847])).
fof(f5332,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f312,f847])).
fof(f5328,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f4253,f847])).
fof(f5420,plain,(
  ( ! [X112,X113] : (s(bool,'T_0') != s(bool,X112) | s(bool,'T_0') = s(bool,sK4(X112,X113,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) ) | $spl10),
  inference(forward_demodulation,[],[f5318,f2198])).
fof(f5318,plain,(
  ( ! [X112,X113] : (s(bool,X112) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X112,X113,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) ) | $spl10),
  inference(superposition,[],[f2279,f847])).
fof(f5313,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f526,f847])).
fof(f5312,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f499,f847])).
fof(f5310,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f458,f847])).
fof(f5257,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f847,f3403])).
fof(f5256,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f847,f3403])).
fof(f5418,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f5254,f2198])).
fof(f5254,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f847,f2454])).
fof(f5250,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f847,f431])).
fof(f5249,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f847,f431])).
fof(f5248,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f847,f263])).
fof(f5242,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5185,f253])).
fof(f5185,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f433,f526])).
fof(f5162,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,X94)))))) )),
  inference(superposition,[],[f526,f332])).
fof(f5161,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f526,f336])).
fof(f5160,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f526,f335])).
fof(f5158,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,X78))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f526,f309])).
fof(f5157,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X77),'const.sum.INL_1'(s(bool,X76))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76)))))) )),
  inference(superposition,[],[f526,f331])).
fof(f5155,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK16(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f526,f351])).
fof(f5068,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5012,f253])).
fof(f5012,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f433,f499])).
fof(f4989,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f499,f332])).
fof(f4988,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) )),
  inference(superposition,[],[f499,f336])).
fof(f4987,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X85,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f499,f335])).
fof(f4985,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f499,f309])).
fof(f4984,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,X73))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,X73)))))) )),
  inference(superposition,[],[f499,f331])).
fof(f4982,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X68,bool),'const.sum.INR_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK16(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f499,f351])).
fof(f4809,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,f(s(bool,X36))) != s(X35,g(s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK12(X33,X34,X35,bool))) ) | $spl10),
  inference(superposition,[],[f345,f3403])).
fof(f4808,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK12(X29,X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f345,f3403])).
fof(f4825,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK12(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4806,f2198])).
fof(f4806,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK12(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f345,f2454])).
fof(f4802,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK12(X7,X8,X9,bool))) )),
  inference(superposition,[],[f345,f431])).
fof(f4801,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK12(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f345,f431])).
fof(f4800,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,X2,bool))) )),
  inference(superposition,[],[f345,f263])).
fof(f4774,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') != s(bool,X88) | s(bool,X89) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X87),'const.sum.INL_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,X89)) ) | $spl10),
  inference(superposition,[],[f4253,f332])).
fof(f4773,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X86) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X85))))) | s(bool,'F_0') = s(bool,X86)) ) | $spl10),
  inference(superposition,[],[f4253,f336])).
fof(f4772,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,X82) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81))))) | s(bool,'F_0') = s(bool,X82)) ) | $spl10),
  inference(superposition,[],[f4253,f335])).
fof(f4770,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') != s(bool,X72) | s(bool,X74) = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4253,f309])).
fof(f4769,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') != s(bool,X70) | s(bool,X71) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,X71)) ) | $spl10),
  inference(superposition,[],[f4253,f331])).
fof(f4767,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X64) | s(bool,X65) = s(bool,happ(s(fun('type.list.list'(X62),bool),sK16(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f4253,f351])).
fof(f4766,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,X61) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60))) | s(bool,'F_0') = s(bool,X61)) ) | $spl10),
  inference(superposition,[],[f4253,f302])).
fof(f4734,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f4253,f431])).
fof(f4343,plain,(
  ( ! [X173,X171,X172] : (~p(s(bool,happ(s(fun(bool,bool),X172),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172)))) | s(bool,X173) = s(bool,sK4(X171,X172,bool)) | s(bool,'F_0') = s(bool,X173)) ) | $spl10),
  inference(superposition,[],[f312,f3403])).
fof(f4701,plain,(
  ( ! [X94,X97,X95,X96] : (s(X94,X96) = s(X94,sK4(X96,X95,X94)) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun(X94,bool),X95))) | s(bool,'F_0') = s(bool,X97)) ) | $spl10),
  inference(subsumption_resolution,[],[f4320,f255])).
fof(f4320,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,'F_0')) | s(X94,X96) = s(X94,sK4(X96,X95,X94)) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun(X94,bool),X95))) | s(bool,'F_0') = s(bool,X97)) ) | $spl10),
  inference(superposition,[],[f311,f3403])).
fof(f4690,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f4286,f255])).
fof(f4286,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f323,f3403])).
fof(f4253,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'F_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f418,f3403])).
fof(f4250,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f268,f3403])).
fof(f4182,plain,(
  ( ! [X185,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X184),s(bool,X185)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X184)))) | s(bool,'F_0') = s(bool,X185) | s(bool,'F_0') = s(bool,sK4(X183,X184,bool))) ) | $spl10),
  inference(superposition,[],[f312,f3403])).
fof(f4079,plain,(
  ( ! [X74,X72,X73] : (s(bool,'F_0') = s(bool,X73) | s(bool,X74) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,X73))))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f3403,f332])).
fof(f4078,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'F_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70))))) | s(bool,'F_0') = s(bool,X71)) ) | $spl10),
  inference(superposition,[],[f3403,f336])).
fof(f4077,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') = s(bool,X65) | s(bool,X67) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),'const.pair.,_2'(s(bool,X65),s(X64,X66))))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3403,f335])).
fof(f4075,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') = s(bool,X57) | s(bool,X59) = s(bool,'const.ind_type.FCONS_3'(s(bool,X57),s(fun('type.num.num',bool),X58),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X59)) ) | $spl10),
  inference(superposition,[],[f3403,f309])).
fof(f4074,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X56) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X54,bool),'const.sum.INR_1'(s(bool,X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl10),
  inference(superposition,[],[f3403,f331])).
fof(f4072,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,happ(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool)),s('type.list.list'(X47),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X50)) ) | $spl10),
  inference(superposition,[],[f3403,f351])).
fof(f4044,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3403,f3403])).
fof(f4010,plain,(
  ( ! [X87,X85,X86] : (s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X86)))))) ) | $spl10),
  inference(superposition,[],[f3403,f332])).
fof(f4009,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X84) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X81,bool),'const.pair.,_2'(s(X81,X82),s(bool,X83)))))) ) | $spl10),
  inference(superposition,[],[f3403,f336])).
fof(f4008,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X78) = s(bool,X80) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X77),'const.pair.,_2'(s(bool,X78),s(X77,X79)))))) ) | $spl10),
  inference(superposition,[],[f3403,f335])).
fof(f4006,plain,(
  ( ! [X70,X72,X71] : (s(bool,X70) = s(bool,X72) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X70),s(fun('type.num.num',bool),X71),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f3403,f309])).
fof(f4005,plain,(
  ( ! [X68,X69,X67] : (s(bool,X68) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))))) ) | $spl10),
  inference(superposition,[],[f3403,f331])).
fof(f4003,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X62) = s(bool,X63) | s(bool,'F_0') = s(bool,X63) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X60),bool),sK16(X61,X62,X60,bool)),s('type.list.list'(X60),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f3403,f351])).
fof(f4002,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)))) ) | $spl10),
  inference(superposition,[],[f3403,f302])).
fof(f3555,plain,(
  ( ! [X94,X93] : (~p(s(bool,happ(s(fun(bool,bool),X94),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X94)))) | s(bool,'F_0') = s(bool,sK4(X93,X94,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f3462,f2198])).
fof(f3462,plain,(
  ( ! [X94,X93] : (~p(s(bool,happ(s(fun(bool,bool),X94),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X94)))) | s(bool,'F_0') = s(bool,sK4(X93,X94,bool))) ) | $spl10),
  inference(superposition,[],[f312,f2454])).
fof(f3541,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36)))) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK4(X37,X36,X35))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3540,f253])).
fof(f3540,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36)))) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK4(X37,X36,X35))))) ) | $spl10),
  inference(forward_demodulation,[],[f3433,f2198])).
fof(f3433,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36)))) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK4(X37,X36,X35))))) ) | $spl10),
  inference(superposition,[],[f312,f2454])).
fof(f3403,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2454,f2454])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f499,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f331,f263])).
fof(f526,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f332,f263])).
fof(f619,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f309,f263])).
fof(f672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f335,f263])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f336,f263])).
fof(f944,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool)),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f351,f263])).
fof(f999,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool))) )),
  inference(superposition,[],[f312,f263])).
fof(f2094,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2048,f302])).
fof(f2048,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f458,f351])).
fof(f2095,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28))))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f2050,f302])).
fof(f2050,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28)))))) )),
  inference(superposition,[],[f458,f331])).
fof(f2096,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2051,f302])).
fof(f2051,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f458,f309])).
fof(f2097,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2052,f302])).
fof(f2052,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f458,f335])).
fof(f2098,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(forward_demodulation,[],[f2053,f302])).
fof(f2053,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(superposition,[],[f458,f336])).
fof(f2099,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f2054,f302])).
fof(f2054,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38)))))) )),
  inference(superposition,[],[f458,f332])).
fof(f2103,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2102,f302])).
fof(f2102,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2069,f253])).
fof(f2069,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f404,f458])).
fof(f2104,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2072,f253])).
fof(f2072,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f433,f458])).
fof(f2452,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2429,f253])).
fof(f2429,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2101,f458])).
fof(f2453,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2430,f253])).
fof(f2430,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2101,f426])).
fof(f1590,plain,(
  ( ! [X118,X116,X117] : (~p(s(bool,happ(s(fun(bool,bool),X117),s(bool,X118)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X117)))) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,sK4(X116,X117,bool))) )),
  inference(superposition,[],[f312,f431])).
fof(f1692,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) | s(bool,X114) = s(bool,sK4(X112,X113,bool)) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f312,f431])).
fof(f1961,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42)))) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK4(X43,X42,X41)))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f1670,f253])).
fof(f1670,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42)))) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK4(X43,X42,X41)))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f312,f431])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f1002,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f992,f302])).
fof(f992,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f312,f302])).
fof(f1003,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f993,f351])).
fof(f993,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f312,f351])).
fof(f1004,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8)))))) )),
  inference(forward_demodulation,[],[f994,f331])).
fof(f994,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10))))))))) )),
  inference(superposition,[],[f312,f331])).
fof(f1005,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13)))) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f995,f309])).
fof(f995,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f312,f309])).
fof(f1006,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f996,f335])).
fof(f996,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f312,f335])).
fof(f1007,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24)))) | ~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21)))))) )),
  inference(forward_demodulation,[],[f997,f336])).
fof(f997,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24))))))))) )),
  inference(superposition,[],[f312,f336])).
fof(f1008,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26)))))) )),
  inference(forward_demodulation,[],[f998,f332])).
fof(f998,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28))))))))) )),
  inference(superposition,[],[f312,f332])).
fof(f2412,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126)))))) | ~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123)))))) )),
  inference(forward_demodulation,[],[f2374,f317])).
fof(f2374,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123))))) | p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))))))))) )),
  inference(superposition,[],[f312,f317])).
fof(f1131,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X69,bool),X70),X71)))))) | ~p(s(bool,happ(s(fun(X69,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X69,bool),X70),X71))),s(X69,sK4(X72,happ(s(fun('type.pair.prod'(fun(X69,bool),X70),fun(X69,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X69,bool),X70),X71)),X69)))))) )),
  inference(forward_demodulation,[],[f1110,f333])).
fof(f1110,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,happ(s(fun(X69,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X69,bool),X70),X71))),s(X69,sK4(X72,happ(s(fun('type.pair.prod'(fun(X69,bool),X70),fun(X69,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X69,bool),X70),X71)),X69))))) | p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),happ(s(fun('type.pair.prod'(fun(X69,bool),X70),fun(X69,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X69,bool),X70),X71))))))) )),
  inference(superposition,[],[f312,f333])).
fof(f1097,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f312,f333])).
fof(f2561,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X44))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X46),s('type.list.list'('type.closSem.v'),X47))) | s('type.list.list'('type.closSem.v'),X47) = s('type.list.list'('type.closSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.closSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.closSem.v'),X44)))))) )),
  inference(superposition,[],[f386,f332])).
fof(f2560,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X40),s('type.list.list'('type.closSem.v'),X39))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X41),s('type.list.list'('type.closSem.v'),X42))) | s('type.list.list'('type.closSem.v'),X42) = s('type.list.list'('type.closSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.closSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.closSem.v'),X39)))))) )),
  inference(superposition,[],[f386,f336])).
fof(f2559,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X34),s('type.list.list'('type.closSem.v'),X32))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X35),s('type.list.list'('type.closSem.v'),X36))) | s('type.list.list'('type.closSem.v'),X36) = s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f386,f335])).
fof(f2558,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X29),s('type.list.list'('type.closSem.v'),X30))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),happ(s(fun('type.num.num','type.list.list'('type.closSem.v')),X26),s('type.num.num',X27))))) | s('type.list.list'('type.closSem.v'),X30) = s('type.list.list'('type.closSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.closSem.v'),X25),s(fun('type.num.num','type.list.list'('type.closSem.v')),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) )),
  inference(superposition,[],[f386,f317])).
fof(f2557,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X22),s('type.list.list'('type.closSem.v'),X20))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X23),s('type.list.list'('type.closSem.v'),X24))) | s('type.list.list'('type.closSem.v'),X24) = s('type.list.list'('type.closSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.closSem.v'),X20),s(fun('type.num.num','type.list.list'('type.closSem.v')),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f309])).
fof(f2556,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X17),s('type.list.list'('type.closSem.v'),X16))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X18),s('type.list.list'('type.closSem.v'),X19))) | s('type.list.list'('type.closSem.v'),X19) = s('type.list.list'('type.closSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X15,'type.list.list'('type.closSem.v')),'const.sum.INR_1'(s('type.list.list'('type.closSem.v'),X16)))))) )),
  inference(superposition,[],[f386,f331])).
fof(f2555,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X13),s('type.list.list'('type.closSem.v'),X14))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X12),s('type.list.list'('type.closSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.closSem.v'),X10),X11))))) | s('type.list.list'('type.closSem.v'),X14) = s('type.list.list'('type.closSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.closSem.v'),X10),'type.list.list'('type.closSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.closSem.v'),X10),X11)))) )),
  inference(superposition,[],[f386,f333])).
fof(f2554,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X7),s('type.list.list'('type.closSem.v'),X6))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X8),s('type.list.list'('type.closSem.v'),X9))) | s('type.list.list'('type.closSem.v'),X9) = s('type.list.list'('type.closSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.closSem.v')),sK16(X5,X6,X4,'type.list.list'('type.closSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f386,f351])).
fof(f2553,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X0))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.list.list'('type.closSem.v'),X3) = s('type.list.list'('type.closSem.v'),happ(s(fun('type.list.list'('type.closSem.v'),'type.list.list'('type.closSem.v')),'_i_0_0'),s('type.list.list'('type.closSem.v'),X0)))) )),
  inference(superposition,[],[f386,f302])).
fof(f2446,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f2420,f253])).
fof(f2420,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2101,f431])).
fof(f2382,plain,(
  ( ! [X158,X159,X157,X160] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X160))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',happ(s(fun('type.num.num','type.integer.int'),X158),s('type.num.num',X159))))) | s('type.integer.int',X160) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X157),s(fun('type.num.num','type.integer.int'),X158),s('type.num.num','const.num.SUC_1'(s('type.num.num',X159)))))) )),
  inference(superposition,[],[f327,f317])).
fof(f2376,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X136),s('type.list.list'('type.closSem.v'),X137))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X134),s('type.num.num',X135))),s('type.list.list'('type.closSem.v'),X138))) | s('type.num.num',X136) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X133),s(fun('type.num.num','type.num.num'),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))))) )),
  inference(superposition,[],[f385,f317])).
fof(f2279,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2252,f253])).
fof(f2252,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f433,f2198])).
fof(f2179,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f351,f2170,f433])).
fof(f1982,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X64),s('type.list.list'('type.closSem.v'),X65))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X66),s('type.list.list'('type.closSem.v'),X67))) | s('type.num.num',X66) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X63),'const.sum.INL_1'(s('type.num.num',X64)))))) )),
  inference(superposition,[],[f385,f332])).
fof(f1981,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X59),s('type.list.list'('type.closSem.v'),X60))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X61),s('type.list.list'('type.closSem.v'),X62))) | s('type.num.num',X61) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X57,'type.num.num'),'const.pair.,_2'(s(X57,X58),s('type.num.num',X59)))))) )),
  inference(superposition,[],[f385,f336])).
fof(f1980,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X52),s('type.list.list'('type.closSem.v'),X54))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X55),s('type.list.list'('type.closSem.v'),X56))) | s('type.num.num',X55) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X51),'const.pair.,_2'(s('type.num.num',X52),s(X51,X53)))))) )),
  inference(superposition,[],[f385,f335])).
fof(f1979,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X46),s('type.list.list'('type.closSem.v'),X48))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X49),s('type.list.list'('type.closSem.v'),X50))) | s('type.num.num',X49) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X46),s(fun('type.num.num','type.num.num'),X47),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f385,f309])).
fof(f1978,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X42),s('type.list.list'('type.closSem.v'),X43))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X44),s('type.list.list'('type.closSem.v'),X45))) | s('type.num.num',X44) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X41,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X42)))))) )),
  inference(superposition,[],[f385,f331])).
fof(f1977,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),X40))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X36),X37))),s('type.list.list'('type.closSem.v'),X38))) | s('type.num.num',X39) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X36),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X36),X37)))) )),
  inference(superposition,[],[f385,f333])).
fof(f1976,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X32),s('type.list.list'('type.closSem.v'),X33))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X34),s('type.list.list'('type.closSem.v'),X35))) | s('type.num.num',X34) = s('type.num.num',happ(s(fun('type.list.list'(X30),'type.num.num'),sK16(X31,X32,X30,'type.num.num')),s('type.list.list'(X30),'const.list.NIL_0')))) )),
  inference(superposition,[],[f385,f351])).
fof(f1975,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X26),s('type.list.list'('type.closSem.v'),X27))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),X29))) | s('type.num.num',X28) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X26)))) )),
  inference(superposition,[],[f385,f302])).
fof(f1973,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X20),s('type.list.list'('type.closSem.v'),X21))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),X19))) | s('type.num.num',X20) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X18)))))) )),
  inference(superposition,[],[f385,f261])).
fof(f1972,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X17))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.list.list'('type.closSem.v'),X15))) | s('type.num.num',X16) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f385,f259])).
fof(f1970,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X7),s('type.list.list'('type.closSem.v'),X8))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X5))))),s('type.list.list'('type.closSem.v'),X6))) | s('type.num.num',X7) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X4)))))) )),
  inference(superposition,[],[f385,f815])).
fof(f1969,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.list.list'('type.closSem.v'),X1))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f385,f259])).
fof(f1959,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1661,f253])).
fof(f1661,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f433,f431])).
fof(f1958,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1660,f253])).
fof(f1660,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f404,f431])).
fof(f1957,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1659,f253])).
fof(f1659,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f326,f431])).
fof(f1526,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X46),'const.sum.INL_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f431,f332])).
fof(f1525,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X44))))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f431,f336])).
fof(f1524,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X38),'const.pair.,_2'(s(bool,X39),s(X38,X40))))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f431,f335])).
fof(f1523,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.ind_type.FCONS_3'(s(bool,X35),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f431,f309])).
fof(f1522,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f431,f331])).
fof(f1520,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f431,f351])).
fof(f1519,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f431,f302])).
fof(f1506,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f431,f431])).
fof(f1484,plain,(
  ( ! [X54,X56,X55] : (s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X54),'const.sum.INL_1'(s(bool,X55)))))) )),
  inference(superposition,[],[f431,f332])).
fof(f1483,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X51),s(bool,X52)))))) )),
  inference(superposition,[],[f431,f336])).
fof(f1482,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,X47),s(X46,X48)))))) )),
  inference(superposition,[],[f431,f335])).
fof(f1481,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X43),s(fun('type.num.num',bool),X44),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f431,f309])).
fof(f1480,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X40,bool),'const.sum.INR_1'(s(bool,X41)))))) )),
  inference(superposition,[],[f431,f331])).
fof(f1478,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)),s('type.list.list'(X33),'const.list.NIL_0')))) )),
  inference(superposition,[],[f431,f351])).
fof(f1477,plain,(
  ( ! [X31,X32] : (s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) )),
  inference(superposition,[],[f431,f302])).
fof(f1219,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f260,f259])).
fof(f1115,plain,(
  ( ! [X85,X83,X84] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X85))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X83),X84))))) | s('type.integer.int',X85) = s('type.integer.int',happ(s(fun('type.pair.prod'('type.integer.int',X83),'type.integer.int'),'const.pair.FST_0'),s('type.pair.prod'('type.integer.int',X83),X84)))) )),
  inference(superposition,[],[f327,f333])).
fof(f969,plain,(
  ( ! [X103,X101,X102,X104] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X104))) | s('type.integer.int',X104) = s('type.integer.int',happ(s(fun('type.list.list'(X101),'type.integer.int'),sK16(X102,X103,X101,'type.integer.int')),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f327,f351])).
fof(f891,plain,(
  ( ! [X17,X18,X16] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X17))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X18))) | s('type.integer.int',X18) = s('type.integer.int','const.sum.OUTL_1'(s('type.sum.sum'('type.integer.int',X16),'const.sum.INL_1'(s('type.integer.int',X17)))))) )),
  inference(superposition,[],[f327,f332])).
fof(f890,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X15))) | s('type.integer.int',X15) = s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'(X12,'type.integer.int'),'const.pair.,_2'(s(X12,X13),s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f327,f336])).
fof(f889,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X9))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X8),'const.pair.,_2'(s('type.integer.int',X9),s(X8,X10)))))) )),
  inference(superposition,[],[f327,f335])).
fof(f888,plain,(
  ( ! [X6,X7,X5] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X5))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X5),s(fun('type.num.num','type.integer.int'),X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f327,f309])).
fof(f887,plain,(
  ( ! [X4,X2,X3] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X3))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.sum.OUTR_1'(s('type.sum.sum'(X2,'type.integer.int'),'const.sum.INR_1'(s('type.integer.int',X3)))))) )),
  inference(superposition,[],[f327,f331])).
fof(f886,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int',happ(s(fun('type.integer.int','type.integer.int'),'_i_0_0'),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f327,f302])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f846,f255])).
fof(f846,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f311,f263])).
fof(f815,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))))) )),
  inference(superposition,[],[f259,f259])).
fof(f789,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f776,f253])).
fof(f776,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f433,f426])).
fof(f480,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f475,f255])).
fof(f475,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f323,f263])).
fof(f433,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f268,f263])).
fof(f426,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f256,f265])).
fof(f418,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f255,f264])).
fof(f392,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Recclosure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X3),s('type.num.num',X4))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))))),'const.closSem.Recclosure_0'),s('type.option.option'('type.num.num'),X0))),s('type.list.list'('type.closSem.v'),X1))),s('type.list.list'('type.closSem.v'),X2))),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X3))),s('type.num.num',X4)))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.closSem.v','const.closSem.Recclosure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X3),s('type.num.num',X4))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))))),'const.closSem.Recclosure_0'),s('type.option.option'('type.num.num'),X0))),s('type.list.list'('type.closSem.v'),X1))),s('type.list.list'('type.closSem.v'),X2))),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X3))),s('type.num.num',X4)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X20,X21,X22,X23] : s('type.closSem.v','const.closSem.Recclosure_5'(s('type.option.option'('type.num.num'),X8),s('type.list.list'('type.closSem.v'),X20),s('type.list.list'('type.closSem.v'),X21),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X22),s('type.num.num',X23))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.num.num','type.closSem.v'))))),'const.closSem.Recclosure_0'),s('type.option.option'('type.num.num'),X8))),s('type.list.list'('type.closSem.v'),X20))),s('type.list.list'('type.closSem.v'),X21))),s('type.list.list'('type.pair.prod'('type.num.num','type.closLang.exp')),X22))),s('type.num.num',X23)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_5_const.closSem.Recclosure')).
fof(f391,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.num.num',X3),s('type.closLang.exp',X4))) = s('type.closSem.v',happ(s(fun('type.closLang.exp','type.closSem.v'),happ(s(fun('type.num.num',fun('type.closLang.exp','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))))),'const.closSem.Closure_0'),s('type.option.option'('type.num.num'),X0))),s('type.list.list'('type.closSem.v'),X1))),s('type.list.list'('type.closSem.v'),X2))),s('type.num.num',X3))),s('type.closLang.exp',X4)))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.num.num',X3),s('type.closLang.exp',X4))) = s('type.closSem.v',happ(s(fun('type.closLang.exp','type.closSem.v'),happ(s(fun('type.num.num',fun('type.closLang.exp','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))))),'const.closSem.Closure_0'),s('type.option.option'('type.num.num'),X0))),s('type.list.list'('type.closSem.v'),X1))),s('type.list.list'('type.closSem.v'),X2))),s('type.num.num',X3))),s('type.closLang.exp',X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X20,X21,X22,X23] : s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X8),s('type.list.list'('type.closSem.v'),X20),s('type.list.list'('type.closSem.v'),X21),s('type.num.num',X22),s('type.closLang.exp',X23))) = s('type.closSem.v',happ(s(fun('type.closLang.exp','type.closSem.v'),happ(s(fun('type.num.num',fun('type.closLang.exp','type.closSem.v')),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))),happ(s(fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v')))),happ(s(fun('type.option.option'('type.num.num'),fun('type.list.list'('type.closSem.v'),fun('type.list.list'('type.closSem.v'),fun('type.num.num',fun('type.closLang.exp','type.closSem.v'))))),'const.closSem.Closure_0'),s('type.option.option'('type.num.num'),X8))),s('type.list.list'('type.closSem.v'),X20))),s('type.list.list'('type.closSem.v'),X21))),s('type.num.num',X22))),s('type.closLang.exp',X23)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_5_const.closSem.Closure')).
fof(f390,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.num.num',X3),s('type.closLang.exp',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X4)))))))))))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X2),s('type.num.num',X3),s('type.closLang.exp',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X4)))))))))))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X16,X17,X59,X60,X61] : s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Closure_5'(s('type.option.option'('type.num.num'),X16),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X59),s('type.num.num',X60),s('type.closLang.exp',X61))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X17))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X59))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X60),s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X61)))))))))))))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_size_def|split|4')).
fof(f388,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X25,X26,X46,X27] : ! [X28,X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X46,X25),X28),s(fun(X26,X46),happ(s(fun(X27,fun(X26,X46)),X29),s(X27,X30))))) = s(fun(X26,X25),happ(s(fun(X27,fun(X26,X25)),'_dst_x00x11_2'(s(fun(X46,X25),X28),s(fun(X27,fun(X26,X46)),X29))),s(X27,X30)))),
  file('thm.closSem.v_|49||49||split|3','_dst_x00x11_ax')).
fof(f386,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3)) )),
  inference(cnf_transformation,[],[f250])).
fof(f387,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.closSem.v'),X1) != s('type.list.list'('type.closSem.v'),X3) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f404,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f368])).
fof(f368,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X32,X49,X39,X50] : ((s(bool,X32) = s(bool,X49) & (p(s(bool,X49)) => s(bool,X39) = s(bool,X50))) => ((p(s(bool,X32)) => p(s(bool,X39))) <=> (p(s(bool,X49)) => p(s(bool,X50)))))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.IMP_CONG')).
fof(f360,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.closSem.v_|49||49||split|3','_dst_x0x1_ax')).
fof(f359,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X25,X26,X27] : ! [X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X30),s(fun(X26,X27),X29))) = s(fun(X26,X25),happ(s(fun(fun(X27,X25),fun(X26,X25)),'_dstsel_0_x01_1'(s(fun(X26,X27),X29))),s(fun(X27,X25),X30)))),
  file('thm.closSem.v_|49||49||split|3','_dstsel_0_x01_ax')).
fof(f351,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X31,X9] : ! [X40,X11] : ? [X41] : (s(X31,X40) = s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X31,happ(s(fun(X31,X31),happ(s(fun('type.list.list'(X9),fun(X31,X31)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X31,X31))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),X17))))))),
  file('thm.closSem.v_|49||49||split|3','thm.list.list_Axiom')).
fof(f352,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f346,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X31,X9] : ! [X10,X54] : (s(fun(X9,X31),X10) = s(fun(X9,X31),X54) <=> ! [X32] : s(X31,happ(s(fun(X9,X31),X10),s(X9,X32))) = s(X31,happ(s(fun(X9,X31),X54),s(X9,X32))))),
  file('thm.closSem.v_|49||49||split|3','thm.bool.FUN_EQ_THM')).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f239])).
fof(f345,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.closSem.v_|49||49||split|3',aHL_EXT)).
fof(f336,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X31] : ! [X32,X39] : s(X31,X39) = s(X31,'const.pair.SND_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X32),s(X31,X39)))))),
  file('thm.closSem.v_|49||49||split|3','thm.pair.SND')).
fof(f335,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31,X9] : ! [X32,X39] : s(X9,X32) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X32),s(X31,X39)))))),
  file('thm.closSem.v_|49||49||split|3','thm.pair.FST')).
fof(f334,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X31,X9] : ! [X42] : ? [X10] : ! [X43,X44,X45] : s(X31,happ(s(fun('type.ind_type.recspace'(X9),X31),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X43),s(X9,X44),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45))))) = s(X31,happ(s(fun(fun('type.num.num',X31),X31),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)))),X42),s('type.num.num',X43))),s(X9,X44))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45))),s(fun('type.num.num',X31),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X31),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45)))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.CONSTR_REC')).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X31] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X31),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X31),X8)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_1_const.pair.FST')).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9,X31] : ! [X32] : s(X31,X32) = s(X31,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X31),'const.sum.INR_1'(s(X31,X32)))))),
  file('thm.closSem.v_|49||49||split|3','thm.sum.OUTR')).
fof(f329,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) & (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_|49||49||split|1')).
fof(f330,plain,(
  ( ! [X0,X1] : (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))) )),
  inference(cnf_transformation,[],[f231])).
fof(f327,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X6))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X7))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_|49||49||split|0')).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f230])).
fof(f323,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f322,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),'const.list.CONS_2'(s('type.closSem.v',X0),s('type.list.list'('type.closSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v',X0))),s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),'const.list.CONS_2'(s('type.closSem.v',X0),s('type.list.list'('type.closSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v',X0))),s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X1)))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X16,X17] : s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),'const.list.CONS_2'(s('type.closSem.v',X16),s('type.list.list'('type.closSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v',X16))),s('type.num.num','const.closSem.v|49|_size_1'(s('type.list.list'('type.closSem.v'),X17)))))))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_size_def|split|7')).
fof(f319,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : ! [X8,X20,X21] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X20),s('type.num.num',X21))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X20))),s('type.num.num',X21)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_3_const.ind_type.FCONS')).
fof(f318,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X20,X21] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X20),s(fun('type.num.num','type.ind_type.recspace'(X9)),X21))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X20))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X21)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_1_3_const.ind_type.CONSTR')).
fof(f317,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X6,X10,X38] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X38))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38)))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.FCONS|split|1')).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.FCONS|split|0')).
fof(f262,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.closSem.v'),'const.closSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.closSem.v'),'const.closSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8))) = s('type.closSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.closSem.v'),'const.closSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_1_const.closSem.Word64')).
fof(f261,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6)))))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_size_def|split|1')).
fof(f260,plain,(
  ( ! [X0] : (s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_size_def|split|3')).
fof(f259,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.closSem.v_size_1'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X6)))))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_size_def|split|0')).
fof(f258,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) = s('type.closSem.v',happ(s(fun('type.integer.int','type.closSem.v'),'const.closSem.Number_0'),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))) = s('type.closSem.v',happ(s(fun('type.integer.int','type.closSem.v'),'const.closSem.Number_0'),s('type.integer.int',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X8))) = s('type.closSem.v',happ(s(fun('type.integer.int','type.closSem.v'),'const.closSem.Number_0'),s('type.integer.int',X8)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_1_const.closSem.Number')).
fof(f257,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),'const.closSem.RefPtr_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),'const.closSem.RefPtr_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X8))) = s('type.closSem.v',happ(s(fun('type.num.num','type.closSem.v'),'const.closSem.RefPtr_0'),s('type.num.num',X8)))),
  file('thm.closSem.v_|49||49||split|3','ARITY_0_1_const.closSem.RefPtr')).
fof(f256,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.closSem.v_|49||49||split|3','thm.ind_type.ISO_REFL')).
fof(f7065,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl20)),
  inference(superposition,[],[f6370,f6296])).
fof(f6296,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f6258,f255])).
fof(f6258,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl20)),
  inference(superposition,[],[f5877,f3403])).
fof(f6370,plain,(
  ( ! [X23,X22] : (s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl20),
  inference(subsumption_resolution,[],[f6325,f253])).
fof(f6325,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl20),
  inference(superposition,[],[f433,f6248])).
fof(f6248,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f418,f5877,f2101])).
fof(f10100,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl28),
  inference(cnf_transformation,[],[f10100_D])).
fof(f10100_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f17439,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl20)),
  inference(duplicate_literal_removal,[],[f17276])).
fof(f17276,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl20)),
  inference(superposition,[],[f7065,f11111])).
fof(f11111,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f11101])).
fof(f11101,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f10165,f431])).
fof(f10165,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2170,f2198,f2179,f2279,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2454,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f3506,f3403,f3541,f3555,f4002,f4003,f4005,f4006,f4008,f4009,f4010,f4044,f4072,f4074,f4075,f4077,f4078,f4079,f4182,f4250,f4253,f4690,f4701,f4343,f4734,f4766,f4767,f4769,f4770,f4772,f4773,f4774,f4800,f4801,f4802,f4825,f4808,f4809,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5418,f5256,f5257,f5310,f5312,f5313,f5420,f5328,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5710,f5787,f5870,f5665,f5534,f5576,f5581,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5845,f5842,f5638,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f4698,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5558,f5555,f5419,f3637,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f4828,f4820,f4821,f5874,f5871,f6191,f5535,f5156,f4983,f4768,f4073,f4004,f3508,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f4771,f4076,f4007,f3511,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6673,f6680,f6681,f6682,f6683,f6685,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7532,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8710,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f8938,f8941,f8945,f9041,f9049,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9560,f9172,f9173,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9324,f9325,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9422,f9423,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9665,f9670,f9672,f9775,f10152,f10154,f10156,f10157,f10159,f10160,f10161,f10162,f10164])).
fof(f10164,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f10163])).
fof(f10163,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f7471,f4253])).
fof(f7471,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f7245,f3403])).
fof(f17094,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) | $spl42),
  inference(cnf_transformation,[],[f17094_D])).
fof(f17094_D,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f17019,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl20)),
  inference(duplicate_literal_removal,[],[f16860])).
fof(f16860,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl20)),
  inference(superposition,[],[f7065,f10932])).
fof(f10932,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10922])).
fof(f10922,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f10160,f431])).
fof(f13477,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl38),
  inference(cnf_transformation,[],[f13477_D])).
fof(f13477_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f13479,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl40),
  inference(cnf_transformation,[],[f13479_D])).
fof(f13479_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f13462,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl36),
  inference(superposition,[],[f12222,f302])).
fof(f2174,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl12),
  inference(cnf_transformation,[],[f2174_D])).
fof(f2174_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12222,plain,(
  ( ! [X53] : (p(s(bool,happ(s(fun(bool,bool),X53),s(bool,sK6('_i_0_0',X53,bool))))) | p(s(bool,sK5('_i_0_0',X53,bool)))) ) | $spl36),
  inference(cnf_transformation,[],[f12222_D])).
fof(f12222_D,plain,(
  ( ! [X53] : (p(s(bool,happ(s(fun(bool,bool),X53),s(bool,sK6('_i_0_0',X53,bool))))) | p(s(bool,sK5('_i_0_0',X53,bool)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f12221,plain,(
  ( ! [X52,X53] : (p(s(bool,sK5('_i_0_0',X53,bool))) | ~p(s(bool,X52)) | p(s(bool,happ(s(fun(bool,bool),X53),s(bool,sK6('_i_0_0',X53,bool)))))) )),
  inference(forward_demodulation,[],[f12193,f302])).
fof(f12193,plain,(
  ( ! [X52,X53] : (~p(s(bool,X52)) | p(s(bool,happ(s(fun(bool,bool),X53),s(bool,sK6('_i_0_0',X53,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X53,bool)))))) )),
  inference(superposition,[],[f313,f302])).
fof(f10572,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561,f10563,f10565,f10567,f10569,f10571])).
fof(f10571,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10437,f253])).
fof(f10437,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X406,X400,X402] : (~p(s(bool,'T_0')) | s(X399,X406) = s(X399,happ(s(fun(X398,X399),happ(s(fun(X400,fun(X398,X399)),'_dst_x0x1_2'(s(fun(X401,fun(X398,X399)),X402),s(fun(X400,X401),X403))),s(X400,X404))),s(X398,happ(s(fun(X399,X398),X405),s(X399,X406)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9438])).
fof(f10569,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10436,f253])).
fof(f10436,plain,(
  ( ! [X420,X422,X416,X418,X424,X421,X423,X417,X419] : (~p(s(bool,'T_0')) | s(X416,X424) = s(X416,happ(s(fun(X417,X416),X423),s(X417,happ(s(fun(X416,X417),happ(s(fun(X418,fun(X416,X417)),'_dst_x0x1_2'(s(fun(X419,fun(X416,X417)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s(X416,X424)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9440])).
fof(f10567,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10435,f253])).
fof(f10435,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (~p(s(bool,'T_0')) | s(X420,X427) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X421,fun(X419,X420)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425))))),s(X419,happ(s(fun(X420,X419),X426),s(X420,X427)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9341])).
fof(f10565,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10434,f253])).
fof(f10434,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (~p(s(bool,'T_0')) | s(X437,X445) = s(X437,happ(s(fun(X438,X437),X444),s(X438,happ(s(fun(X437,X438),happ(s(fun(X439,fun(X437,X438)),X440),s(X439,happ(s(fun(X441,X439),X442),s(X441,X443))))),s(X437,X445)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9343])).
fof(f10563,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10433,f253])).
fof(f10433,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X415,X409,X411] : (~p(s(bool,'T_0')) | s(X407,X415) = s(X407,happ(s(fun(X408,X407),X414),s(X408,happ(s(fun(X407,X408),happ(s(fun(X409,fun(X407,X408)),'_dst_x0x1_2'(s(fun(X410,fun(X407,X408)),X411),s(fun(X409,X410),X412))),s(X409,X413))),s(X407,X415)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9439])).
fof(f10561,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10432,f253])).
fof(f10432,plain,(
  ( ! [X428,X430,X426,X432,X429,X431,X425,X427,X433] : (~p(s(bool,'T_0')) | s(X426,X433) = s(X426,happ(s(fun(X425,X426),happ(s(fun(X427,fun(X425,X426)),'_dst_x0x1_2'(s(fun(X428,fun(X425,X426)),X429),s(fun(X427,X428),X430))),s(X427,X431))),s(X425,happ(s(fun(X426,X425),X432),s(X426,X433)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9441])).
fof(f10559,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10431,f253])).
fof(f10431,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (~p(s(bool,'T_0')) | s(X428,X436) = s(X428,happ(s(fun(X429,X428),X435),s(X429,happ(s(fun(X428,X429),happ(s(fun(X430,fun(X428,X429)),X431),s(X430,happ(s(fun(X432,X430),X433),s(X432,X434))))),s(X428,X436)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9342])).
fof(f10557,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10430,f253])).
fof(f10430,plain,(
  ( ! [X447,X453,X449,X451,X446,X452,X454,X448,X450] : (~p(s(bool,'T_0')) | s(X447,X454) = s(X447,happ(s(fun(X446,X447),happ(s(fun(X448,fun(X446,X447)),X449),s(X448,happ(s(fun(X450,X448),X451),s(X450,X452))))),s(X446,happ(s(fun(X447,X446),X453),s(X447,X454)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9344])).
fof(f10553,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10421,f253])).
fof(f10421,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f349])).
fof(f10551,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10420,f253])).
fof(f10420,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f348])).
fof(f10549,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10419,f253])).
fof(f10419,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f344])).
fof(f10547,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10418,f253])).
fof(f10418,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f343])).
fof(f10478,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl26),
  inference(subsumption_resolution,[],[f10188,f10099])).
fof(f10188,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f268])).
fof(f2354,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f263,f317])).
fof(f10175,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f6604,f255])).
fof(f6604,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f406,f263])).
fof(f10570,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561,f10563,f10565,f10567,f10569])).
fof(f10568,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561,f10563,f10565,f10567])).
fof(f10566,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561,f10563,f10565])).
fof(f10564,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561,f10563])).
fof(f10562,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559,f10561])).
fof(f10560,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557,f10559])).
fof(f10558,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553,f10557])).
fof(f10554,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551,f10553])).
fof(f10552,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549,f10551])).
fof(f10550,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547,f10549])).
fof(f10548,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478,f10547])).
fof(f10493,plain,(
  $false | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f10353,f10099])).
fof(f10353,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl26)),
  inference(backward_demodulation,[],[f10099,f10165])).
fof(f10492,plain,(
  $false | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f10352,f10099])).
fof(f10352,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl26)),
  inference(backward_demodulation,[],[f10099,f10160])).
fof(f10491,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10490,f253])).
fof(f10490,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(forward_demodulation,[],[f10489,f10099])).
fof(f10489,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10339,f10099])).
fof(f10339,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114))))) | ~p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9391])).
fof(f10488,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10487,f253])).
fof(f10487,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(forward_demodulation,[],[f10486,f10099])).
fof(f10486,plain,(
  ( ! [X125,X123,X121,X124,X122] : (~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10328,f10099])).
fof(f10328,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X123,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X123,X121),X124))),s(X123,X125))) | ~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X123,X121),X124),s(X123,X125))))))) ) | $spl26),
  inference(backward_demodulation,[],[f10099,f9293])).
fof(f10479,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f252,f251,f253,f256,f257,f258,f259,f260,f261,f262,f263,f265,f264,f302,f309,f317,f318,f319,f320,f321,f322,f326,f323,f328,f327,f330,f329,f331,f332,f333,f334,f335,f336,f345,f347,f346,f352,f351,f359,f360,f404,f387,f386,f385,f388,f390,f391,f392,f418,f426,f431,f433,f480,f789,f815,f847,f886,f887,f888,f889,f890,f891,f969,f1001,f1115,f1219,f1477,f1478,f1480,f1481,f1482,f1483,f1484,f1506,f1519,f1520,f1522,f1523,f1524,f1525,f1526,f1957,f1958,f1959,f1969,f1970,f1972,f1973,f1975,f1976,f1977,f1978,f1979,f1980,f1981,f1982,f2167,f2173,f2376,f2382,f2446,f2518,f2521,f2553,f2554,f2555,f2556,f2557,f2558,f2559,f2560,f2561,f2737,f2739,f1097,f1131,f2412,f1008,f1007,f1006,f1005,f1004,f1003,f1002,f312,f1961,f1692,f1590,f2453,f2452,f2104,f2103,f2101,f2099,f2098,f2097,f2096,f2095,f2094,f999,f944,f729,f672,f619,f526,f499,f458,f269,f268,f255,f4800,f4801,f4802,f4982,f4984,f4985,f4987,f4988,f4989,f5068,f5155,f5157,f5158,f5160,f5161,f5162,f5242,f5248,f5249,f5250,f5310,f5312,f5313,f5332,f5340,f5346,f5353,f5424,f5562,f5664,f5658,f5787,f5870,f5665,f5534,f5576,f5262,f5265,f5269,f5268,f5267,f5266,f5264,f5263,f5261,f5260,f5842,f5635,f350,f349,f348,f344,f343,f389,f2152,f2730,f2733,f2736,f2735,f2734,f2732,f2731,f2729,f2728,f406,f304,f342,f339,f5555,f5419,f2171,f341,f340,f314,f313,f310,f4812,f4813,f4814,f5871,f6191,f5535,f5156,f4983,f2175,f2049,f1521,f1479,f1088,f5159,f4986,f2368,f2341,f2339,f5348,f5335,f5591,f5601,f5798,f5808,f5334,f5331,f5471,f5317,f5316,f5311,f5309,f5304,f5303,f5301,f5300,f5299,f5298,f5297,f5333,f845,f844,f843,f842,f2375,f841,f840,f965,f1111,f5423,f5422,f311,f6670,f6680,f6681,f6682,f6686,f6690,f6693,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f7154,f7155,f7157,f7158,f7159,f7160,f7162,f7245,f7253,f7352,f7353,f7355,f7356,f7357,f7358,f7360,f7444,f7452,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f8040,f8041,f8066,f8067,f8092,f8093,f8017,f8028,f8672,f8673,f8674,f8675,f8676,f8677,f8678,f8679,f8680,f8744,f8745,f8766,f9155,f9479,f9480,f9481,f9143,f9142,f9141,f9128,f9127,f9501,f9502,f9503,f9106,f9158,f9159,f9160,f9161,f9162,f9163,f9164,f9165,f9166,f9177,f9288,f9289,f9291,f9292,f9293,f9294,f9295,f9297,f9298,f9303,f9304,f9305,f9306,f9307,f9310,f9311,f9312,f9313,f9326,f9329,f9330,f9334,f9335,f9336,f9337,f9338,f9339,f9341,f9342,f9343,f9344,f9351,f9357,f9359,f9367,f9369,f9385,f9386,f9387,f9389,f9390,f9391,f9392,f9393,f9395,f9396,f9401,f9402,f9403,f9404,f9405,f9408,f9409,f9410,f9411,f9424,f9427,f9428,f9431,f9432,f9433,f9434,f9435,f9436,f9438,f9439,f9440,f9441,f9448,f9454,f9456,f9464,f9466,f9775,f10156,f10157,f10161,f10162,f6626,f6625,f10175,f6567,f6587,f5533,f5553,f10176,f3301,f2354,f10186,f10099,f10478])).
fof(f10477,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10187,f253])).
fof(f10187,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(backward_demodulation,[],[f10099,f255])).
fof(f10177,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl34),
  inference(cnf_transformation,[],[f10177_D])).
fof(f10177_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1010,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1010_D])).
fof(f1010_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f10186,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2060,f253])).
fof(f2060,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f312,f458])).
fof(f1011,plain,(
  ( ! [X3] : (~p(s(bool,sK4(X3,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1011_D])).
fof(f1011_D,plain,(
  ( ! [X3] : (~p(s(bool,sK4(X3,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3301,plain,(
  ( ! [X23] : (~p(s(bool,sK4(X23,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f312,f302])).
fof(f10176,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f4876,f253])).
fof(f4876,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f312,f458])).
fof(f2168,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2168_D])).
fof(f2168_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5553,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f5520,f253])).
fof(f5520,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f5419,f458])).
fof(f5533,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f5419,f302])).
fof(f6587,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6554,f253])).
fof(f6554,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f5419,f458])).
fof(f6567,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f5419,f302])).
fof(f2738,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),sK2(X7,X6))))) ) | $spl18),
  inference(cnf_transformation,[],[f2738_D])).
fof(f2738_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),sK2(X7,X6))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2520,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl16),
  inference(cnf_transformation,[],[f2520_D])).
fof(f2520_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6625,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),sK2(X6,X5))))) )),
  inference(subsumption_resolution,[],[f6606,f264])).
fof(f6606,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),sK2(X6,X5))))) )),
  inference(superposition,[],[f406,f431])).
fof(f6626,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8))))) )),
  inference(subsumption_resolution,[],[f6607,f264])).
fof(f6607,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f406,f431])).
fof(f10155,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl32),
  inference(cnf_transformation,[],[f10155_D])).
fof(f10155_D,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f10162,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f7479,f7245])).
fof(f7479,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f7245,f2095])).
fof(f10161,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f7490,f7444])).
fof(f7490,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f7245,f2099])).
fof(f10153,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl30),
  inference(cnf_transformation,[],[f10153_D])).
fof(f10153_D,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f10157,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f8656,f7245])).
fof(f8656,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f7444,f2095])).
fof(f10156,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f8667,f7444])).
fof(f8667,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f7444,f2099])).
fof(f10154,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(subsumption_resolution,[],[f8996,f9041])).
fof(f8996,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f7245,f3506])).
fof(f10152,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(subsumption_resolution,[],[f8997,f9041])).
fof(f8997,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f7444,f3506])).
fof(f10033,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9896])).
fof(f9896,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2523,f8592])).
fof(f8592,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f8578])).
fof(f8578,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f7530,f431])).
fof(f7530,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f7458,f1025])).
fof(f1025,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1011,f269])).
fof(f7458,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1026,f7245])).
fof(f1026,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1011,f264])).
fof(f2523,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f2522,f255])).
fof(f2522,plain,(
  ( ! [X14,X12] : (p(s(bool,'F_0')) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(forward_demodulation,[],[f2488,f1025])).
fof(f2488,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,sK4(X13,'_i_0_0',bool))) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f2103,f1076])).
fof(f1076,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,sK4(X15,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(subsumption_resolution,[],[f1057,f255])).
fof(f1057,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,sK4(X15,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f480,f1025])).
fof(f6688,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl24),
  inference(cnf_transformation,[],[f6688_D])).
fof(f6688_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6687,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f6653,f268])).
fof(f6653,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f5576,f2455])).
fof(f2455,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2432,f253])).
fof(f2432,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f2101,f1192])).
fof(f1192,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1148,f265])).
fof(f1148,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1026,f311])).
fof(f6684,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl22),
  inference(cnf_transformation,[],[f6684_D])).
fof(f6684_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6683,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6650,f268])).
fof(f6650,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f5576,f2454])).
fof(f6190,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5889,f2174])).
fof(f5889,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f418,f2174,f323])).
fof(f6189,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5891,f2174])).
fof(f5891,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1026,f2174,f323])).
fof(f6188,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5899,f2174])).
fof(f5899,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f418,f2174,f323])).
fof(f6187,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5901,f2174])).
fof(f5901,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1026,f2174,f323])).
fof(f6186,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6177,f2174])).
fof(f6177,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(backward_demodulation,[],[f5962,f5665])).
fof(f5962,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f332,f2174,f2103])).
fof(f6088,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6087,f418])).
fof(f6087,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5965,f426])).
fof(f5965,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f426,f2174,f2103])).
fof(f6086,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6085,f418])).
fof(f6085,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f5966,f1192])).
fof(f5966,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1192,f2174,f2103])).
fof(f6084,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f6083,f418])).
fof(f6083,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5967,f5664])).
fof(f5967,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f5664,f2174,f2103])).
fof(f6082,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f6081,f418])).
fof(f6081,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f5968,f2198])).
fof(f5968,plain,(
  s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f2198,f2174,f2103])).
fof(f5986,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f2174,f5534])).
fof(f5984,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1026,f2174,f311])).
fof(f5980,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f426,f418,f2174,f2103])).
fof(f5950,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1026,f1026,f2174,f480])).
fof(f5948,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f418,f418,f2174,f480])).
fof(f5926,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f332,f2174,f404])).
fof(f5924,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f336,f2174,f404])).
fof(f5922,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f335,f2174,f404])).
fof(f5920,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f309,f2174,f404])).
fof(f5918,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f331,f2174,f404])).
fof(f5916,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f351,f2174,f404])).
fof(f5914,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f302,f2174,f404])).
fof(f5911,plain,(
  $false | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f253,f2198,f2174,f404])).
fof(f5909,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f5664,f2174,f404])).
fof(f5907,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f253,f1192,f2174,f404])).
fof(f5905,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f426,f2174,f404])).
fof(f5902,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2174,f1026,f2174,f323])).
fof(f5900,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f255,f418,f2174,f323])).
fof(f5898,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2174,f1026,f2174,f323])).
fof(f5896,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2174,f418,f2174,f323])).
fof(f5892,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2174,f1026,f2174,f323])).
fof(f5890,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f255,f418,f2174,f323])).
fof(f5888,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2174,f1026,f2174,f323])).
fof(f5886,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2174,f418,f2174,f323])).
fof(f5882,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f418,f2174,f269])).
fof(f5880,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f426,f2174,f264])).
fof(f5878,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f253,f2174])).
fof(f5877,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl20),
  inference(cnf_transformation,[],[f5877_D])).
fof(f5877_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5870,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f5535,f335])).
fof(f2990,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2749,f253])).
fof(f2749,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f326])).
fof(f2989,plain,(
  $false | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f2751,f253])).
fof(f2751,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1026,f2520,f326])).
fof(f2988,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2753,f2520])).
fof(f2753,plain,(
  ~p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f326])).
fof(f2987,plain,(
  $false | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f2755,f2520])).
fof(f2755,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1026,f2520,f326])).
fof(f2986,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2759,f253])).
fof(f2759,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f326])).
fof(f2985,plain,(
  $false | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f2761,f253])).
fof(f2761,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1026,f2520,f326])).
fof(f2984,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2763,f2520])).
fof(f2763,plain,(
  ~p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f326])).
fof(f2983,plain,(
  $false | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f2765,f2520])).
fof(f2765,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1026,f2520,f326])).
fof(f2976,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2975,f2520])).
fof(f2975,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f2870,f2836])).
fof(f2836,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f2101])).
fof(f2870,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f2836,f268])).
fof(f2974,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2869,f253])).
fof(f2869,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f2836,f255])).
fof(f2837,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f418,f418,f2520,f2101])).
fof(f2831,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1025,f1026,f2520,f433])).
fof(f2829,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1025,f418,f2520,f433])).
fof(f2799,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f332,f2520,f404])).
fof(f2797,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f336,f2520,f404])).
fof(f2795,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f335,f2520,f404])).
fof(f2793,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f309,f2520,f404])).
fof(f2791,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f331,f2520,f404])).
fof(f2789,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f351,f2520,f404])).
fof(f2787,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f302,f2520,f404])).
fof(f2785,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f255,f1025,f2520,f404])).
fof(f2766,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f2520,f1026,f2520,f326])).
fof(f2764,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f2520,f418,f2520,f326])).
fof(f2762,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f253,f1026,f2520,f326])).
fof(f2760,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f253,f418,f2520,f326])).
fof(f2756,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f2520,f1026,f2520,f326])).
fof(f2754,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f2520,f418,f2520,f326])).
fof(f2752,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f253,f1026,f2520,f326])).
fof(f2750,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f253,f418,f2520,f326])).
fof(f2746,plain,(
  $false | ($spl6 | $spl16)),
  inference(unit_resulting_resolution,[],[f1025,f2520,f268])).
fof(f2744,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f418,f2520,f265])).
fof(f2742,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f255,f2520])).
fof(f2739,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2(X10,X9))))) )),
  inference(subsumption_resolution,[],[f2725,f264])).
fof(f2725,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2(X10,X9))))) )),
  inference(superposition,[],[f406,f431])).
fof(f2737,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),sK2(X7,X6))))) )),
  inference(subsumption_resolution,[],[f2724,f264])).
fof(f2724,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),sK2(X7,X6)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f406,f431])).
fof(f2519,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'F_0') = s(bool,X8)) ) | $spl14),
  inference(cnf_transformation,[],[f2519_D])).
fof(f2519_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2521,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f2487,f264])).
fof(f2487,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f2103,f431])).
fof(f2518,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2486,f264])).
fof(f2486,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f2103,f431])).
fof(f2170,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2170_D])).
fof(f2170_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2173,plain,(
  ( ! [X24] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f2151,f302])).
fof(f2151,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f310,f302])).
fof(f2167,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2166,f302])).
fof(f2166,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2144,f253])).
fof(f2144,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f310,f458])).
fof(f1001,plain,(
  ( ! [X3] : (~p(s(bool,sK4(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f312,f302])).
fof(f417,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f416])).
fof(f416,plain,(
  s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f413,f414])).
fof(f412,plain,(
  s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f414,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f251,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f195])).
fof(f195,plain,(
  ? [X0,X1] : ((s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ? [X0,X1] : (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X6))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X6))) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.closSem.v_|49||49||split|3','thm.closSem.v_|49||49||split|3')).
fof(f413,plain,(
  s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f415,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f252,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK0))) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
% SZS output end Proof for thm.closSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 192
% Passive clauses: 2005
% Generated clauses: 22851
% Final active clauses: 127
% Final passive clauses: 1179
% Input formulas: 69
% Initial clauses: 142
% 
% Duplicate literals: 743
% Trivial inequalities: 183
% Fw subsumption resolutions: 1004
% Fw demodulations: 4270
% Bw demodulations: 514
% Condensations: 962
% Global subsumptions: 16
% 
% Simple tautologies: 285
% Equational tautologies: 154
% Forward subsumptions: 12186
% Fw demodulations to eq. taut.: 920
% Bw demodulations to eq. taut.: 65
% 
% Unit resulting resolution: 2516
% Factoring: 1
% Forward superposition: 6445
% Backward superposition: 5131
% Self superposition: 152
% Equality factoring: 78
% Equality resolution: 32
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 2179
% SAT solver unit clauses: 187
% SAT solver binary clauses: 826
% 
% Sat splits: 33
% Sat splitting refutations: 90
% 
% Memory used [KB]: 24434
% Time elapsed: 0.958 s
% ------------------------------
% ------------------------------
% Success in time 2.3 s
