Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.typeSound
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSound
% SZS output start Proof for thm.typeSound
fof(f3415,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3397,f56,f88,f310,f315,f313,f82,f86,f382])).
fof(f382,plain,(
  $false | ($spl2 | $spl5 | $spl12)),
  inference(subsumption_resolution,[],[f331,f86])).
fof(f331,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl12)),
  inference(backward_demodulation,[],[f313,f88])).
fof(f86,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f82,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X10] : (s(X6,X10) = s(X6,X10) <=> p(s(bool,'T_0')))),
  file('thm.typeSound.typeSound_45_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f313,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl12),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f315,plain,(
  s(bool,'F_0') = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) | $spl14),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  s(bool,'F_0') = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f310,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f274])).
fof(f274,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl2),
  inference(superposition,[],[f255,f57])).
fof(f57,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.typeSound.typeSound_45_prove_ltb',aHL_BOOL_CASES)).
fof(f255,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | $spl2),
  inference(subsumption_resolution,[],[f182,f86])).
fof(f182,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f53,f90])).
fof(f90,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f57,f57])).
fof(f53,plain,(
  ~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))))),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(sK0,sK1),sK2),s('type.finite_map.fmap'(sK0,sK1),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f35])).
fof(f35,plain,(
  ? [X0,X1,X2] : ~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X2))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X2))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1] : ! [X2] : p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X6,X7),X8),s('type.finite_map.fmap'(X6,X7),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X6,X7),X8),s('type.finite_map.fmap'(X6,X7),X8))))),
  file('thm.typeSound.typeSound_45_prove_ltb','thm.typeSound.typeSound_45_prove_ltb')).
fof(f88,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f56,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.typeSound.typeSound_45_prove_ltb',aHL_FALSITY)).
fof(f3397,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f3364,f88])).
fof(f3364,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(superposition,[],[f3355,f315])).
fof(f3355,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X2))))) )),
  inference(equality_resolution,[],[f3352])).
fof(f3352,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.finite_map.fmap'(X0,X1),X2) != s('type.finite_map.fmap'(X0,X1),X3) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))) )),
  inference(subsumption_resolution,[],[f3347,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s(X0,X2)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ! [X2] : s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s(X0,X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7,X6] : ! [X12] : s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X7,X6),'const.finite_map.FEMPTY_0'),s(X7,X12)))),
  file('thm.typeSound.typeSound_45_prove_ltb','thm.finite_map.FLOOKUP_EMPTY')).
fof(f3347,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.finite_map.fmap'(X0,X1),X2) != s('type.finite_map.fmap'(X0,X1),X3) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s(X0,sK5('const.finite_map.FEMPTY_0',X3,X2,X1,X0)))) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))) )),
  inference(superposition,[],[f80,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'(X0,X1),X2) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s('type.finite_map.fmap'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'(X0,X1),X2) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s('type.finite_map.fmap'(X0,X1),X2)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ! [X2] : s('type.finite_map.fmap'(X0,X1),X2) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FEMPTY_0'),s('type.finite_map.fmap'(X0,X1),X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ! [X11] : s('type.finite_map.fmap'(X6,X7),X11) = s('type.finite_map.fmap'(X6,X7),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X6,X7),'const.finite_map.FEMPTY_0'),s('type.finite_map.fmap'(X6,X7),X11)))),
  file('thm.typeSound.typeSound_45_prove_ltb','thm.finite_map.FUNION_FEMPTY_|49|')).
fof(f80,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.finite_map.fmap'(X0,X1),X3) != s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) | s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,sK5(X4,X3,X2,X1,X0)))) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3)))) | (s('type.finite_map.fmap'(X0,X1),X3) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),sK4(X3,X2,X1,X0)),s('type.finite_map.fmap'(X0,X1),X2))) & ! [X7] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),sK4(X3,X2,X1,X0)),s(X0,X7))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X7)))))) & (! [X4] : (s('type.finite_map.fmap'(X0,X1),X3) != s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) | (s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,sK5(X4,X3,X2,X1,X0)))) & s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))))) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3)))) | ? [X6] : (s('type.finite_map.fmap'(X0,X1),X3) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X6),s('type.finite_map.fmap'(X0,X1),X2))) & ! [X7] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X6),s(X0,X7))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X7)))))) & (! [X4] : (s('type.finite_map.fmap'(X0,X1),X3) != s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) | ? [X5] : (s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,X5))) & s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X5))))) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))))),
  inference(rectify,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3)))) | ? [X4] : (s('type.finite_map.fmap'(X0,X1),X3) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) & ! [X5] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,X5))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X5)))))) & (! [X4] : (s('type.finite_map.fmap'(X0,X1),X3) != s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) | ? [X5] : (s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,X5))) & s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X5))))) | p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3)))) <=> ? [X4] : (s('type.finite_map.fmap'(X0,X1),X3) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) & ! [X5] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,X5))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X5))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.finite_map.fmap'(X0,X1),X3)))) <=> ? [X4] : (s('type.finite_map.fmap'(X0,X1),X3) = s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.finite_map.fmap'(X0,X1),X2))) & ! [X5] : (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X4),s(X0,X5))) | s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X0,X1),X2),s(X0,X5))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : ! [X13,X14] : (p(s(bool,'const.typeSound.store_type_extension_2'(s('type.finite_map.fmap'(X6,X7),X13),s('type.finite_map.fmap'(X6,X7),X14)))) <=> ? [X15] : (s('type.finite_map.fmap'(X6,X7),X14) = s('type.finite_map.fmap'(X6,X7),'const.finite_map.FUNION_2'(s('type.finite_map.fmap'(X6,X7),X15),s('type.finite_map.fmap'(X6,X7),X13))) & ! [X16] : (s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X6,X7),X15),s(X6,X16))) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),'const.finite_map.FLOOKUP_2'(s('type.finite_map.fmap'(X6,X7),X13),s(X6,X16))))))),
  file('thm.typeSound.typeSound_45_prove_ltb','thm.typeSound.store_type_extension_def')).
% SZS output end Proof for thm.typeSound
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 34
% Passive clauses: 576
% Generated clauses: 3224
% Final active clauses: 27
% Final passive clauses: 515
% Input formulas: 15
% Initial clauses: 29
% 
% Duplicate literals: 377
% Trivial inequalities: 41
% Fw subsumption resolutions: 175
% Fw demodulations: 12
% Bw demodulations: 4
% 
% Simple tautologies: 115
% Equational tautologies: 83
% Forward subsumptions: 1373
% 
% Binary resolution: 23
% Forward superposition: 978
% Backward superposition: 1101
% Self superposition: 240
% Equality factoring: 212
% Equality resolution: 3
% 
% Split clauses: 165
% Split components: 330
% SAT solver clauses: 484
% SAT solver unit clauses: 210
% SAT solver binary clauses: 115
% SAT solver learnt clauses: 1
% 
% Sat splits: 165
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 474
% TWLsolver calls for satisfiability: 19
% 
% Memory used [KB]: 1791
% Time elapsed: 0.246 s
% ------------------------------
% ------------------------------
% Success in time 0.271 s
