Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ml_translator
% (132296)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 689
% Passive clauses: 7566
% Generated clauses: 17624
% Final active clauses: 522
% Final passive clauses: 6306
% Input formulas: 35
% Initial clauses: 168
% 
% Duplicate literals: 361
% Trivial inequalities: 91
% Fw subsumption resolutions: 1907
% Fw demodulations: 217
% Bw demodulations: 152
% 
% Simple tautologies: 207
% Equational tautologies: 60
% Forward subsumptions: 6660
% 
% Binary resolution: 8932
% Factoring: 1166
% Forward superposition: 2147
% Backward superposition: 2067
% Self superposition: 26
% Equality factoring: 28
% Equality resolution: 25
% 
% Split clauses: 564
% Split components: 1136
% SAT solver clauses: 987
% SAT solver unit clauses: 27
% SAT solver binary clauses: 174
% 
% Sat splits: 564
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 980
% TWLsolver calls for satisfiability: 391
% 
% Memory used [KB]: 7164
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ml_translator
% (132320)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 333
% Passive clauses: 1488
% Generated clauses: 7369
% Final active clauses: 323
% Final passive clauses: 1138
% Input formulas: 35
% Initial clauses: 168
% 
% Duplicate literals: 197
% Trivial inequalities: 31
% Fw subsumption resolutions: 738
% Fw demodulations: 302
% Bw demodulations: 21
% Condensations: 219
% Global subsumptions: 13
% 
% Simple tautologies: 233
% Equational tautologies: 19
% Forward subsumptions: 2432
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2300
% Factoring: 166
% Forward superposition: 709
% Backward superposition: 952
% Self superposition: 28
% Equality factoring: 30
% Equality resolution: 13
% 
% Split clauses: 32
% Split components: 69
% SAT solver clauses: 1888
% SAT solver unit clauses: 176
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 32
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1836
% TWLsolver calls for satisfiability: 1678
% 
% Memory used [KB]: 8059
% Time elapsed: 0.710 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ml_translator
% (132342)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 1243
% Passive clauses: 6299
% Generated clauses: 94368
% Final active clauses: 515
% Final passive clauses: 547
% Input formulas: 35
% Initial clauses: 168
% 
% Duplicate literals: 3815
% Trivial inequalities: 248
% Fw subsumption resolutions: 4741
% Fw demodulations: 9189
% Bw demodulations: 2561
% Condensations: 8768
% Global subsumptions: 3972
% 
% Simple tautologies: 267
% Equational tautologies: 1614
% Forward subsumptions: 49404
% Bw demodulations to eq. taut.: 378
% 
% Unit resulting resolution: 21903
% Factoring: 1
% Forward superposition: 17175
% Backward superposition: 9074
% Self superposition: 282
% Equality factoring: 220
% Equality resolution: 8
% 
% Split clauses: 325
% Split components: 673
% SAT solver clauses: 16803
% SAT solver unit clauses: 1890
% SAT solver binary clauses: 2684
% 
% Sat splits: 325
% Sat splitting refutations: 5880
% 
% Memory used [KB]: 99017
% Time elapsed: 10.515 s
% ------------------------------
% ------------------------------
% remaining time: 31 next slice time: 31
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ml_translator
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_translator
% SZS output start Proof for thm.ml_translator
fof(f31884,plain,(
  $false),
  inference(subsumption_resolution,[],[f31883,f377])).
fof(f377,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK12),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13)))))))))),
  inference(resolution,[],[f204,f207])).
fof(f207,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK16(X2,X1,X0))))))))) | ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK16(X2,X1,X0))))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',sK16(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | ? [X4] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X4)))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(rectify,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | ? [X3] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3)))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6,X9] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.ast.exp',X6),s(fun('type.semanticPrimitives.v',bool),X9)))) <=> ? [X20] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X6),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X20)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X9),s('type.semanticPrimitives.v',X20))))))),
  file('thm.ml_translator.Eval_WEAKEN','thm.ml_translator.Eval_def')).
fof(f204,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.ast.exp',sK12),s(fun('type.semanticPrimitives.v',bool),sK15))))),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.ast.exp',sK12),s(fun('type.semanticPrimitives.v',bool),sK15)))) & ! [X4] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK15),s('type.semanticPrimitives.v',X4)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK14),s('type.semanticPrimitives.v',X4))))) & ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.ast.exp',sK12),s(fun('type.semanticPrimitives.v',bool),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f87])).
fof(f87,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X3)))) & ! [X4] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X3),s('type.semanticPrimitives.v',X4)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X4))))) & ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X2)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X3)))) & (! [X4] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X3),s('type.semanticPrimitives.v',X4)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X4))))) & ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X3)))) => (! [X4] : (p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X3),s('type.semanticPrimitives.v',X4)))) => p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X4))))) => p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X1),s('type.ast.exp',X0),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.ast.exp',X6),s(fun('type.semanticPrimitives.v',bool),X9)))) => (! [X10] : (p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X9),s('type.semanticPrimitives.v',X10)))) => p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X8),s('type.semanticPrimitives.v',X10))))) => p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.ast.exp',X6),s(fun('type.semanticPrimitives.v',bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.ast.exp',X6),s(fun('type.semanticPrimitives.v',bool),X9)))) => (! [X10] : (p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X9),s('type.semanticPrimitives.v',X10)))) => p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X8),s('type.semanticPrimitives.v',X10))))) => p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X7),s('type.ast.exp',X6),s(fun('type.semanticPrimitives.v',bool),X8))))))),
  file('thm.ml_translator.Eval_WEAKEN','thm.ml_translator.Eval_WEAKEN')).
fof(f31883,plain,(
  ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK12),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13)))))))))),
  inference(subsumption_resolution,[],[f31857,f343])).
fof(f343,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.ml_translator.Eval_WEAKEN','thm.bool.IMP_CLAUSES|split|2')).
fof(f31857,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK12),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13)))))))))),
  inference(superposition,[],[f499,f1557])).
fof(f1557,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK14),s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13))))),
  inference(resolution,[],[f1177,f353])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.ml_translator.Eval_WEAKEN','thm.bool.EQ_CLAUSES|split|1')).
fof(f1177,plain,(
  p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK14),s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13)))))),
  inference(resolution,[],[f205,f378])).
fof(f378,plain,(
  p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK15),s('type.semanticPrimitives.v',sK16(sK15,sK12,sK13)))))),
  inference(resolution,[],[f204,f208])).
fof(f208,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',sK16(X2,X1,X0))))) | ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f205,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK15),s('type.semanticPrimitives.v',X4)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK14),s('type.semanticPrimitives.v',X4))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f499,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),sK14),s('type.semanticPrimitives.v',X0)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK12),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) )),
  inference(resolution,[],[f206,f209])).
fof(f209,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3))))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f206,plain,(
  ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK13),s('type.ast.exp',sK12),s(fun('type.semanticPrimitives.v',bool),sK14))))),
  inference(cnf_transformation,[],[f131])).
% SZS output end Proof for thm.ml_translator
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 851
% Passive clauses: 5464
% Generated clauses: 31456
% Final active clauses: 831
% Final passive clauses: 4213
% Input formulas: 35
% Initial clauses: 168
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 2252
% Trivial inequalities: 209
% Fw subsumption resolutions: 4621
% Fw demodulations: 2488
% Bw demodulations: 7
% 
% Simple tautologies: 3328
% Equational tautologies: 210
% Forward subsumptions: 12488
% 
% Binary resolution: 17407
% Factoring: 62
% Forward superposition: 2841
% Backward superposition: 1473
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 30
% 
% Split clauses: 817
% Split components: 1888
% SAT solver clauses: 953
% SAT solver unit clauses: 43
% SAT solver binary clauses: 216
% 
% Sat splits: 817
% Sat splitting refutations: 46
% 
% Memory used [KB]: 8059
% Time elapsed: 0.801 s
% ------------------------------
% ------------------------------
% Success in time 12.679 s
