Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.LazyPairingHeap
% (120886)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 607
% Passive clauses: 4795
% Generated clauses: 14589
% Final active clauses: 597
% Final passive clauses: 3915
% Input formulas: 133
% Initial clauses: 363
% 
% Duplicate literals: 183
% Trivial inequalities: 10
% Fw subsumption resolutions: 718
% Fw demodulations: 1154
% Bw demodulations: 9
% 
% Simple tautologies: 252
% Equational tautologies: 45
% Forward subsumptions: 7292
% Fw demodulations to eq. taut.: 20
% 
% Binary resolution: 5065
% Factoring: 931
% Forward superposition: 2811
% Backward superposition: 3216
% Self superposition: 48
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 175
% Split components: 352
% SAT solver clauses: 344
% SAT solver unit clauses: 11
% SAT solver binary clauses: 67
% 
% Sat splits: 175
% 
% TWLsolver clauses: 344
% TWLsolver calls for satisfiability: 273
% 
% Memory used [KB]: 8699
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.LazyPairingHeap
% (120897)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 401
% Passive clauses: 2162
% Generated clauses: 11682
% Final active clauses: 390
% Final passive clauses: 1685
% Input formulas: 133
% Initial clauses: 363
% 
% Duplicate literals: 100
% Trivial inequalities: 9
% Fw subsumption resolutions: 524
% Fw demodulations: 1096
% Bw demodulations: 84
% Condensations: 127
% Global subsumptions: 22
% 
% Simple tautologies: 235
% Equational tautologies: 7
% Forward subsumptions: 4987
% Fw demodulations to eq. taut.: 20
% 
% Binary resolution: 2220
% Factoring: 306
% Forward superposition: 2467
% Backward superposition: 2098
% Self superposition: 60
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 2736
% SAT solver unit clauses: 386
% SAT solver binary clauses: 108
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 6
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2648
% TWLsolver calls for satisfiability: 2763
% 
% Memory used [KB]: 17142
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.LazyPairingHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.LazyPairingHeap
% SZS output start Proof for thm.LazyPairingHeap
fof(f12358,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f567,f961,f959,f566,f960,f958,f9575,f9589,f9587,f12311,f12355,f12357])).
fof(f12357,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f12356,f7266])).
fof(f7266,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X0,X1),X2),s(fun(X1,fun(X1,bool)),X3),s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0'))))) )),
  inference(forward_demodulation,[],[f7254,f673])).
fof(f673,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X7] : ! [X15,X25] : s(X7,X15) = s(X7,'const.ind_type.FCONS_3'(s(X7,X15),s(fun('type.num.num',X7),X25),s('type.num.num','const.num.|48|_0')))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.ind_type.FCONS|split|0')).
fof(f7254,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X0,X1),X2),s(fun(X1,fun(X1,bool)),X3),s('type.LazyPairingHeap.heap'(X0),'const.ind_type.FCONS_3'(s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0'),s(fun('type.num.num','type.LazyPairingHeap.heap'(X0)),X4),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f673,f780])).
fof(f780,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4) | ~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4))))) )),
  inference(cnf_transformation,[],[f516])).
fof(f516,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | (s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,sK55(X4,X3,X2,X1,X0)),s('type.LazyPairingHeap.heap'(X1),sK56(X4,X3,X2,X1,X0)),s('type.LazyPairingHeap.heap'(X1),sK57(X4,X3,X2,X1,X0)))) & s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))) & (! [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) != s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4) | p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK55,sK56,sK57])],[f515])).
fof(f515,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | (? [X8,X9,X10] : s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X8),s('type.LazyPairingHeap.heap'(X1),X9),s('type.LazyPairingHeap.heap'(X1),X10))) & s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))) & (! [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) != s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4) | p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4))))))),
  inference(rectify,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | (? [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) & s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))) & (! [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) != s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4) | p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4))))))),
  inference(flattening,[],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | (? [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) & s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))) & ((! [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) != s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4)) | p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4))))))),
  inference(nnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) <=> (? [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) & s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) <=> (? [X5,X6,X7] : s('type.LazyPairingHeap.heap'(X1),X4) = s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7))) & ~s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X12,X13] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X7,X6),X11),s(fun(X6,fun(X6,bool)),X12),s('type.LazyPairingHeap.heap'(X7),X13)))) <=> (? [X14,X15,X16] : s('type.LazyPairingHeap.heap'(X7),X13) = s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Tree_3'(s(X7,X14),s('type.LazyPairingHeap.heap'(X7),X15),s('type.LazyPairingHeap.heap'(X7),X16))) & ~s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X7),X13)))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.LazyPairingHeap.delete_min_side_def')).
fof(f12356,plain,(
  p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0')))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f958,f961])).
fof(f12355,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f12310,f960])).
fof(f12310,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f12307,f653])).
fof(f653,plain,(
  ( ! [X0,X1] : (s('type.LazyPairingHeap.heap'(X0),X1) = s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Tree_3'(s(X0,sK18(X1,X0)),s('type.LazyPairingHeap.heap'(X0),sK19(X1,X0)),s('type.LazyPairingHeap.heap'(X0),sK20(X1,X0)))) | s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1] : (s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X0),X1) | s('type.LazyPairingHeap.heap'(X0),X1) = s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Tree_3'(s(X0,sK18(X1,X0)),s('type.LazyPairingHeap.heap'(X0),sK19(X1,X0)),s('type.LazyPairingHeap.heap'(X0),sK20(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20])],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X0),X1) | ? [X2,X3,X4] : s('type.LazyPairingHeap.heap'(X0),X1) = s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Tree_3'(s(X0,X2),s('type.LazyPairingHeap.heap'(X0),X3),s('type.LazyPairingHeap.heap'(X0),X4))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : ! [X1] : (s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X0),X1) | ? [X2,X3,X4] : s('type.LazyPairingHeap.heap'(X0),X1) = s('type.LazyPairingHeap.heap'(X0),'const.LazyPairingHeap.Tree_3'(s(X0,X2),s('type.LazyPairingHeap.heap'(X0),X3),s('type.LazyPairingHeap.heap'(X0),X4))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X29] : (s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X7),X29) | ? [X15,X10,X30] : s('type.LazyPairingHeap.heap'(X7),X29) = s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Tree_3'(s(X7,X15),s('type.LazyPairingHeap.heap'(X7),X10),s('type.LazyPairingHeap.heap'(X7),X30))))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.LazyPairingHeap.heap_nchotomy')).
fof(f12307,plain,(
  ( ! [X2,X0,X1] : (s('type.LazyPairingHeap.heap'(sK14),sK17) != s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Tree_3'(s(sK14,X0),s('type.LazyPairingHeap.heap'(sK14),X1),s('type.LazyPairingHeap.heap'(sK14),X2)))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f12306,f960])).
fof(f12306,plain,(
  ( ! [X2,X0,X1] : (s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) | s('type.LazyPairingHeap.heap'(sK14),sK17) != s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Tree_3'(s(sK14,X0),s('type.LazyPairingHeap.heap'(sK14),X1),s('type.LazyPairingHeap.heap'(sK14),X2)))) ) | $spl1),
  inference(subsumption_resolution,[],[f12286,f962])).
fof(f962,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f572,f568])).
fof(f568,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  p(s(bool,'T_0'))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.bool.TRUTH')).
fof(f572,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f417])).
fof(f417,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f39])).
fof(f39,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f12286,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) | s('type.LazyPairingHeap.heap'(sK14),sK17) != s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Tree_3'(s(sK14,X0),s('type.LazyPairingHeap.heap'(sK14),X1),s('type.LazyPairingHeap.heap'(sK14),X2)))) ) | $spl1),
  inference(superposition,[],[f781,f1041])).
fof(f1041,plain,(
  s(bool,'F_0') = s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17))) | $spl1),
  inference(unit_resulting_resolution,[],[f959,f602])).
fof(f602,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X43] : (s(bool,'F_0') = s(bool,X43) <=> ~p(s(bool,X43)))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f781,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4) | s('type.LazyPairingHeap.heap'(X1),X4) != s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Tree_3'(s(X1,X5),s('type.LazyPairingHeap.heap'(X1),X6),s('type.LazyPairingHeap.heap'(X1),X7)))) )),
  inference(cnf_transformation,[],[f516])).
fof(f12311,plain,(
  $false | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f960,f12307,f653])).
fof(f9587,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) | $spl4),
  inference(cnf_transformation,[],[f9587_D])).
fof(f9587_D,plain,(
  p(s(bool,'const.relation.reflexive_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9589,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK27('const.prim_rec.<_0','type.num.num')),s('type.num.num',sK27('const.prim_rec.<_0','type.num.num'))))) | $spl7),
  inference(cnf_transformation,[],[f9589_D])).
fof(f9589_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK27('const.prim_rec.<_0','type.num.num')),s('type.num.num',sK27('const.prim_rec.<_0','type.num.num'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f9575,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK27('const.prim_rec.<_0','type.num.num')),s('type.num.num',sK27('const.prim_rec.<_0','type.num.num'))))) | p(s(bool,'const.relation.reflexive_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(superposition,[],[f669,f715])).
fof(f715,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X66,X67] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X66),s('type.num.num',X67))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s('type.num.num',X66))),s('type.num.num',X67)))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','ARITY_0_2_const.prim_rec.<')).
fof(f669,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK27(X1,X0)))),s(X0,sK27(X1,X0))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK27(X1,X0)))),s(X0,sK27(X1,X0))))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f467])).
fof(f467,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X7] : ! [X71] : (p(s(bool,'const.relation.reflexive_1'(s(fun(X7,fun(X7,bool)),X71)))) <=> ! [X18] : p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X71),s(X7,X18))),s(X7,X18)))))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.relation.reflexive_def')).
fof(f958,plain,(
  p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) | $spl0),
  inference(cnf_transformation,[],[f958_D])).
fof(f958_D,plain,(
  p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f960,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(sK14),sK17) | $spl3),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(sK14),sK17) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f566,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(sK14),sK17) | p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f416])).
fof(f416,plain,(
  (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) | s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(sK14),sK17)) & (~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) | s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f415])).
fof(f415,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4)) & (~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) | s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4)))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) <~> s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) <=> s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') != s('type.LazyPairingHeap.heap'(X1),X4))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LazyPairingHeap.heap'(X1),X4)))) <=> ~s('type.LazyPairingHeap.heap'(X1),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X1),X4))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.LazyPairingHeap.heap'(X7),X10)))) <=> ~s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X7),X10))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.LazyPairingHeap.heap'(X7),X10)))) <=> ~s('type.LazyPairingHeap.heap'(X7),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(X7),X10))),
  file('thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb','thm.LazyPairingHeap.LazyPairingHeap_69_prove_ltb')).
fof(f959,plain,(
  ~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) | $spl1),
  inference(cnf_transformation,[],[f959_D])).
fof(f959_D,plain,(
  ~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f961,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) | $spl2),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f567,plain,(
  s('type.LazyPairingHeap.heap'(sK14),'const.LazyPairingHeap.Empty_0') = s('type.LazyPairingHeap.heap'(sK14),sK17) | ~p(s(bool,'const.LazyPairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LazyPairingHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f416])).
% SZS output end Proof for thm.LazyPairingHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 468
% Passive clauses: 2023
% Generated clauses: 11787
% Final active clauses: 369
% Final passive clauses: 1357
% Input formulas: 133
% Initial clauses: 363
% 
% Duplicate literals: 9
% Trivial inequalities: 6
% Fw subsumption resolutions: 70
% Fw demodulations: 1708
% Bw demodulations: 21
% Condensations: 25
% 
% Simple tautologies: 113
% Equational tautologies: 17
% Forward subsumptions: 7775
% Fw demodulations to eq. taut.: 35
% 
% Unit resulting resolution: 2290
% Factoring: 1
% Forward superposition: 3311
% Backward superposition: 3905
% Self superposition: 67
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 2034
% SAT solver unit clauses: 674
% SAT solver binary clauses: 658
% 
% Sat splits: 3
% Sat splitting refutations: 3
% 
% Memory used [KB]: 14967
% Time elapsed: 0.592 s
% ------------------------------
% ------------------------------
% Success in time 1.92 s
