Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (73204)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1491
% Passive clauses: 10463
% Generated clauses: 31836
% Final active clauses: 1468
% Final passive clauses: 8527
% Input formulas: 133
% Initial clauses: 619
% 
% Duplicate literals: 284
% Trivial inequalities: 51
% Fw subsumption resolutions: 2270
% Fw demodulations: 2223
% Bw demodulations: 20
% 
% Simple tautologies: 347
% Equational tautologies: 51
% Forward subsumptions: 15886
% 
% Binary resolution: 10690
% Factoring: 1143
% Forward superposition: 5371
% Backward superposition: 9383
% Self superposition: 33
% Equality factoring: 2
% Equality resolution: 32
% 
% Split clauses: 417
% Split components: 840
% SAT solver clauses: 709
% SAT solver unit clauses: 10
% SAT solver binary clauses: 124
% 
% Sat splits: 417
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 709
% TWLsolver calls for satisfiability: 637
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (73234)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 511
% Passive clauses: 3044
% Generated clauses: 11134
% Final active clauses: 487
% Final passive clauses: 2520
% Input formulas: 133
% Initial clauses: 619
% 
% Duplicate literals: 121
% Trivial inequalities: 6
% Fw subsumption resolutions: 490
% Fw demodulations: 513
% Bw demodulations: 21
% Condensations: 105
% Global subsumptions: 22
% 
% Simple tautologies: 224
% Equational tautologies: 18
% Forward subsumptions: 3261
% 
% Binary resolution: 2920
% Factoring: 262
% Forward superposition: 2004
% Backward superposition: 969
% Self superposition: 16
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 3959
% SAT solver unit clauses: 636
% SAT solver binary clauses: 234
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 19
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3871
% TWLsolver calls for satisfiability: 9536
% 
% Memory used [KB]: 16758
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f20282,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f772,f1406,f1404,f771,f1405,f1403,f4078,f4088,f4087,f4968,f4973,f4985,f4986,f4987,f4990,f4993,f4996,f4999,f5005,f6135,f6138,f6140,f6142,f6144,f6148,f6150,f6154,f6156,f6166,f6168,f6178,f6180,f6188,f6196,f6216,f6290,f6302,f6308,f6310,f6338,f6354,f6365,f6367,f6373,f6379,f6381,f6391,f6491,f6503,f6509,f6511,f6539,f6555,f6566,f6568,f6574,f6580,f6582,f6592,f6639,f6641,f6643,f6653,f6655,f6657,f6680,f6682,f6693,f6695,f6718,f6720,f6731,f6733,f6756,f6758,f6769,f6771,f6788,f6790,f7456,f7457,f6805,f6807,f6811,f6813,f6818,f6820,f6824,f6826,f6837,f6845,f6855,f6857,f6859,f6870,f6872,f6883,f6885,f6896,f6898,f6909,f6911,f6921,f6932,f6943,f7008,f7010,f7012,f7023,f7025,f7027,f7029,f7031,f7033,f7063,f7065,f7067,f7069,f7079,f7081,f7083,f7085,f7087,f7089,f7102,f7112,f7124,f7134,f7136,f7155,f7157,f7166,f7168,f7187,f7189,f7198,f7200,f7217,f7219,f7228,f7230,f7241,f7252,f7272,f7274,f7285,f7287,f7298,f7300,f7311,f7331,f7333,f7343,f7345,f7347,f7359,f7361,f7368,f7374,f7380,f7391,f7398,f7400,f7405,f7407,f7409,f7411,f7421,f7423,f7433,f7435,f7440,f7469,f7470,f7471,f7472,f7473,f7474,f7475,f7476,f7477,f7478,f7479,f7480,f7481,f7482,f7483,f7484,f7485,f7486,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f7504,f7505,f7506,f7507,f7508,f7509,f7510,f7511,f7512,f7513,f7514,f7515,f7516,f7517,f7518,f7519,f7520,f7521,f7522,f7523,f7524,f7525,f7526,f7527,f7528,f7529,f7530,f7531,f7532,f7533,f7534,f7535,f7536,f7537,f7538,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7547,f7548,f7549,f7550,f7551,f7552,f7553,f7554,f7555,f7556,f7557,f7558,f7559,f7560,f7561,f7562,f7563,f7564,f7565,f7566,f7567,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7577,f7578,f7579,f7580,f7581,f7582,f7583,f7584,f7585,f7586,f7587,f7588,f7589,f7590,f7591,f7592,f7593,f7594,f7595,f7596,f7597,f7598,f7599,f7600,f7601,f7602,f7603,f7604,f7605,f7606,f7607,f7608,f7609,f7610,f7611,f7612,f7613,f7614,f7615,f7616,f7617,f7618,f7619,f7620,f7621,f7622,f7623,f7624,f7625,f7626,f7627,f7628,f7629,f7630,f7631,f7632,f7633,f7634,f7635,f7636,f7637,f7638,f7639,f7640,f7641,f7642,f7643,f7644,f7645,f7646,f7647,f7648,f7649,f7650,f7651,f7652,f7653,f7654,f7655,f7656,f7657,f7658,f7659,f7660,f7661,f7662,f7663,f7664,f7665,f7666,f7667,f7668,f7669,f7670,f7671,f7672,f7673,f7674,f7675,f7676,f7677,f7678,f7679,f7680,f7681,f7683,f7684,f7685,f7686,f7687,f7688,f7689,f7690,f7691,f7692,f7693,f7694,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7706,f7707,f7708,f7709,f7710,f7711,f7712,f7713,f7714,f7715,f7716,f7717,f7718,f7719,f7720,f7721,f7722,f7723,f7724,f7725,f7726,f7727,f7728,f7729,f7730,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7739,f7740,f7741,f7742,f7743,f7744,f7745,f7746,f7747,f7748,f7749,f7750,f7751,f7752,f7753,f7754,f7755,f7756,f7757,f7758,f7759,f7760,f7761,f7762,f7763,f7764,f7765,f7766,f7767,f7768,f7769,f7770,f7771,f7772,f7773,f7774,f7775,f7776,f7777,f7778,f7779,f7780,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7788,f7789,f7790,f7791,f7792,f7793,f7794,f7795,f7796,f7797,f7798,f7799,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7815,f7816,f7817,f7818,f7819,f7820,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7829,f7830,f7831,f7832,f7833,f7834,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f8208,f8210,f8759,f8761,f9666,f9667,f9711,f9712,f9758,f9759,f9805,f9806,f15865,f15866,f15993,f15994,f16055,f16056,f16113,f16114,f16172,f16173,f6136,f20172,f20174,f20191,f20208,f20225])).
fof(f20225,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20224,f9279])).
fof(f9279,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.exhSem.v',X0),'const.pair.,_2'(s('type.exhSem.v',sK52),s(X0,X1))) != s('type.pair.prod'('type.exhSem.v',X0),'const.pair.,_2'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))),s(X0,X2)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1406,f981])).
fof(f981,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X45,X44] : ! [X21,X49,X52,X96] : (s('type.pair.prod'(X45,X44),'const.pair.,_2'(s(X45,X21),s(X44,X49))) = s('type.pair.prod'(X45,X44),'const.pair.,_2'(s(X45,X52),s(X44,X96))) <=> (s(X45,X21) = s(X45,X52) & s(X44,X49) = s(X44,X96)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.pair.CLOSED_PAIR_EQ')).
fof(f20224,plain,(
  s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v',sK52),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) = s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) | $spl0),
  inference(forward_demodulation,[],[f19869,f19921])).
fof(f19921,plain,(
  s('type.num.num',sK54) = s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)) | $spl0),
  inference(unit_resulting_resolution,[],[f16883,f924])).
fof(f924,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0,X1] : ((s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X52,X92] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X92))) <=> s('type.num.num',X52) = s('type.num.num',X92))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_|49||49||split|4')).
fof(f16883,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))) | $spl0),
  inference(unit_resulting_resolution,[],[f3080,f5032,f8627,f10473,f13614,f955,f1031])).
fof(f1031,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220(X2,X1,X0)))) | sP15(X0,X1,X2) | sP14(X0,X1,X2) | sP16(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK219(X2,X1,X0)))) | sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK219(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK219(X2,X1,X0))))) | sP16(X0,X1,X2) | sP14(X0,X1,X2) | sP15(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220(X2,X1,X0))))) | sP17(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP16(X0,X1,X2) & ~sP14(X0,X1,X2) & ~sP15(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP17(X0,X1,X2)) | sP18(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK219,sK220])],[f620])).
fof(f620,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | ? [X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X6))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X6)))) | sP16(X0,X1,X2) | sP14(X0,X1,X2) | sP15(X0,X1,X2) | ? [X5] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP17(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP16(X0,X1,X2) & ~sP14(X0,X1,X2) & ~sP15(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP17(X0,X1,X2)) | sP18(X0,X1,X2)))),
  inference(rectify,[],[f619])).
fof(f619,plain,(
  ! [X0,X2,X1] : ((~sP18(X0,X2,X1) | ? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP16(X0,X2,X1) | sP14(X0,X2,X1) | sP15(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP17(X0,X2,X1)) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP16(X0,X2,X1) & ~sP14(X0,X2,X1) & ~sP15(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP17(X0,X2,X1)) | sP18(X0,X2,X1)))),
  inference(flattening,[],[f618])).
fof(f618,plain,(
  ! [X0,X2,X1] : ((~sP18(X0,X2,X1) | (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP16(X0,X2,X1) | sP14(X0,X2,X1) | sP15(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP17(X0,X2,X1))) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP16(X0,X2,X1) & ~sP14(X0,X2,X1) & ~sP15(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP17(X0,X2,X1)) | sP18(X0,X2,X1)))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0,X2,X1] : (sP18(X0,X2,X1) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP16(X0,X2,X1) | sP14(X0,X2,X1) | sP15(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP17(X0,X2,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f955,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X0))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X92,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X92))) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X52)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_distinct|split|3')).
fof(f13614,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,'const.conSem.Loc_1'(s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f1370,f1049])).
fof(f1049,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK227(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK229(X2,X1,X0)),s('type.list.list'('type.string.char'),sK228(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK227(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK229(X2,X1,X0)),s('type.list.list'('type.string.char'),sK228(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK226(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK230(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK229(X2,X1,X0)))),s('type.list.list'('type.string.char'),sK228(X2,X1,X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK230(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK227(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK226(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP15(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK226,sK227,sK228,sK229,sK230])],[f629])).
fof(f629,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11),s('type.list.list'('type.string.char'),X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11))),s('type.list.list'('type.string.char'),X10))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f628])).
fof(f628,plain,(
  ! [X0,X2,X1] : ((~sP15(X0,X2,X1) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6)))))) & (! [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP15(X0,X2,X1)))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0,X2,X1] : (sP15(X0,X2,X1) <=> ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f1370,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2,X3] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X35,X34,X33,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X33),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X34),s('type.list.list'('type.string.char'),X35)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_distinct|split|12')).
fof(f10473,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,'const.conSem.Loc_1'(s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f1368,f1054])).
fof(f1054,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK231(X2,X1,X0)),s('type.list.list'('type.string.char'),sK232(X2,X1,X0)),s('type.conLang.exp',sK233(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f633])).
fof(f633,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK231(X2,X1,X0)),s('type.list.list'('type.string.char'),sK232(X2,X1,X0)),s('type.conLang.exp',sK233(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK234(X2,X1,X0)),s('type.list.list'('type.string.char'),sK232(X2,X1,X0)),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK235(X2,X1,X0)),s('type.conLang.exp',sK233(X2,X1,X0)))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK235(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK231(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK234(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP14(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK231,sK232,sK233,sK234,sK235])],[f632])).
fof(f632,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.string.char'),X9),s('type.conLang.exp',X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11),s('type.list.list'('type.string.char'),X9),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.conLang.exp',X10))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP14(X0,X1,X2)))),
  inference(rectify,[],[f631])).
fof(f631,plain,(
  ! [X0,X2,X1] : ((~sP14(X0,X2,X1) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14)))))) & (! [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | sP14(X0,X2,X1)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X2,X1] : (sP14(X0,X2,X1) <=> ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f1368,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2,X3] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2,X3] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X35,X34,X33,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X33),s('type.list.list'('type.string.char'),X34),s('type.conLang.exp',X35)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_distinct|split|10')).
fof(f8627,plain,(
  ( ! [X2,X0,X1] : (~sP16(X0,X1,'const.conSem.Loc_1'(s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f1339,f1045])).
fof(f1045,plain,(
  ( ! [X2,X0,X1] : (~sP16(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK223(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK224(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK223(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK224(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK223(X2,X1,X0)))),s('type.list.list'('type.exhSem.v'),sK225(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK224(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK225(X2,X1,X0))))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP16(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK223,sK224,sK225])],[f626])).
fof(f626,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | ? [X6,X7,X8] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6),s('type.list.list'('type.conSem.v'),X7))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6))),s('type.list.list'('type.exhSem.v'),X8))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X8)))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f625])).
fof(f625,plain,(
  ! [X0,X2,X1] : ((~sP16(X0,X2,X1) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18)))))) & (! [X16,X17,X18] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | sP16(X0,X2,X1)))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X2,X1] : (sP16(X0,X2,X1) <=> ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f1339,plain,(
  ( ! [X2,X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X34,X33,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X33),s('type.list.list'('type.conSem.v'),X34)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_distinct|split|7')).
fof(f5032,plain,(
  sP18(sK53,sK52,'const.conSem.Loc_1'(s('type.num.num',sK54))) | $spl0),
  inference(unit_resulting_resolution,[],[f1403,f1059])).
fof(f1059,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | sP18(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | sP18(X0,X2,X1)) & (~sP18(X0,X2,X1) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> sP18(X0,X2,X1))),
  inference(definition_folding,[],[f278,f430,f429,f428,f427,f426])).
fof(f429,plain,(
  ! [X0,X2,X1] : (sP17(X0,X2,X1) <=> ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f278,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X33,X34,X35] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.conSem.v',X34),s('type.exhSem.v',X35)))) <=> (? [X8] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X8))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X8)))) | ? [X17,X18,X19] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.conSem.v'),X18))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X17))),s('type.list.list'('type.exhSem.v'),X19))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.conSem.v'),X18),s('type.list.list'('type.exhSem.v'),X19))))) | ? [X20,X21,X22,X23,X24] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.string.char'),X21),s('type.conLang.exp',X22))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X21),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.conLang.exp',X22))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X23,X20,X21,X25,X24] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25),s('type.list.list'('type.string.char'),X21))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25))),s('type.list.list'('type.string.char'),X21))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X8] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X8))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X8)))) | ? [X18,X19] : (s('type.conSem.v',X34) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X18))) & s('type.exhSem.v',X35) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X19))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.conSem.v'),X18),s('type.list.list'('type.exhSem.v'),X19)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.dec_to_exhProof.v_rel_cases|split|0')).
fof(f3080,plain,(
  ( ! [X2,X0,X1] : (~sP17(X0,X1,'const.conSem.Loc_1'(s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f957,f1041])).
fof(f1041,plain,(
  ( ! [X2,X0,X1] : (~sP17(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK221(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK221(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK222(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK221(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK222(X2,X1,X0))))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP17(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK221,sK222])],[f623])).
fof(f623,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | ? [X5,X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP17(X0,X1,X2)))),
  inference(rectify,[],[f622])).
fof(f622,plain,(
  ! [X0,X2,X1] : ((~sP17(X0,X2,X1) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))) & (! [X3,X4] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP17(X0,X2,X1)))),
  inference(nnf_transformation,[],[f429])).
fof(f957,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X92,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X92)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.conSem.v_distinct|split|14')).
fof(f19869,plain,(
  s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v',sK52),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) = s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f16905,f16883,f983])).
fof(f983,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f581])).
fof(f16905,plain,(
  s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))) | $spl0),
  inference(unit_resulting_resolution,[],[f3080,f5032,f8627,f10473,f13614,f955,f1033])).
fof(f1033,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220(X2,X1,X0)))) | sP15(X0,X1,X2) | sP14(X0,X1,X2) | sP16(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK219(X2,X1,X0)))) | sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f621])).
fof(f20208,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20207,f9279])).
fof(f20207,plain,(
  s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v',sK52),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) = s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) | $spl0),
  inference(forward_demodulation,[],[f19886,f19921])).
fof(f19886,plain,(
  s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v',sK52),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))))) = s('type.pair.prod'('type.exhSem.v','type.conSem.v'),'const.pair.,_2'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))))) | $spl0),
  inference(unit_resulting_resolution,[],[f16905,f16883,f983])).
fof(f20191,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20190,f9368])).
fof(f9368,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,'type.exhSem.v'),'const.pair.,_2'(s(X0,X1),s('type.exhSem.v',sK52))) != s('type.pair.prod'(X0,'type.exhSem.v'),'const.pair.,_2'(s(X0,X2),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1406,f982])).
fof(f982,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f581])).
fof(f20190,plain,(
  s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))) = s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))))) | $spl0),
  inference(forward_demodulation,[],[f19903,f19921])).
fof(f19903,plain,(
  s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))) = s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f16905,f16883,f983])).
fof(f20174,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20173,f9368])).
fof(f20173,plain,(
  s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))) = s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))))) | $spl0),
  inference(forward_demodulation,[],[f19920,f19921])).
fof(f19920,plain,(
  s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))),s('type.exhSem.v',sK52))) = s('type.pair.prod'('type.conSem.v','type.exhSem.v'),'const.pair.,_2'(s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK220('const.conSem.Loc_1'(s('type.num.num',sK54)),sK52,sK53)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f16905,f16883,f983])).
fof(f20172,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f20155,f1406])).
fof(f20155,plain,(
  s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) | $spl0),
  inference(backward_demodulation,[],[f19921,f16905])).
fof(f6136,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f6136_D])).
fof(f6136_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f16173,plain,(
  ( ! [X5] : (p(s(bool,sK134('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134('_i_0_0',X5,bool)))))) ) | $spl32),
  inference(cnf_transformation,[],[f16173_D])).
fof(f16173_D,plain,(
  ( ! [X5] : (p(s(bool,sK134('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134('_i_0_0',X5,bool)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f16172,plain,(
  ( ! [X4,X5] : (p(s(bool,sK134('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134('_i_0_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f16157,f869])).
fof(f869,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X99] : ! [X101] : s(X99,X101) = s(X99,happ(s(fun(X99,X99),'_i_0_0'),s(X99,X101)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','_i0_ax')).
fof(f16157,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134('_i_0_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK134('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f906,f869])).
fof(f906,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK134(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK134(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK132(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK133(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK134(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK134(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK132,sK133,sK134])],[f548])).
fof(f548,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f547])).
fof(f547,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X45] : ! [X47,X46] : (? [X21] : (p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) | p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))) <=> (? [X21] : p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) | ? [X21] : p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.EXISTS_OR_THM')).
fof(f16114,plain,(
  ( ! [X5] : (p(s(bool,sK134(X5,'_i_0_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134(X5,'_i_0_0',bool)))))) ) | $spl30),
  inference(cnf_transformation,[],[f16114_D])).
fof(f16114_D,plain,(
  ( ! [X5] : (p(s(bool,sK134(X5,'_i_0_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134(X5,'_i_0_0',bool)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f16113,plain,(
  ( ! [X4,X5] : (p(s(bool,sK134(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134(X5,'_i_0_0',bool)))))) )),
  inference(forward_demodulation,[],[f16098,f869])).
fof(f16098,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK134(X5,'_i_0_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK134(X5,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f905,f869])).
fof(f905,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK134(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK134(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f549])).
fof(f16056,plain,(
  ( ! [X5] : (p(s(bool,sK133('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK132('_i_0_0',X5,bool)))))) ) | $spl28),
  inference(cnf_transformation,[],[f16056_D])).
fof(f16056_D,plain,(
  ( ! [X5] : (p(s(bool,sK133('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK132('_i_0_0',X5,bool)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f16055,plain,(
  ( ! [X4,X5] : (p(s(bool,sK133('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK132('_i_0_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f16040,f869])).
fof(f16040,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK132('_i_0_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK133('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f904,f869])).
fof(f904,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK132(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK133(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f549])).
fof(f15994,plain,(
  ( ! [X5] : (p(s(bool,sK132(X5,'_i_0_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK133(X5,'_i_0_0',bool)))))) ) | $spl26),
  inference(cnf_transformation,[],[f15994_D])).
fof(f15994_D,plain,(
  ( ! [X5] : (p(s(bool,sK132(X5,'_i_0_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK133(X5,'_i_0_0',bool)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f15993,plain,(
  ( ! [X4,X5] : (p(s(bool,sK132(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK133(X5,'_i_0_0',bool)))))) )),
  inference(forward_demodulation,[],[f15978,f869])).
fof(f15978,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK132(X5,'_i_0_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK133(X5,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f903,f869])).
fof(f903,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK132(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK133(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f549])).
fof(f15866,plain,(
  ( ! [X5] : (p(s(bool,sK123('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK124('_i_0_0',X5,bool)))))) ) | $spl24),
  inference(cnf_transformation,[],[f15866_D])).
fof(f15866_D,plain,(
  ( ! [X5] : (p(s(bool,sK123('_i_0_0',X5,bool))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK124('_i_0_0',X5,bool)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f15865,plain,(
  ( ! [X4,X5] : (p(s(bool,sK123('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK124('_i_0_0',X5,bool)))))) )),
  inference(forward_demodulation,[],[f15850,f869])).
fof(f15850,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,sK124('_i_0_0',X5,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK123('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f884,f869])).
fof(f884,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK124(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK123(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK123(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK123(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK124(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK123,sK124])],[f528])).
fof(f528,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f362])).
fof(f362,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f361])).
fof(f361,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X45] : ! [X46,X47] : (! [X21] : (p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) => p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))) => (? [X21] : p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) => ? [X21] : p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.MONO_EXISTS')).
fof(f9806,plain,(
  ( ! [X5] : (p(s(bool,sK128(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f9806_D])).
fof(f9806_D,plain,(
  ( ! [X5] : (p(s(bool,sK128(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f9805,plain,(
  ( ! [X4,X5] : (p(s(bool,sK128(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9793,f869])).
fof(f9793,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK128(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f894,f869])).
fof(f894,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK128(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK127(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK128(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK127,sK128])],[f537])).
fof(f537,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f536])).
fof(f536,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X45] : ! [X47,X46] : ((? [X21] : p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) | p(s(bool,X46))) <=> ? [X21] : (p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) | p(s(bool,X46))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f9759,plain,(
  ( ! [X5] : (p(s(bool,sK127(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl20),
  inference(cnf_transformation,[],[f9759_D])).
fof(f9759_D,plain,(
  ( ! [X5] : (p(s(bool,sK127(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9758,plain,(
  ( ! [X4,X5] : (p(s(bool,sK127(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9746,f869])).
fof(f9746,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK127(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f892,f869])).
fof(f892,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK127(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f538])).
fof(f9712,plain,(
  ( ! [X5] : (p(s(bool,sK126('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f9712_D])).
fof(f9712_D,plain,(
  ( ! [X5] : (p(s(bool,sK126('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9711,plain,(
  ( ! [X4,X5] : (p(s(bool,sK126('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9699,f869])).
fof(f9699,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK126('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f891,f869])).
fof(f891,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK126(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK125(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK126(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK125,sK126])],[f533])).
fof(f533,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f531])).
fof(f531,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X45] : ! [X47,X46] : ((p(s(bool,X47)) | ? [X21] : p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))) <=> ? [X21] : (p(s(bool,X47)) | p(s(bool,happ(s(fun(X45,bool),X46),s(X45,X21))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f9667,plain,(
  ( ! [X5] : (p(s(bool,sK125('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl16),
  inference(cnf_transformation,[],[f9667_D])).
fof(f9667_D,plain,(
  ( ! [X5] : (p(s(bool,sK125('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f9666,plain,(
  ( ! [X4,X5] : (p(s(bool,sK125('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9654,f869])).
fof(f9654,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK125('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f889,f869])).
fof(f889,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK125(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f534])).
fof(f8761,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl14),
  inference(cnf_transformation,[],[f8761_D])).
fof(f8761_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8759,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f8748,f869])).
fof(f8748,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f881,f869])).
fof(f881,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X45] : ! [X47,X21] : (p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) => p(s(bool,happ(s(fun(X45,bool),X47),s(X45,'const.min.@_1'(s(fun(X45,bool),X47)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.SELECT_AX')).
fof(f8210,plain,(
  p(s(bool,sK121('_i_0_0',bool))) | $spl12),
  inference(cnf_transformation,[],[f8210_D])).
fof(f8210_D,plain,(
  p(s(bool,sK121('_i_0_0',bool))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f8208,plain,(
  ( ! [X3] : (p(s(bool,sK121('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f8197,f869])).
fof(f8197,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK121('_i_0_0',bool)))))) )),
  inference(superposition,[],[f875,f869])).
fof(f875,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK121(X1,X0)))))) )),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK120(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK121(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK120,sK121])],[f521])).
fof(f521,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f520])).
fof(f520,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X45] : ! [X47] : (~? [X21] : p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.NOT_EXISTS_THM')).
fof(f7845,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6141,f6152])).
fof(f6152,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f6136,f916])).
fof(f916,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X42,X43] : ((p(s(bool,X42)) => p(s(bool,X43))) => ((p(s(bool,X43)) => p(s(bool,X42))) => s(bool,X42) = s(bool,X43)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f6141,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f813])).
fof(f813,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f7844,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6149,f1407])).
fof(f1407,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f776,f773])).
fof(f773,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2',aHL_TRUTH)).
fof(f776,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f48])).
fof(f48,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f6149,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1445,f6136,f916])).
fof(f1445,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1407,f813])).
fof(f7843,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6155,f1407])).
fof(f6155,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1445,f6136,f916])).
fof(f7842,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6157,f1407])).
fof(f6157,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1469,f6136,f945])).
fof(f945,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1)))),
  inference(rectify,[],[f562])).
fof(f562,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0)))),
  inference(flattening,[],[f561])).
fof(f561,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP2(X1,X0)))),
  inference(nnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X1,X0] : (sP2(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1469,plain,(
  sP2('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1407,f773,f947])).
fof(f947,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f563])).
fof(f7841,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6159,f6136])).
fof(f6159,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1495,f6136,f945])).
fof(f1495,plain,(
  ( ! [X0] : (sP2('F_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1407,f1476])).
fof(f1476,plain,(
  ( ! [X0,X1] : (sP2(X1,'const.bool.BOUNDED_1'(s(bool,X0))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1475,f773])).
fof(f1475,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP2(X1,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(superposition,[],[f950,f780])).
fof(f780,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X6)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.BOUNDED_THM')).
fof(f950,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f563])).
fof(f7840,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6160,f6136])).
fof(f6160,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1646,f6136,f945])).
fof(f1646,plain,(
  ( ! [X0] : (sP2('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f950])).
fof(f781,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X45] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X45,X45),'_i_0_0'),s(fun(X45,X45),'_i_0_0'))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.ind_type.ISO_REFL')).
fof(f7839,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6161,f6136])).
fof(f6161,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4102,f6136,f945])).
fof(f4102,plain,(
  sP2('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f950])).
fof(f7838,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6162,f1407])).
fof(f6162,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1490,f6136,f945])).
fof(f1490,plain,(
  ( ! [X0] : (sP2('const.bool.BOUNDED_1'(s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f1471])).
fof(f1471,plain,(
  ( ! [X0,X1] : (sP2('const.bool.BOUNDED_1'(s(bool,X0)),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1470,f773])).
fof(f1470,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP2('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f947,f780])).
fof(f7837,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6163,f1407])).
fof(f6163,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1645,f6136,f945])).
fof(f1645,plain,(
  ( ! [X0] : (sP2('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f947])).
fof(f7836,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6164,f1407])).
fof(f6164,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4101,f6136,f945])).
fof(f4101,plain,(
  sP2('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f947])).
fof(f7835,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6170,f1407])).
fof(f6170,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1474,f6136,f945])).
fof(f1474,plain,(
  sP2('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1407,f773,f950])).
fof(f7834,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6171,f1407])).
fof(f6171,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1495,f6136,f945])).
fof(f7833,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6172,f1407])).
fof(f6172,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1646,f6136,f945])).
fof(f7832,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6173,f1407])).
fof(f6173,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4102,f6136,f945])).
fof(f7831,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6174,f6136])).
fof(f6174,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1490,f6136,f945])).
fof(f7830,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6175,f6136])).
fof(f6175,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1645,f6136,f945])).
fof(f7829,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6176,f6136])).
fof(f6176,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4101,f6136,f945])).
fof(f7828,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6215,f6486])).
fof(f6486,plain,(
  ( ! [X0] : (sP45('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1411,f1414,f6136,f1270])).
fof(f1270,plain,(
  ( ! [X2,X0,X1] : (~sP43(X1,X2,X0) | p(s(bool,X0)) | ~sP44(X0,X2,X1) | sP45(X0,X1,X2) | ~sP42(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f728])).
fof(f728,plain,(
  ! [X0,X1,X2] : ((~sP45(X0,X1,X2) | (sP42(X2,X1,X0) & sP43(X1,X2,X0) & sP44(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP42(X2,X1,X0) | ~sP43(X1,X2,X0) | ~sP44(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP45(X0,X1,X2)))),
  inference(rectify,[],[f727])).
fof(f727,plain,(
  ! [X2,X1,X0] : ((~sP45(X2,X1,X0) | (sP42(X0,X1,X2) & sP43(X1,X0,X2) & sP44(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP42(X0,X1,X2) | ~sP43(X1,X0,X2) | ~sP44(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP45(X2,X1,X0)))),
  inference(flattening,[],[f726])).
fof(f726,plain,(
  ! [X2,X1,X0] : ((~sP45(X2,X1,X0) | (sP42(X0,X1,X2) & sP43(X1,X0,X2) & sP44(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP42(X0,X1,X2) | ~sP43(X1,X0,X2) | ~sP44(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP45(X2,X1,X0)))),
  inference(nnf_transformation,[],[f463])).
fof(f463,plain,(
  ! [X2,X1,X0] : (sP45(X2,X1,X0) <=> (sP42(X0,X1,X2) & sP43(X1,X0,X2) & sP44(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP45])])).
fof(f1414,plain,(
  ( ! [X0,X1] : (sP43(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1276])).
fof(f1276,plain,(
  ( ! [X2,X0,X1] : (sP43(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0,X1,X2] : ((~sP43(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP43(X0,X1,X2)))),
  inference(rectify,[],[f733])).
fof(f733,plain,(
  ! [X1,X0,X2] : ((~sP43(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP43(X1,X0,X2)))),
  inference(flattening,[],[f732])).
fof(f732,plain,(
  ! [X1,X0,X2] : ((~sP43(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP43(X1,X0,X2)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X1,X0,X2] : (sP43(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP43])])).
fof(f1411,plain,(
  ( ! [X0,X1] : (sP44(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1272])).
fof(f1272,plain,(
  ( ! [X2,X0,X1] : (sP44(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f731])).
fof(f731,plain,(
  ! [X0,X1,X2] : ((~sP44(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP44(X0,X1,X2)))),
  inference(rectify,[],[f730])).
fof(f730,plain,(
  ! [X2,X0,X1] : ((~sP44(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP44(X2,X0,X1)))),
  inference(flattening,[],[f729])).
fof(f729,plain,(
  ! [X2,X0,X1] : ((~sP44(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP44(X2,X0,X1)))),
  inference(nnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X2,X0,X1] : (sP44(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP44])])).
fof(f1417,plain,(
  ( ! [X0,X1] : (sP42(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1280])).
fof(f1280,plain,(
  ( ! [X2,X0,X1] : (sP42(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f736])).
fof(f736,plain,(
  ! [X0,X1,X2] : ((~sP42(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP42(X0,X1,X2)))),
  inference(flattening,[],[f735])).
fof(f735,plain,(
  ! [X0,X1,X2] : ((~sP42(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP42(X0,X1,X2)))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1,X2] : (sP42(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP42])])).
fof(f6215,plain,(
  ( ! [X0] : (~sP45('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f773,f6136,f1267])).
fof(f1267,plain,(
  ( ! [X2,X0,X1] : (~sP45(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f728])).
fof(f7827,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6217,f6656])).
fof(f6656,plain,(
  ( ! [X2,X0,X1] : (sP43(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1278])).
fof(f1278,plain,(
  ( ! [X2,X0,X1] : (sP43(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f734])).
fof(f6217,plain,(
  ~sP43('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f6136,f1269])).
fof(f1269,plain,(
  ( ! [X2,X0,X1] : (~sP43(X1,X2,X0) | p(s(bool,X1)) | ~sP44(X0,X2,X1) | sP45(X0,X1,X2) | ~sP42(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f728])).
fof(f2203,plain,(
  ~sP45('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f773,f1445,f1285])).
fof(f1285,plain,(
  ( ! [X2,X0,X1] : (~sP45(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP45(X2,X1,X0)) & (~sP45(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP45(X2,X1,X0))),
  inference(definition_folding,[],[f287,f463,f462,f461,f460])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X89,X90,X91] : ((p(s(bool,X91)) <=> s(bool,X89) = s(bool,X90)) <=> ((p(s(bool,X91)) | p(s(bool,X90)) | p(s(bool,X89))) & (p(s(bool,X91)) | ~p(s(bool,X89)) | ~p(s(bool,X90))) & (p(s(bool,X90)) | ~p(s(bool,X89)) | ~p(s(bool,X91))) & (p(s(bool,X89)) | ~p(s(bool,X90)) | ~p(s(bool,X91)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sat.sat_2|49|_prove_ltb')).
fof(f1418,plain,(
  ( ! [X0,X1] : (sP42(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f773,f1281])).
fof(f1281,plain,(
  ( ! [X2,X0,X1] : (sP42(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f736])).
fof(f7826,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6223,f6656])).
fof(f6223,plain,(
  ~sP43('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f6136,f1269])).
fof(f4174,plain,(
  ~sP45('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1445,f4087,f1285])).
fof(f7825,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6224,f6656])).
fof(f6224,plain,(
  ( ! [X0] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2220,f1647,f6136,f1269])).
fof(f1647,plain,(
  ( ! [X2,X0,X1] : (sP44(X0,X1,'const.ind_type.ISO_2'(s(fun(X2,X2),'_i_0_0'),s(fun(X2,X2),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f781,f1272])).
fof(f2220,plain,(
  ( ! [X0] : (~sP45('F_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1407,f869,f1286])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (~sP45(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f737])).
fof(f1650,plain,(
  ( ! [X2,X0,X1] : (sP42(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f781,f1281])).
fof(f7824,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6225,f6656])).
fof(f6225,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2221,f1647,f6136,f1269])).
fof(f2221,plain,(
  ( ! [X0,X1] : (~sP45('F_0',X0,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f1407,f958,f1286])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X44,X45] : ! [X21] : s(X45,X21) = s(X45,'const.sum.OUTL_1'(s('type.sum.sum'(X45,X44),'const.sum.INL_1'(s(X45,X21)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sum.OUTL')).
fof(f7823,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6226,f6656])).
fof(f6226,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2222,f1647,f6136,f1269])).
fof(f2222,plain,(
  ( ! [X0,X1] : (~sP45('F_0',X0,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f1407,f959,f1286])).
fof(f959,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X45,X44] : ! [X21] : s(X44,X21) = s(X44,'const.sum.OUTR_1'(s('type.sum.sum'(X45,X44),'const.sum.INR_1'(s(X44,X21)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sum.OUTR')).
fof(f7822,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6227,f6656])).
fof(f6227,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f6136,f1269])).
fof(f4137,plain,(
  ( ! [X0,X1] : (sP44(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1272])).
fof(f4105,plain,(
  ~sP45('T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f773,f1407,f4087,f1267])).
fof(f4154,plain,(
  ( ! [X0,X1] : (sP42(X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1281])).
fof(f7821,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6228,f6656])).
fof(f6228,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2220,f4137,f6136,f1269])).
fof(f7820,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6229,f6656])).
fof(f6229,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2221,f4137,f6136,f1269])).
fof(f7819,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6230,f6656])).
fof(f6230,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2222,f4137,f6136,f1269])).
fof(f7818,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6231,f6656])).
fof(f6231,plain,(
  ~sP43('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f6136,f1269])).
fof(f1412,plain,(
  ( ! [X0,X1] : (sP44(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1273])).
fof(f1273,plain,(
  ( ! [X2,X0,X1] : (sP44(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f731])).
fof(f7817,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6232,f6656])).
fof(f6232,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f6136,f1269])).
fof(f7816,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6233,f6656])).
fof(f6233,plain,(
  ~sP43('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f6136,f1269])).
fof(f7815,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6235,f6656])).
fof(f6235,plain,(
  ( ! [X0] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2220,f1413,f6136,f1269])).
fof(f1413,plain,(
  ( ! [X0,X1] : (sP44('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1274])).
fof(f1274,plain,(
  ( ! [X2,X0,X1] : (sP44(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f731])).
fof(f7814,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6236,f6656])).
fof(f6236,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2220,f1413,f6136,f1269])).
fof(f7813,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6238,f6656])).
fof(f6238,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2221,f1413,f6136,f1269])).
fof(f7812,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6239,f6656])).
fof(f6239,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2221,f1413,f6136,f1269])).
fof(f7811,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6241,f6656])).
fof(f6241,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2222,f1413,f6136,f1269])).
fof(f7810,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6242,f6656])).
fof(f6242,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2222,f1413,f6136,f1269])).
fof(f7809,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6246,f6656])).
fof(f6246,plain,(
  ( ! [X0] : (~sP43(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2214,f1413,f6136,f1269])).
fof(f2214,plain,(
  ( ! [X0] : (~sP45('F_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1407,f869,f1286])).
fof(f1651,plain,(
  ( ! [X2,X0,X1] : (sP42('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f781,f1282])).
fof(f1282,plain,(
  ( ! [X2,X0,X1] : (sP42(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f736])).
fof(f7808,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6247,f6656])).
fof(f6247,plain,(
  ~sP43(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2214,f1413,f6136,f1269])).
fof(f4155,plain,(
  ( ! [X0,X1] : (sP42('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1282])).
fof(f7807,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6250,f6656])).
fof(f6250,plain,(
  ( ! [X0,X1] : (~sP43('const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))))),'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2216,f1413,f6136,f1269])).
fof(f2216,plain,(
  ( ! [X0,X1] : (~sP45('F_0','const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f959,f1286])).
fof(f7806,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6251,f6656])).
fof(f6251,plain,(
  ( ! [X0] : (~sP43('const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2216,f1413,f6136,f1269])).
fof(f7805,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6253,f6656])).
fof(f6253,plain,(
  ( ! [X0,X1] : (~sP43('const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))))),'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2215,f1413,f6136,f1269])).
fof(f2215,plain,(
  ( ! [X0,X1] : (~sP45('F_0','const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f958,f1286])).
fof(f7804,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6254,f6656])).
fof(f6254,plain,(
  ( ! [X0] : (~sP43('const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2215,f1413,f6136,f1269])).
fof(f7803,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6256,f6642])).
fof(f6642,plain,(
  ( ! [X2,X0,X1] : (sP44(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1274])).
fof(f6256,plain,(
  ~sP44('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1414,f6136,f1269])).
fof(f7802,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6257,f6642])).
fof(f6257,plain,(
  ~sP44('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1414,f6136,f1269])).
fof(f7801,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6330,f6642])).
fof(f6330,plain,(
  ~sP44('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f4152,f6136,f1269])).
fof(f4152,plain,(
  ( ! [X0,X1] : (sP43(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1276])).
fof(f7800,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6355,f6642])).
fof(f6355,plain,(
  ~sP44('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1415,f6136,f1269])).
fof(f1415,plain,(
  ( ! [X0,X1] : (sP43(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1277])).
fof(f1277,plain,(
  ( ! [X2,X0,X1] : (sP43(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f734])).
fof(f7799,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6356,f6642])).
fof(f6356,plain,(
  ~sP44('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1415,f6136,f1269])).
fof(f7798,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6357,f6642])).
fof(f6357,plain,(
  ~sP44('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1415,f6136,f1269])).
fof(f7797,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6418,f6656])).
fof(f6418,plain,(
  ~sP43('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f6136,f1270])).
fof(f7796,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6424,f6656])).
fof(f6424,plain,(
  ~sP43('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f6136,f1270])).
fof(f7795,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6425,f6656])).
fof(f6425,plain,(
  ( ! [X0] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2220,f1647,f6136,f1270])).
fof(f7794,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6426,f6656])).
fof(f6426,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2221,f1647,f6136,f1270])).
fof(f7793,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6427,f6656])).
fof(f6427,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2222,f1647,f6136,f1270])).
fof(f7792,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6428,f6656])).
fof(f6428,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f6136,f1270])).
fof(f7791,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6429,f6656])).
fof(f6429,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2220,f4137,f6136,f1270])).
fof(f7790,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6430,f6656])).
fof(f6430,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2221,f4137,f6136,f1270])).
fof(f7789,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6431,f6656])).
fof(f6431,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2222,f4137,f6136,f1270])).
fof(f7788,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6432,f6656])).
fof(f6432,plain,(
  ~sP43('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f6136,f1270])).
fof(f7787,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6433,f6656])).
fof(f6433,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f6136,f1270])).
fof(f7786,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6434,f6656])).
fof(f6434,plain,(
  ~sP43('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f6136,f1270])).
fof(f7785,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6436,f6656])).
fof(f6436,plain,(
  ( ! [X0] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2220,f1413,f6136,f1270])).
fof(f7784,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6437,f6656])).
fof(f6437,plain,(
  ~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2220,f1413,f6136,f1270])).
fof(f7783,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6439,f6656])).
fof(f6439,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2221,f1413,f6136,f1270])).
fof(f7782,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6440,f6656])).
fof(f6440,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2221,f1413,f6136,f1270])).
fof(f7781,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6442,f6656])).
fof(f6442,plain,(
  ( ! [X0,X1] : (~sP43('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f2222,f1413,f6136,f1270])).
fof(f7780,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6443,f6656])).
fof(f6443,plain,(
  ( ! [X0] : (~sP43('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f2222,f1413,f6136,f1270])).
fof(f7779,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6447,f6656])).
fof(f6447,plain,(
  ( ! [X0] : (~sP43(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))),'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2214,f1413,f6136,f1270])).
fof(f7778,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6448,f6656])).
fof(f6448,plain,(
  ~sP43(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2214,f1413,f6136,f1270])).
fof(f7777,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6451,f6656])).
fof(f6451,plain,(
  ( ! [X0,X1] : (~sP43('const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))))),'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2216,f1413,f6136,f1270])).
fof(f7776,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6452,f6656])).
fof(f6452,plain,(
  ( ! [X0] : (~sP43('const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2216,f1413,f6136,f1270])).
fof(f7775,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6454,f6656])).
fof(f6454,plain,(
  ( ! [X0,X1] : (~sP43('const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))))),'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f2215,f1413,f6136,f1270])).
fof(f7774,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6455,f6656])).
fof(f6455,plain,(
  ( ! [X0] : (~sP43('const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f2215,f1413,f6136,f1270])).
fof(f7773,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6457,f6642])).
fof(f6457,plain,(
  ~sP44('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1414,f6136,f1270])).
fof(f7772,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6458,f6642])).
fof(f6458,plain,(
  ~sP44('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1414,f6136,f1270])).
fof(f7771,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6531,f6642])).
fof(f6531,plain,(
  ~sP44('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f4152,f6136,f1270])).
fof(f7770,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6556,f6642])).
fof(f6556,plain,(
  ~sP44('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1415,f6136,f1270])).
fof(f7769,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6557,f6642])).
fof(f6557,plain,(
  ~sP44('T_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1415,f6136,f1270])).
fof(f7768,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6558,f6642])).
fof(f6558,plain,(
  ~sP44('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1415,f6136,f1270])).
fof(f7767,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6588,f4105])).
fof(f6588,plain,(
  sP45('T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1417,f4137,f1415,f6136,f1270])).
fof(f7766,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6590,f4106])).
fof(f4106,plain,(
  ~sP45('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1407,f4087,f1267])).
fof(f6590,plain,(
  sP45('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f4137,f1415,f6136,f1270])).
fof(f4153,plain,(
  ( ! [X0,X1] : (sP42(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1280])).
fof(f7765,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6638,f6642])).
fof(f6638,plain,(
  ( ! [X0] : (~sP44('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f773,f6136,f1271])).
fof(f1271,plain,(
  ( ! [X2,X0,X1] : (~sP44(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f731])).
fof(f7764,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6652,f6656])).
fof(f6652,plain,(
  ( ! [X0] : (~sP43('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f773,f6136,f1275])).
fof(f1275,plain,(
  ( ! [X2,X0,X1] : (~sP43(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f734])).
fof(f7763,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6659,f6136])).
fof(f6659,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f1407,f6136,f1279])).
fof(f1279,plain,(
  ( ! [X2,X0,X1] : (~sP42(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f736])).
fof(f7762,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6660,f6136])).
fof(f6660,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f1407,f6136,f1279])).
fof(f7761,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6661,f6136])).
fof(f6661,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f1407,f6136,f1279])).
fof(f1419,plain,(
  ( ! [X0,X1] : (sP42('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f773,f1282])).
fof(f7760,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6662,f6136])).
fof(f6662,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f1407,f6136,f1279])).
fof(f7759,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6663,f6136])).
fof(f6663,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f1407,f6136,f1279])).
fof(f7758,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6664,f6136])).
fof(f6664,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f1279])).
fof(f7757,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6665,f6136])).
fof(f6665,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f1649,plain,(
  ( ! [X2,X0,X1] : (sP42(X0,X1,'const.ind_type.ISO_2'(s(fun(X2,X2),'_i_0_0'),s(fun(X2,X2),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f781,f1280])).
fof(f7756,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6666,f6136])).
fof(f6666,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7755,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6668,f6136])).
fof(f6668,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7754,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6669,f6136])).
fof(f6669,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7753,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6670,f6136])).
fof(f6670,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6136,f6136,f1279])).
fof(f7752,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6671,f6136])).
fof(f6671,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7751,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6672,f6136])).
fof(f6672,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7750,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6674,f6136])).
fof(f6674,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f1407,f6136,f1279])).
fof(f7749,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6675,f6136])).
fof(f6675,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f1407,f6136,f1279])).
fof(f7748,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6676,f6136])).
fof(f6676,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f1407,f6136,f1279])).
fof(f7747,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6677,f6136])).
fof(f6677,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f1407,f6136,f1279])).
fof(f7746,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6678,f6136])).
fof(f6678,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f1407,f6136,f1279])).
fof(f7745,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6684,f6136])).
fof(f6684,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f7744,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6685,f6136])).
fof(f6685,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7743,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6686,f6136])).
fof(f6686,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f6136,f6136,f1279])).
fof(f7742,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6687,f6136])).
fof(f6687,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7741,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6688,f6136])).
fof(f6688,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7740,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6689,f6136])).
fof(f6689,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6136,f6136,f1279])).
fof(f7739,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6690,f6136])).
fof(f6690,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7738,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6691,f6136])).
fof(f6691,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7737,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6697,f6136])).
fof(f6697,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f1407,f6136,f1279])).
fof(f7736,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6698,f6136])).
fof(f6698,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f1407,f6136,f1279])).
fof(f7735,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6699,f6136])).
fof(f6699,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f1407,f6136,f1279])).
fof(f7734,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6700,f6136])).
fof(f6700,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f1407,f6136,f1279])).
fof(f7733,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6701,f6136])).
fof(f6701,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f1407,f6136,f1279])).
fof(f7732,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6703,f6136])).
fof(f6703,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f7731,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6704,f6136])).
fof(f6704,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7730,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6705,f6136])).
fof(f6705,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f6136,f6136,f1279])).
fof(f7729,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6706,f6136])).
fof(f6706,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7728,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6707,f6136])).
fof(f6707,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7727,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6708,f6136])).
fof(f6708,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6136,f6136,f1279])).
fof(f7726,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6709,f6136])).
fof(f6709,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7725,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6710,f6136])).
fof(f6710,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7724,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6711,f6136])).
fof(f6711,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f1407,f6136,f1279])).
fof(f7723,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6712,f6136])).
fof(f6712,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f1407,f6136,f1279])).
fof(f7722,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6713,f6136])).
fof(f6713,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f1407,f6136,f1279])).
fof(f7721,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6715,f6136])).
fof(f6715,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f1407,f6136,f1279])).
fof(f7720,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6716,f6136])).
fof(f6716,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f1407,f6136,f1279])).
fof(f7719,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6721,f6136])).
fof(f6721,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f1279])).
fof(f7718,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6722,f6136])).
fof(f6722,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f7717,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6723,f6136])).
fof(f6723,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7716,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6724,f6136])).
fof(f6724,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f6136,f6136,f1279])).
fof(f7715,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6725,f6136])).
fof(f6725,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7714,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6726,f6136])).
fof(f6726,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7713,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6728,f6136])).
fof(f6728,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7712,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6729,f6136])).
fof(f6729,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7711,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6734,f6136])).
fof(f6734,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1407,f6136,f1279])).
fof(f7710,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6735,f6136])).
fof(f6735,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f1407,f6136,f1279])).
fof(f7709,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6736,f6136])).
fof(f6736,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f1407,f6136,f1279])).
fof(f7708,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6738,f6136])).
fof(f6738,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f1407,f6136,f1279])).
fof(f7707,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6739,f6136])).
fof(f6739,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f1407,f6136,f1279])).
fof(f7706,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6740,f6136])).
fof(f6740,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f1279])).
fof(f7705,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6741,f6136])).
fof(f6741,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f7704,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6742,f6136])).
fof(f6742,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7703,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6744,f6136])).
fof(f6744,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7702,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6745,f6136])).
fof(f6745,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7701,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6746,f6136])).
fof(f6746,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6136,f6136,f1279])).
fof(f7700,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6747,f6136])).
fof(f6747,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7699,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6748,f6136])).
fof(f6748,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7698,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6749,f6136])).
fof(f6749,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1407,f6136,f1279])).
fof(f7697,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6750,f6136])).
fof(f6750,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f1407,f6136,f1279])).
fof(f7696,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6751,f6136])).
fof(f6751,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f1407,f6136,f1279])).
fof(f7695,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6753,f6136])).
fof(f6753,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f1407,f6136,f1279])).
fof(f7694,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6754,f6136])).
fof(f6754,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f1407,f6136,f1279])).
fof(f7693,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6759,f6136])).
fof(f6759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f1279])).
fof(f7692,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6760,f6136])).
fof(f6760,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1649,f6136,f6136,f1279])).
fof(f7691,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6761,f6136])).
fof(f6761,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4153,f6136,f6136,f1279])).
fof(f7690,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6762,f6136])).
fof(f6762,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f6136,f6136,f1279])).
fof(f7689,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6763,f6136])).
fof(f6763,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1650,f6136,f6136,f1279])).
fof(f7688,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6764,f6136])).
fof(f6764,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f6136,f6136,f1279])).
fof(f7687,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6766,f6136])).
fof(f6766,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1651,f6136,f6136,f1279])).
fof(f7686,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6767,f6136])).
fof(f6767,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4155,f6136,f6136,f1279])).
fof(f7685,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6772,f1445])).
fof(f6772,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f2200,f6136,f1283])).
fof(f1283,plain,(
  ( ! [X2,X0,X1] : (sP45(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f737])).
fof(f2200,plain,(
  ~sP45('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f773,f1445,f1285])).
fof(f7684,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6773,f1445])).
fof(f6773,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f2203,f6136,f1283])).
fof(f7683,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7682,f1445])).
fof(f7682,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f6774,f4093])).
fof(f4093,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f773,f4087,f919])).
fof(f919,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f364])).
fof(f6774,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4105,f6136,f1283])).
fof(f7681,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6785,f1445])).
fof(f6785,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4174,f6136,f1283])).
fof(f7680,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6786,f1445])).
fof(f6786,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4176,f6136,f1283])).
fof(f4176,plain,(
  ~sP45('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1445,f4087,f1285])).
fof(f7679,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6793,f869])).
fof(f6793,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2174,f6136,f1286])).
fof(f2174,plain,(
  ( ! [X0] : (sP45('T_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f773,f869,f1284])).
fof(f1284,plain,(
  ( ! [X2,X0,X1] : (sP45(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f737])).
fof(f7678,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6794,f958])).
fof(f6794,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2175,f6136,f1286])).
fof(f2175,plain,(
  ( ! [X0,X1] : (sP45('T_0',X0,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f773,f958,f1284])).
fof(f7677,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6795,f959])).
fof(f6795,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2176,f6136,f1286])).
fof(f2176,plain,(
  ( ! [X0,X1] : (sP45('T_0',X0,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f773,f959,f1284])).
fof(f7676,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6796,f780])).
fof(f6796,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2172,f6136,f1286])).
fof(f2172,plain,(
  ( ! [X0] : (sP45('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f773,f780,f1284])).
fof(f7675,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6798,f4093])).
fof(f6798,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4113,f6136,f1286])).
fof(f4113,plain,(
  sP45('T_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1418,f1411,f1414,f4087,f1268])).
fof(f1268,plain,(
  ( ! [X2,X0,X1] : (~sP43(X1,X2,X0) | ~p(s(bool,X2)) | ~sP44(X0,X2,X1) | sP45(X0,X1,X2) | ~sP42(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f728])).
fof(f7674,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6800,f869])).
fof(f6800,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2169,f6136,f1286])).
fof(f2169,plain,(
  ( ! [X0] : (sP45('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f773,f869,f1284])).
fof(f7673,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6801,f780])).
fof(f6801,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2167,f6136,f1286])).
fof(f2167,plain,(
  ( ! [X0] : (sP45('T_0','const.bool.BOUNDED_1'(s(bool,X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f780,f1284])).
fof(f7672,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6802,f959])).
fof(f6802,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2171,f6136,f1286])).
fof(f2171,plain,(
  ( ! [X0,X1] : (sP45('T_0','const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f773,f959,f1284])).
fof(f7671,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6803,f958])).
fof(f6803,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2170,f6136,f1286])).
fof(f2170,plain,(
  ( ! [X0,X1] : (sP45('T_0','const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f773,f958,f1284])).
fof(f7670,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6836,f7392])).
fof(f7392,plain,(
  ( ! [X0,X1] : (sP47('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1420,f6136,f1392])).
fof(f1392,plain,(
  ( ! [X2,X0,X1] : (sP47(X0,X1,X2) | p(s(bool,X0)) | ~sP46(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1293])).
fof(f1293,plain,(
  ( ! [X2,X0,X1] : (sP47(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP46(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X0,X1,X2] : ((~sP47(X0,X1,X2) | (sP46(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP46(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP47(X0,X1,X2)))),
  inference(rectify,[],[f739])).
fof(f739,plain,(
  ! [X2,X0,X1] : ((~sP47(X2,X0,X1) | (sP46(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP46(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP47(X2,X0,X1)))),
  inference(flattening,[],[f738])).
fof(f738,plain,(
  ! [X2,X0,X1] : ((~sP47(X2,X0,X1) | (sP46(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP46(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP47(X2,X0,X1)))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X2,X0,X1] : (sP47(X2,X0,X1) <=> (sP46(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP47])])).
fof(f1420,plain,(
  ( ! [X0,X1] : (sP46(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1295])).
fof(f1295,plain,(
  ( ! [X2,X0,X1] : (sP46(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f742])).
fof(f742,plain,(
  ! [X0,X1,X2] : ((~sP46(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP46(X0,X1,X2)))),
  inference(flattening,[],[f741])).
fof(f741,plain,(
  ! [X0,X1,X2] : ((~sP46(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP46(X0,X1,X2)))),
  inference(nnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2] : (sP46(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP46])])).
fof(f6836,plain,(
  ( ! [X0,X1] : (~sP47('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f1288])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (~sP47(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f740])).
fof(f7669,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6844,f7392])).
fof(f6844,plain,(
  ( ! [X0,X1] : (~sP47('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f1289])).
fof(f1289,plain,(
  ( ! [X2,X0,X1] : (~sP47(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f740])).
fof(f7668,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6854,f6858])).
fof(f6858,plain,(
  ( ! [X2,X0,X1] : (sP46(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1297])).
fof(f1297,plain,(
  ( ! [X2,X0,X1] : (sP46(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f742])).
fof(f6854,plain,(
  ( ! [X0] : (~sP46('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f773,f6136,f1294])).
fof(f1294,plain,(
  ( ! [X2,X0,X1] : (~sP46(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f742])).
fof(f7667,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6862,f1407])).
fof(f6862,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2347,f6136,f1298])).
fof(f1298,plain,(
  ( ! [X2,X0,X1] : (sP47(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP47(X2,X0,X1)) & (~sP47(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f743])).
fof(f743,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP47(X2,X0,X1)) & (~sP47(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP47(X2,X0,X1))),
  inference(definition_folding,[],[f289,f466,f465])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X89,X90,X91] : ((p(s(bool,X91)) <=> (p(s(bool,X90)) & p(s(bool,X89)))) <=> ((p(s(bool,X91)) | ~p(s(bool,X90)) | ~p(s(bool,X89))) & (p(s(bool,X90)) | ~p(s(bool,X91))) & (p(s(bool,X89)) | ~p(s(bool,X91)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sat.sat_23_prove_ltb')).
fof(f2347,plain,(
  ~sP47('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2339,f1287])).
fof(f1287,plain,(
  ( ! [X2,X0,X1] : (~sP47(X0,X1,X2) | sP46(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f740])).
fof(f2339,plain,(
  ~sP46('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f773,f1407,f773,f1294])).
fof(f7666,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6863,f1407])).
fof(f6863,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4221,f6136,f1298])).
fof(f4221,plain,(
  ~sP47('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f773,f4087,f1303])).
fof(f1303,plain,(
  ( ! [X2,X0,X1] : (~sP47(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f744])).
fof(f7665,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6864,f1407])).
fof(f6864,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4225,f6136,f1298])).
fof(f4225,plain,(
  ~sP47('F_0','T_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f773,f4087,f1303])).
fof(f7664,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6865,f6136])).
fof(f6865,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1652,f6136,f1298])).
fof(f1652,plain,(
  ( ! [X0,X1] : (~sP47('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f1288])).
fof(f7663,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6866,f6136])).
fof(f6866,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1653,f6136,f1298])).
fof(f1653,plain,(
  ( ! [X0,X1] : (~sP47('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f1289])).
fof(f7662,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6867,f6136])).
fof(f6867,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4178,f6136,f1298])).
fof(f4178,plain,(
  ( ! [X0] : (~sP47('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f1288])).
fof(f7661,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6868,f6136])).
fof(f6868,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4179,f6136,f1298])).
fof(f4179,plain,(
  ( ! [X0] : (~sP47('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f1289])).
fof(f7660,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6873,f1407])).
fof(f6873,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f1298])).
fof(f1501,plain,(
  ( ! [X0] : (~sP47('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1407,f1288])).
fof(f7659,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6874,f6136])).
fof(f6874,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f6136,f1298])).
fof(f1502,plain,(
  ( ! [X0] : (~sP47('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f1407,f1289])).
fof(f7658,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6877,f6136])).
fof(f6877,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4225,f6136,f1298])).
fof(f7657,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6878,f1407])).
fof(f6878,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1652,f6136,f1298])).
fof(f7656,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6879,f6136])).
fof(f6879,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1653,f6136,f1298])).
fof(f7655,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6880,f1407])).
fof(f6880,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4178,f6136,f1298])).
fof(f7654,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6881,f6136])).
fof(f6881,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4179,f6136,f1298])).
fof(f7653,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6888,f1407])).
fof(f6888,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2347,f6136,f1299])).
fof(f1299,plain,(
  ( ! [X2,X0,X1] : (sP47(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f744])).
fof(f7652,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6889,f1407])).
fof(f6889,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4221,f6136,f1299])).
fof(f7651,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6890,f1407])).
fof(f6890,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4225,f6136,f1299])).
fof(f7650,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6891,f6136])).
fof(f6891,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1652,f6136,f1299])).
fof(f7649,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6892,f6136])).
fof(f6892,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1653,f6136,f1299])).
fof(f7648,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6893,f6136])).
fof(f6893,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4178,f6136,f1299])).
fof(f7647,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6894,f6136])).
fof(f6894,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4179,f6136,f1299])).
fof(f7646,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6899,f6136])).
fof(f6899,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f1299])).
fof(f7645,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6900,f1407])).
fof(f6900,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f6136,f1299])).
fof(f7644,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6902,f6136])).
fof(f6902,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4221,f6136,f1299])).
fof(f7643,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6904,f6136])).
fof(f6904,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1652,f6136,f1299])).
fof(f7642,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6905,f1407])).
fof(f6905,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1653,f6136,f1299])).
fof(f7641,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6906,f6136])).
fof(f6906,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4178,f6136,f1299])).
fof(f7640,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6907,f1407])).
fof(f6907,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4179,f6136,f1299])).
fof(f7639,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6931,f7229])).
fof(f7229,plain,(
  ( ! [X2,X0,X1] : (sP49(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f6136,f6136,f1315])).
fof(f1315,plain,(
  ( ! [X2,X0,X1] : (sP49(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP49(X2,X0,X1)) & (~sP49(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f751])).
fof(f751,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP49(X2,X0,X1)) & (~sP49(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP49(X2,X0,X1))),
  inference(definition_folding,[],[f291,f469,f468])).
fof(f468,plain,(
  ! [X2,X0,X1] : (sP48(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP48])])).
fof(f469,plain,(
  ! [X2,X0,X1] : (sP49(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP48(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP49])])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X89,X90,X91] : ((p(s(bool,X91)) <=> (p(s(bool,X90)) | p(s(bool,X89)))) <=> ((p(s(bool,X91)) | ~p(s(bool,X90))) & (p(s(bool,X91)) | ~p(s(bool,X89))) & (p(s(bool,X90)) | p(s(bool,X89)) | ~p(s(bool,X91)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sat.sat_25_prove_ltb')).
fof(f6931,plain,(
  ( ! [X0,X1] : (~sP49(X0,X1,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f1304])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (~sP49(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f747])).
fof(f747,plain,(
  ! [X0,X1,X2] : ((~sP49(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP48(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP48(X0,X1,X2) | sP49(X0,X1,X2)))),
  inference(rectify,[],[f746])).
fof(f746,plain,(
  ! [X2,X0,X1] : ((~sP49(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP48(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP48(X2,X0,X1) | sP49(X2,X0,X1)))),
  inference(flattening,[],[f745])).
fof(f745,plain,(
  ! [X2,X0,X1] : ((~sP49(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP48(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP48(X2,X0,X1)) | sP49(X2,X0,X1)))),
  inference(nnf_transformation,[],[f469])).
fof(f7638,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6942,f7229])).
fof(f6942,plain,(
  ( ! [X0,X1] : (~sP49(X0,'T_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f1305])).
fof(f1305,plain,(
  ( ! [X2,X0,X1] : (~sP49(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f747])).
fof(f7637,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6984,f6136])).
fof(f6984,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1655,f6136,f1310])).
fof(f1310,plain,(
  ( ! [X2,X0,X1] : (sP49(X0,X1,X2) | ~sP48(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f747])).
fof(f1655,plain,(
  ( ! [X0,X1] : (~sP49('F_0',X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f1304])).
fof(f1425,plain,(
  ( ! [X0,X1] : (sP48('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1314])).
fof(f1314,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0,X1,X2] : ((~sP48(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP48(X0,X1,X2)))),
  inference(rectify,[],[f749])).
fof(f749,plain,(
  ! [X2,X0,X1] : ((~sP48(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP48(X2,X0,X1)))),
  inference(flattening,[],[f748])).
fof(f748,plain,(
  ! [X2,X0,X1] : ((~sP48(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP48(X2,X0,X1)))),
  inference(nnf_transformation,[],[f468])).
fof(f7636,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6985,f6136])).
fof(f6985,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1663,f1655,f6136,f1310])).
fof(f1663,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,X1,'const.ind_type.ISO_2'(s(fun(X2,X2),'_i_0_0'),s(fun(X2,X2),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f781,f1312])).
fof(f1312,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f750])).
fof(f7635,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6986,f6136])).
fof(f6986,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f1655,f6136,f1310])).
fof(f1424,plain,(
  ( ! [X0,X1] : (sP48(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f773,f1313])).
fof(f1313,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f750])).
fof(f7634,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6987,f6136])).
fof(f6987,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1664,f1655,f6136,f1310])).
fof(f1664,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f781,f1313])).
fof(f7633,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6988,f6136])).
fof(f6988,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1425,f4229,f6136,f1310])).
fof(f4229,plain,(
  ( ! [X0] : (~sP49('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f1304])).
fof(f7632,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6989,f6136])).
fof(f6989,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1516,f6136,f1310])).
fof(f1516,plain,(
  ( ! [X0] : (~sP49('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f1407,f1305])).
fof(f7631,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6990,f6136])).
fof(f6990,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f1516,f6136,f1310])).
fof(f7630,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6992,f6136])).
fof(f6992,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1663,f1516,f6136,f1310])).
fof(f7629,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6993,f6136])).
fof(f6993,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4248,f1516,f6136,f1310])).
fof(f4248,plain,(
  ( ! [X0,X1] : (sP48(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1312])).
fof(f7628,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6994,f6136])).
fof(f6994,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1659,f6136,f1310])).
fof(f1659,plain,(
  ( ! [X0,X1] : (~sP49('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f1305])).
fof(f7627,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6995,f6136])).
fof(f6995,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1664,f1659,f6136,f1310])).
fof(f7626,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6997,f6136])).
fof(f6997,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1663,f1659,f6136,f1310])).
fof(f7625,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6998,f6136])).
fof(f6998,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4248,f1659,f6136,f1310])).
fof(f7624,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6999,f6136])).
fof(f6999,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1425,f4234,f6136,f1310])).
fof(f4234,plain,(
  ( ! [X0] : (~sP49('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1407,f4087,f1305])).
fof(f7623,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7001,f6136])).
fof(f7001,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1663,f4234,f6136,f1310])).
fof(f7622,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7002,f6136])).
fof(f7002,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4248,f4234,f6136,f1310])).
fof(f7621,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7007,f7135])).
fof(f7135,plain,(
  ( ! [X2,X0,X1] : (sP48(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1314])).
fof(f7007,plain,(
  ~sP48('F_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1407,f6136,f1310])).
fof(f7620,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7009,f7135])).
fof(f7009,plain,(
  ( ! [X0] : (~sP48('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1407,f6136,f1310])).
fof(f7619,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7011,f7135])).
fof(f7011,plain,(
  ~sP48('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),'F_0') | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1407,f6136,f1310])).
fof(f7618,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7022,f7135])).
fof(f7022,plain,(
  ( ! [X0] : (~sP48('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1310])).
fof(f1515,plain,(
  ( ! [X0] : (~sP49('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1407,f1304])).
fof(f7617,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7024,f7135])).
fof(f7024,plain,(
  ( ! [X0,X1] : (~sP48('F_0',X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1310])).
fof(f7616,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7026,f7135])).
fof(f7026,plain,(
  ( ! [X0] : (~sP48('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1310])).
fof(f7615,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7028,f7135])).
fof(f7028,plain,(
  ( ! [X0] : (~sP48('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1310])).
fof(f7614,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7030,f7135])).
fof(f7030,plain,(
  ( ! [X0,X1] : (~sP48('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1310])).
fof(f7613,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7032,f7135])).
fof(f7032,plain,(
  ( ! [X0] : (~sP48('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1310])).
fof(f7612,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7035,f6136])).
fof(f7035,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1423,f6136,f1310])).
fof(f1423,plain,(
  ( ! [X0,X1] : (sP48(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f1312])).
fof(f7611,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7037,f6136])).
fof(f7037,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1423,f6136,f1310])).
fof(f7610,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7038,f6136])).
fof(f7038,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1423,f6136,f1310])).
fof(f7609,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7039,f6136])).
fof(f7039,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1663,f6136,f1310])).
fof(f7608,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7041,f6136])).
fof(f7041,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1663,f6136,f1310])).
fof(f7607,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7042,f6136])).
fof(f7042,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1663,f6136,f1310])).
fof(f7606,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7043,f6136])).
fof(f7043,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f4248,f6136,f1310])).
fof(f7605,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7045,f6136])).
fof(f7045,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1659,f4248,f6136,f1310])).
fof(f7604,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7046,f6136])).
fof(f7046,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f4248,f6136,f1310])).
fof(f7603,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7048,f6136])).
fof(f7048,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1664,f6136,f1310])).
fof(f7602,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7049,f6136])).
fof(f7049,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1664,f6136,f1310])).
fof(f7601,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7050,f6136])).
fof(f7050,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1664,f6136,f1310])).
fof(f7600,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7051,f6136])).
fof(f7051,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1664,f6136,f1310])).
fof(f7599,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7052,f6136])).
fof(f7052,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1425,f6136,f1310])).
fof(f7598,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7053,f6136])).
fof(f7053,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1425,f6136,f1310])).
fof(f7597,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7054,f6136])).
fof(f7054,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1425,f6136,f1310])).
fof(f7596,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7056,f6136])).
fof(f7056,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1425,f6136,f1310])).
fof(f7595,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7057,f6136])).
fof(f7057,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1425,f6136,f1310])).
fof(f7594,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7058,f7135])).
fof(f7058,plain,(
  ~sP48('F_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f6136,f1310])).
fof(f7593,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7059,f7135])).
fof(f7059,plain,(
  ( ! [X0] : (~sP48('F_0','F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1407,f6136,f1310])).
fof(f7592,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7060,f7135])).
fof(f7060,plain,(
  ~sP48('F_0','F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1407,f6136,f1310])).
fof(f7591,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7071,f7135])).
fof(f7071,plain,(
  ( ! [X0] : (~sP48('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1310])).
fof(f7590,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7072,f7135])).
fof(f7072,plain,(
  ( ! [X0,X1] : (~sP48('F_0',X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1310])).
fof(f7589,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7073,f7135])).
fof(f7073,plain,(
  ( ! [X0] : (~sP48('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1310])).
fof(f7588,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7074,f7135])).
fof(f7074,plain,(
  ( ! [X0] : (~sP48('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1310])).
fof(f7587,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7075,f7135])).
fof(f7075,plain,(
  ( ! [X0,X1] : (~sP48('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1310])).
fof(f7586,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7076,f7135])).
fof(f7076,plain,(
  ( ! [X0] : (~sP48('F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1310])).
fof(f7585,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7093,f6136])).
fof(f7093,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1663,f773,f6136,f1311])).
fof(f1311,plain,(
  ( ! [X2,X0,X1] : (~sP48(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f750])).
fof(f7584,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7094,f6136])).
fof(f7094,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4248,f773,f6136,f1311])).
fof(f7583,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7095,f6136])).
fof(f7095,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f773,f6136,f1311])).
fof(f7582,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7096,f6136])).
fof(f7096,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1664,f773,f6136,f1311])).
fof(f7581,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7101,f7135])).
fof(f7101,plain,(
  ( ! [X0] : (~sP48('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f1407,f6136,f1311])).
fof(f7580,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7111,f7135])).
fof(f7111,plain,(
  ( ! [X0,X1] : (~sP48('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f6136,f1311])).
fof(f7579,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7113,f6136])).
fof(f7113,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f773,f6136,f1311])).
fof(f7578,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7114,f6136])).
fof(f7114,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1663,f773,f6136,f1311])).
fof(f7577,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7115,f6136])).
fof(f7115,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4248,f773,f6136,f1311])).
fof(f7576,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7117,f6136])).
fof(f7117,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1664,f773,f6136,f1311])).
fof(f7575,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7123,f7135])).
fof(f7123,plain,(
  ( ! [X0] : (~sP48('T_0','F_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f1407,f6136,f1311])).
fof(f7574,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7133,f7135])).
fof(f7133,plain,(
  ( ! [X0,X1] : (~sP48('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136,f6136,f1311])).
fof(f7573,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7138,f6136])).
fof(f7138,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1407,f6136,f1315])).
fof(f7572,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7139,f6136])).
fof(f7139,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1407,f6136,f1315])).
fof(f7571,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7140,f6136])).
fof(f7140,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1407,f6136,f1315])).
fof(f7570,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7141,f6136])).
fof(f7141,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1407,f6136,f1315])).
fof(f7569,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7142,f6136])).
fof(f7142,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1407,f6136,f1315])).
fof(f7568,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7144,f6136])).
fof(f7144,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7567,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7145,f6136])).
fof(f7145,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7566,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7146,f6136])).
fof(f7146,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1315])).
fof(f7565,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7147,f6136])).
fof(f7147,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7564,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7148,f6136])).
fof(f7148,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7563,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7149,f1407])).
fof(f7149,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f6136,f6136,f1315])).
fof(f2480,plain,(
  ~sP49('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2473,f1306])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (~sP49(X0,X1,X2) | sP48(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f747])).
fof(f2473,plain,(
  ~sP48('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f773,f1407,f1407,f1311])).
fof(f7562,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7150,f1407])).
fof(f7150,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1407,f6136,f1315])).
fof(f7561,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7151,f1407])).
fof(f7151,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1407,f6136,f1315])).
fof(f7560,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7152,f1407])).
fof(f7152,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1407,f6136,f1315])).
fof(f7559,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7158,f1407])).
fof(f7158,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1315])).
fof(f7558,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7159,f1407])).
fof(f7159,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7557,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7160,f1407])).
fof(f7160,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7556,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7161,f1407])).
fof(f7161,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1315])).
fof(f7555,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7162,f1407])).
fof(f7162,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7554,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7163,f1407])).
fof(f7163,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7553,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7169,f6136])).
fof(f7169,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f6136,f1315])).
fof(f7552,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7170,f6136])).
fof(f7170,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1407,f6136,f1315])).
fof(f7551,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7171,f6136])).
fof(f7171,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1407,f6136,f1315])).
fof(f7550,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7173,f6136])).
fof(f7173,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1407,f6136,f1315])).
fof(f7549,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7174,f6136])).
fof(f7174,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1407,f6136,f1315])).
fof(f7548,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7175,f6136])).
fof(f7175,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1315])).
fof(f7547,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7176,f6136])).
fof(f7176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7546,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7177,f6136])).
fof(f7177,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7545,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7179,f6136])).
fof(f7179,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7544,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7180,f6136])).
fof(f7180,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7543,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7181,f1407])).
fof(f7181,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f6136,f6136,f1315])).
fof(f7542,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7182,f1407])).
fof(f7182,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f6136,f1315])).
fof(f7541,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7183,f1407])).
fof(f7183,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1407,f6136,f1315])).
fof(f7540,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7184,f1407])).
fof(f7184,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1407,f6136,f1315])).
fof(f7539,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7190,f1407])).
fof(f7190,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1315])).
fof(f7538,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7191,f1407])).
fof(f7191,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7537,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7192,f1407])).
fof(f7192,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7536,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7193,f1407])).
fof(f7193,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1315])).
fof(f7535,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7194,f1407])).
fof(f7194,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7534,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7195,f1407])).
fof(f7195,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7533,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7202,f6136])).
fof(f7202,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1407,f6136,f1315])).
fof(f7532,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7203,f6136])).
fof(f7203,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f1407,f6136,f1315])).
fof(f7531,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7204,f1407])).
fof(f7204,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f1407,f6136,f1315])).
fof(f7530,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7205,f6136])).
fof(f7205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f1315])).
fof(f7529,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7206,f6136])).
fof(f7206,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7528,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7207,f6136])).
fof(f7207,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7527,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7209,f6136])).
fof(f7209,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7526,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7210,f6136])).
fof(f7210,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7525,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7211,f1407])).
fof(f7211,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f6136,f6136,f1315])).
fof(f7524,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7213,f6136])).
fof(f7213,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1407,f6136,f1315])).
fof(f7523,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7214,f6136])).
fof(f7214,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f1407,f6136,f1315])).
fof(f7522,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7215,f1407])).
fof(f7215,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f1407,f6136,f1315])).
fof(f7521,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7221,f6136])).
fof(f7221,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f6136,f6136,f1315])).
fof(f7520,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7222,f6136])).
fof(f7222,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f6136,f6136,f1315])).
fof(f7519,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7223,f6136])).
fof(f7223,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f6136,f1315])).
fof(f7518,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7224,f6136])).
fof(f7224,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f6136,f6136,f1315])).
fof(f7517,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7225,f6136])).
fof(f7225,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4234,f6136,f6136,f1315])).
fof(f7516,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7226,f1407])).
fof(f7226,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f6136,f6136,f1315])).
fof(f7515,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7232,f6136])).
fof(f7232,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1665,f773,f6136,f1318])).
fof(f1318,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP49(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f752])).
fof(f1665,plain,(
  ( ! [X0,X1] : (sP49('T_0',X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f773,f781,f1316])).
fof(f1316,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP49(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f752])).
fof(f7514,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7233,f6136])).
fof(f7233,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1530,f773,f6136,f1318])).
fof(f1530,plain,(
  ( ! [X0] : (sP49('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f773,f1317])).
fof(f1317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP49(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f752])).
fof(f7513,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7234,f6136])).
fof(f7234,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1670,f773,f6136,f1318])).
fof(f1670,plain,(
  ( ! [X0,X1] : (sP49('T_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f773,f781,f1317])).
fof(f7512,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7267,f6136])).
fof(f7267,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1519,f773,f6136,f1318])).
fof(f1519,plain,(
  ( ! [X0] : (sP49('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f773,f1316])).
fof(f7511,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7268,f6136])).
fof(f7268,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1665,f773,f6136,f1318])).
fof(f7510,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7270,f6136])).
fof(f7270,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1670,f773,f6136,f1318])).
fof(f7509,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7275,f1407])).
fof(f7275,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1553,f6136,f1321])).
fof(f1321,plain,(
  ( ! [X2,X0,X1] : (~sP51(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f755])).
fof(f755,plain,(
  ! [X0,X1,X2] : ((~sP51(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP50(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP50(X0,X1,X2) | sP51(X0,X1,X2)))),
  inference(rectify,[],[f754])).
fof(f754,plain,(
  ! [X2,X0,X1] : ((~sP51(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP50(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP50(X2,X0,X1) | sP51(X2,X0,X1)))),
  inference(flattening,[],[f753])).
fof(f753,plain,(
  ! [X2,X0,X1] : ((~sP51(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP50(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP50(X2,X0,X1)) | sP51(X2,X0,X1)))),
  inference(nnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X2,X0,X1] : (sP51(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP50(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP51])])).
fof(f1553,plain,(
  ( ! [X0] : (sP51('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f773,f1333])).
fof(f1333,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP51(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f760])).
fof(f760,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP51(X2,X0,X1)) & (~sP51(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f759])).
fof(f759,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP51(X2,X0,X1)) & (~sP51(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP51(X2,X0,X1))),
  inference(definition_folding,[],[f391,f472,f471])).
fof(f471,plain,(
  ! [X2,X0,X1] : (sP50(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP50])])).
fof(f391,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X89,X90,X91] : ((p(s(bool,X91)) <=> (p(s(bool,X90)) => p(s(bool,X89)))) <=> ((p(s(bool,X91)) | p(s(bool,X90))) & (p(s(bool,X91)) | ~p(s(bool,X89))) & (~p(s(bool,X90)) | p(s(bool,X89)) | ~p(s(bool,X91)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.sat.sat_27_prove_ltb')).
fof(f7508,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7276,f6136])).
fof(f7276,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1564,f6136,f1321])).
fof(f1564,plain,(
  ( ! [X0] : (sP51('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f773,f1334])).
fof(f1334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP51(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f760])).
fof(f7507,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7277,f6136])).
fof(f7277,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1681,f6136,f1321])).
fof(f1681,plain,(
  ( ! [X0,X1] : (sP51('T_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f773,f781,f1334])).
fof(f7506,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7279,f1407])).
fof(f7279,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1605,f6136,f1321])).
fof(f1605,plain,(
  ( ! [X0,X1] : (sP51('const.bool.BOUNDED_1'(s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f1555])).
fof(f1555,plain,(
  ( ! [X2,X0,X1] : (sP51('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1554,f773])).
fof(f1554,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP51('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1)) )),
  inference(superposition,[],[f1333,f780])).
fof(f7505,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7280,f6136])).
fof(f7280,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1623,f6136,f1321])).
fof(f1623,plain,(
  ( ! [X0,X1] : (sP51('const.bool.BOUNDED_1'(s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f773,f1566])).
fof(f1566,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP51('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(subsumption_resolution,[],[f1565,f773])).
fof(f1565,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP51('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(superposition,[],[f1334,f780])).
fof(f7504,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7281,f6136])).
fof(f7281,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1625,f6136,f1321])).
fof(f1625,plain,(
  ( ! [X2,X0,X1] : (sP51('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(subsumption_resolution,[],[f1624,f773])).
fof(f1624,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP51('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(superposition,[],[f1566,f780])).
fof(f7503,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7282,f1407])).
fof(f7282,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1686,f6136,f1321])).
fof(f1686,plain,(
  ( ! [X0,X1] : (sP51('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1426,f781,f1390])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP50(X0,X1,X2) | sP51(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1324])).
fof(f1324,plain,(
  ( ! [X2,X0,X1] : (sP51(X0,X1,X2) | ~sP50(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f755])).
fof(f1426,plain,(
  ( ! [X0,X1] : (sP50(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1407,f1329])).
fof(f1329,plain,(
  ( ! [X2,X0,X1] : (sP50(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f758])).
fof(f758,plain,(
  ! [X0,X1,X2] : ((~sP50(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP50(X0,X1,X2)))),
  inference(rectify,[],[f757])).
fof(f757,plain,(
  ! [X2,X0,X1] : ((~sP50(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP50(X2,X0,X1)))),
  inference(flattening,[],[f756])).
fof(f756,plain,(
  ! [X2,X0,X1] : ((~sP50(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP50(X2,X0,X1)))),
  inference(nnf_transformation,[],[f471])).
fof(f7502,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7283,f6136])).
fof(f7283,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1688,f6136,f1321])).
fof(f1688,plain,(
  ( ! [X0,X1] : (sP51('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1427,f781,f1390])).
fof(f1427,plain,(
  ( ! [X0,X1] : (sP50(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f773,f1330])).
fof(f1330,plain,(
  ( ! [X2,X0,X1] : (sP50(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f758])).
fof(f7501,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7291,f1407])).
fof(f7291,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2584,f6136,f1321])).
fof(f2584,plain,(
  sP51('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1428,f773,f1407,f1327])).
fof(f1327,plain,(
  ( ! [X2,X0,X1] : (sP51(X0,X1,X2) | ~sP50(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f755])).
fof(f1428,plain,(
  ( ! [X0,X1] : (sP50('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1331])).
fof(f1331,plain,(
  ( ! [X2,X0,X1] : (sP50(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f758])).
fof(f7500,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7292,f6136])).
fof(f7292,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1605,f6136,f1321])).
fof(f7499,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7293,f6136])).
fof(f7293,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1623,f6136,f1321])).
fof(f7498,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7294,f6136])).
fof(f7294,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1625,f6136,f1321])).
fof(f7497,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7295,f6136])).
fof(f7295,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1686,f6136,f1321])).
fof(f7496,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7296,f6136])).
fof(f7296,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1688,f6136,f1321])).
fof(f7495,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7330,f7346])).
fof(f7346,plain,(
  ( ! [X2,X0,X1] : (sP50(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1331])).
fof(f7330,plain,(
  ~sP50('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1548,f773,f6136,f1327])).
fof(f1548,plain,(
  ( ! [X0] : (~sP51('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f1407,f1322])).
fof(f1322,plain,(
  ( ! [X2,X0,X1] : (~sP51(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f755])).
fof(f7494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7332,f7346])).
fof(f7332,plain,(
  ( ! [X0] : (~sP50('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1675,f773,f6136,f1327])).
fof(f1675,plain,(
  ( ! [X0,X1] : (~sP51('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f781,f1322])).
fof(f7493,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7342,f7346])).
fof(f7342,plain,(
  ( ! [X0] : (~sP50('T_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f773,f773,f6136,f1328])).
fof(f1328,plain,(
  ( ! [X2,X0,X1] : (~sP50(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f758])).
fof(f7492,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7355,f1407])).
fof(f7355,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1548,f773,f6136,f1332])).
fof(f1332,plain,(
  ( ! [X2,X0,X1] : (sP51(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f760])).
fof(f7491,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7356,f1407])).
fof(f7356,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1675,f773,f6136,f1332])).
fof(f7490,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7363,f6136])).
fof(f7363,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1675,f773,f6136,f1332])).
fof(f7489,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7364,f1407])).
fof(f7364,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2601,f773,f6136,f1332])).
fof(f2601,plain,(
  ~sP51('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2593,f1323])).
fof(f1323,plain,(
  ( ! [X2,X0,X1] : (~sP51(X0,X1,X2) | sP50(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f755])).
fof(f2593,plain,(
  ~sP50('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f773,f773,f1407,f1328])).
fof(f7488,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7397,f1420])).
fof(f7397,plain,(
  ( ! [X0] : (~sP46(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f1392])).
fof(f7487,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7399,f1420])).
fof(f7399,plain,(
  ( ! [X0] : (~sP46('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f6136,f1392])).
fof(f7486,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7404,f1654])).
fof(f1654,plain,(
  ( ! [X2,X0,X1] : (sP46(X0,X1,'const.ind_type.ISO_2'(s(fun(X2,X2),'_i_0_0'),s(fun(X2,X2),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f781,f1295])).
fof(f7404,plain,(
  ( ! [X0,X1] : (~sP46(X0,'F_0','const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1652,f6136,f1392])).
fof(f7485,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7406,f1654])).
fof(f7406,plain,(
  ( ! [X0,X1] : (~sP46('F_0',X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1653,f6136,f1392])).
fof(f7484,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7408,f4206])).
fof(f4206,plain,(
  ( ! [X0,X1] : (sP46(X0,X1,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f1295])).
fof(f7408,plain,(
  ( ! [X0] : (~sP46(X0,'F_0','|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4178,f6136,f1392])).
fof(f7483,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7410,f4206])).
fof(f7410,plain,(
  ( ! [X0] : (~sP46('F_0',X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4179,f6136,f1392])).
fof(f7482,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7413,f6136])).
fof(f7413,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1433,f6136,f1408])).
fof(f1408,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f930,f1407])).
fof(f930,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1433,plain,(
  ( ! [X2,X3] : (sP1(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(subsumption_resolution,[],[f1431,f773])).
fof(f1431,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP1(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(superposition,[],[f931,f780])).
fof(f931,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f556])).
fof(f7481,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7414,f6136])).
fof(f7414,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1639,f6136,f1408])).
fof(f1639,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f781,f931])).
fof(f7480,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7415,f6136])).
fof(f7415,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4095,f6136,f1408])).
fof(f4095,plain,(
  ( ! [X0] : (sP1(X0,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f931])).
fof(f7479,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7416,f6136])).
fof(f7416,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1410,f6136,f1408])).
fof(f1410,plain,(
  ( ! [X0] : (sP1('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f773,f932])).
fof(f932,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f556])).
fof(f7478,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7417,f6136])).
fof(f7417,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1432,f6136,f1408])).
fof(f1432,plain,(
  ( ! [X0,X1] : (sP1('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(subsumption_resolution,[],[f1430,f773])).
fof(f1430,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP1('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f932,f780])).
fof(f7477,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7418,f6136])).
fof(f7418,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1640,f6136,f1408])).
fof(f1640,plain,(
  ( ! [X0,X1] : (sP1('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f781,f932])).
fof(f7476,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7419,f6136])).
fof(f7419,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4096,f6136,f1408])).
fof(f4096,plain,(
  ( ! [X0] : (sP1('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4087,f932])).
fof(f7475,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7424,f6136])).
fof(f7424,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f6136,f1408])).
fof(f1409,plain,(
  ( ! [X0] : (sP1(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f773,f931])).
fof(f7474,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7425,f6136])).
fof(f7425,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1433,f6136,f1408])).
fof(f7473,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7426,f6136])).
fof(f7426,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1639,f6136,f1408])).
fof(f7472,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7427,f6136])).
fof(f7427,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4095,f6136,f1408])).
fof(f7471,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7429,f6136])).
fof(f7429,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1432,f6136,f1408])).
fof(f7470,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7430,f6136])).
fof(f7430,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1640,f6136,f1408])).
fof(f7469,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7431,f6136])).
fof(f7431,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4096,f6136,f1408])).
fof(f7440,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f5032,f6136,f1060])).
fof(f1060,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X2,X1) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))) )),
  inference(cnf_transformation,[],[f634])).
fof(f7435,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f6136,f6136,f1408])).
fof(f7433,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f1407,f6136,f1408])).
fof(f7423,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f6136,f6136,f1408])).
fof(f7421,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1410,f1407,f6136,f1408])).
fof(f7411,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1422,f4179,f6136,f1392])).
fof(f1422,plain,(
  ( ! [X0,X1] : (sP46('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1297])).
fof(f7409,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1421,f4178,f6136,f1392])).
fof(f1421,plain,(
  ( ! [X0,X1] : (sP46(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1407,f1296])).
fof(f1296,plain,(
  ( ! [X2,X0,X1] : (sP46(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f742])).
fof(f7407,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1422,f1653,f6136,f1392])).
fof(f7405,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1421,f1652,f6136,f1392])).
fof(f7400,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1422,f1502,f6136,f1392])).
fof(f7398,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1421,f1501,f6136,f1392])).
fof(f7391,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1564,f773,f773,f6136,f1335])).
fof(f1335,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP51(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f760])).
fof(f7380,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2601,f773,f6136,f1333])).
fof(f7374,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1548,f773,f6136,f6136,f1332])).
fof(f7368,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2601,f773,f1407,f6136,f1332])).
fof(f7361,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1548,f6136,f773,f6136,f1332])).
fof(f7359,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1548,f1407,f773,f6136,f1332])).
fof(f7347,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2593,f6136,f1331])).
fof(f7345,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2593,f6136,f1329])).
fof(f7343,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1427,f773,f773,f6136,f1328])).
fof(f7333,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1428,f1675,f773,f6136,f1327])).
fof(f7331,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1428,f1548,f773,f6136,f1327])).
fof(f7311,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1564,f773,f6136,f1322])).
fof(f7300,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1553,f6136,f6136,f1321])).
fof(f7298,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2584,f1407,f6136,f1321])).
fof(f7287,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1553,f6136,f6136,f1321])).
fof(f7285,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1553,f1407,f6136,f1321])).
fof(f7274,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1519,f6136,f773,f6136,f1318])).
fof(f7272,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1519,f1407,f773,f6136,f1318])).
fof(f7252,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1519,f773,f6136,f6136,f1318])).
fof(f7241,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1530,f773,f1407,f6136,f1318])).
fof(f7230,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f6136,f1315])).
fof(f7228,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1407,f6136,f6136,f1315])).
fof(f7219,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f1407,f6136,f1315])).
fof(f7217,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2480,f1407,f1407,f6136,f1315])).
fof(f7200,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f6136,f1315])).
fof(f7198,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f6136,f6136,f1315])).
fof(f7189,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f1407,f6136,f1315])).
fof(f7187,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f1407,f6136,f1315])).
fof(f7168,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f6136,f6136,f6136,f1315])).
fof(f7166,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1407,f6136,f6136,f1315])).
fof(f7157,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f6136,f1407,f6136,f1315])).
fof(f7155,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1407,f1407,f6136,f1315])).
fof(f7136,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2473,f6136,f1314])).
fof(f7134,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f773,f6136,f6136,f1311])).
fof(f7124,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f773,f1407,f6136,f1311])).
fof(f7112,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f773,f6136,f6136,f1311])).
fof(f7102,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f773,f1407,f6136,f1311])).
fof(f7089,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1425,f6136,f6136,f1310])).
fof(f7087,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1659,f1664,f6136,f6136,f1310])).
fof(f7085,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1516,f1424,f6136,f6136,f1310])).
fof(f7083,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f4248,f6136,f6136,f1310])).
fof(f7081,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1663,f6136,f6136,f1310])).
fof(f7079,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1423,f6136,f6136,f1310])).
fof(f7069,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1425,f1407,f6136,f1310])).
fof(f7067,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4229,f4248,f1407,f6136,f1310])).
fof(f7065,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1655,f1663,f1407,f6136,f1310])).
fof(f7063,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1515,f1423,f1407,f6136,f1310])).
fof(f7033,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1425,f4234,f6136,f6136,f1310])).
fof(f7031,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1659,f6136,f6136,f1310])).
fof(f7029,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1516,f6136,f6136,f1310])).
fof(f7027,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1425,f4229,f6136,f6136,f1310])).
fof(f7025,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1655,f6136,f6136,f1310])).
fof(f7023,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1515,f6136,f6136,f1310])).
fof(f7012,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1425,f4234,f1407,f6136,f1310])).
fof(f7010,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1659,f1407,f6136,f1310])).
fof(f7008,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f1516,f1407,f6136,f1310])).
fof(f6943,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1530,f773,f6136,f1305])).
fof(f6932,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1519,f773,f6136,f1304])).
fof(f6921,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2330,f773,f773,f6136,f1303])).
fof(f2330,plain,(
  sP47('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1420,f773,f773,f1290])).
fof(f1290,plain,(
  ( ! [X2,X0,X1] : (sP47(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP46(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f740])).
fof(f6911,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f6136,f1299])).
fof(f6909,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f1407,f6136,f1299])).
fof(f6898,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f6136,f1299])).
fof(f6896,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2347,f1407,f6136,f1299])).
fof(f6885,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f6136,f1298])).
fof(f6883,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f1407,f6136,f1298])).
fof(f6872,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1501,f6136,f6136,f1298])).
fof(f6870,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2347,f1407,f6136,f1298])).
fof(f6859,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2339,f6136,f1297])).
fof(f6857,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2339,f6136,f1296])).
fof(f6855,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1420,f773,f773,f6136,f1294])).
fof(f6845,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2330,f773,f6136,f1289])).
fof(f6837,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2330,f773,f6136,f1288])).
fof(f6826,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2175,f958,f6136,f1286])).
fof(f6824,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2176,f959,f6136,f1286])).
fof(f6820,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2174,f869,f6136,f1286])).
fof(f6818,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2172,f780,f6136,f1286])).
fof(f6813,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2170,f958,f6136,f1286])).
fof(f6811,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2171,f959,f6136,f1286])).
fof(f6807,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2169,f869,f6136,f1286])).
fof(f6805,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2167,f780,f6136,f1286])).
fof(f7457,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6799])).
fof(f6799,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3468,f6136,f1286])).
fof(f3468,plain,(
  sP45('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1417,f1412,f1415,f1407,f1269])).
fof(f7456,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6797])).
fof(f6797,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3393,f6136,f1286])).
fof(f3393,plain,(
  sP45('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1419,f1411,f1414,f773,f1268])).
fof(f6790,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2203,f1445,f6136,f1283])).
fof(f6788,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2200,f1445,f6136,f1283])).
fof(f6771,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f6136,f1279])).
fof(f6769,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1407,f6136,f6136,f1279])).
fof(f6758,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f1407,f6136,f1279])).
fof(f6756,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1407,f1407,f6136,f1279])).
fof(f6733,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f6136,f1279])).
fof(f6731,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f1407,f6136,f6136,f1279])).
fof(f6720,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f6136,f1407,f6136,f1279])).
fof(f6718,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f1407,f1407,f6136,f1279])).
fof(f6695,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f6136,f6136,f1279])).
fof(f6693,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f1407,f6136,f6136,f1279])).
fof(f6682,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6136,f1407,f6136,f1279])).
fof(f6680,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f1407,f1407,f6136,f1279])).
fof(f6657,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2309,f6136,f1278])).
fof(f2309,plain,(
  ~sP43('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f773,f1407,f773,f1275])).
fof(f6655,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2309,f6136,f1277])).
fof(f6653,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1414,f773,f773,f6136,f1275])).
fof(f6643,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f6136,f1274])).
fof(f2294,plain,(
  ~sP44('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f773,f1407,f773,f1271])).
fof(f6641,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f6136,f1273])).
fof(f6639,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2177,f773,f773,f6136,f1271])).
fof(f2177,plain,(
  ( ! [X0] : (sP44('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f2167,f1266])).
fof(f1266,plain,(
  ( ! [X2,X0,X1] : (~sP45(X0,X1,X2) | sP44(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f728])).
fof(f6592,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f1415,f6136,f1270])).
fof(f6582,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f1415,f6136,f1270])).
fof(f6580,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f1415,f6136,f1270])).
fof(f6574,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f1415,f6136,f1270])).
fof(f6568,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f1415,f6136,f1270])).
fof(f6566,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f1415,f6136,f1270])).
fof(f6555,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f4152,f6136,f1270])).
fof(f6539,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f4152,f6136,f1270])).
fof(f6511,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f1414,f6136,f1270])).
fof(f6509,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f1414,f6136,f1270])).
fof(f6503,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f1414,f6136,f1270])).
fof(f6491,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f1414,f6136,f1270])).
fof(f6391,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f1415,f6136,f1269])).
fof(f6381,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f1415,f6136,f1269])).
fof(f6379,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f1415,f6136,f1269])).
fof(f6373,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f1415,f6136,f1269])).
fof(f6367,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f1415,f6136,f1269])).
fof(f6365,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f1415,f6136,f1269])).
fof(f6354,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1412,f4152,f6136,f1269])).
fof(f6338,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1418,f4174,f1411,f4152,f6136,f1269])).
fof(f6310,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f1412,f1414,f6136,f1269])).
fof(f6308,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1412,f1414,f6136,f1269])).
fof(f6302,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f4154,f4105,f4137,f1414,f6136,f1269])).
fof(f6290,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1418,f2203,f1411,f1414,f6136,f1269])).
fof(f6216,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2172,f773,f773,f6136,f1267])).
fof(f6196,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1466,f773,f6136,f950])).
fof(f1466,plain,(
  ~sP2('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f773,f773,f946])).
fof(f946,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f563])).
fof(f6188,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1466,f773,f6136,f947])).
fof(f6180,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1469,f6136,f6136,f945])).
fof(f6178,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1474,f1407,f6136,f945])).
fof(f6168,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1469,f6136,f6136,f945])).
fof(f6166,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1469,f1407,f6136,f945])).
fof(f6156,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1407,f1445,f6136,f916])).
fof(f6154,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1445,f6136,f916])).
fof(f6150,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1407,f1445,f6136,f916])).
fof(f6148,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6136,f1445,f6136,f916])).
fof(f6144,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1445,f6136,f818])).
fof(f818,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f6142,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f780,f6136,f813])).
fof(f6140,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f773,f6136])).
fof(f6138,plain,(
  p(s(bool,sK120('_i_0_0',bool))) | $spl10),
  inference(cnf_transformation,[],[f6138_D])).
fof(f6138_D,plain,(
  p(s(bool,sK120('_i_0_0',bool))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f6135,plain,(
  ( ! [X3] : (p(s(bool,sK120('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f6124,f869])).
fof(f6124,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK120('_i_0_0',bool)))))) )),
  inference(superposition,[],[f874,f869])).
fof(f874,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK120(X1,X0)))))) )),
  inference(cnf_transformation,[],[f522])).
fof(f5005,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4747,f4638])).
fof(f4638,plain,(
  ( ! [X0] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))))) ) | $spl2),
  inference(superposition,[],[f914,f1405])).
fof(f914,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X8] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X8))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X8))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.dec_to_exhProof.v_rel_rules|split|4')).
fof(f4747,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1435,f915,f1271])).
fof(f915,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7,X8] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X8))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X8))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.dec_to_exhProof.v_rel_rules|split|0')).
fof(f1435,plain,(
  ( ! [X0,X1] : (sP44(X0,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1273])).
fof(f4999,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4757,f4638])).
fof(f4757,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1436,f915,f1271])).
fof(f1436,plain,(
  ( ! [X0,X1] : (sP44('const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1274])).
fof(f4996,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4765,f4638])).
fof(f4765,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1437,f915,f1275])).
fof(f1437,plain,(
  ( ! [X0,X1] : (sP43(X0,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1277])).
fof(f4993,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4772,f4638])).
fof(f4772,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1438,f915,f1275])).
fof(f1438,plain,(
  ( ! [X0,X1] : (sP43('const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1278])).
fof(f4990,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4819,f4638])).
fof(f4819,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1439,f915,f1294])).
fof(f1439,plain,(
  ( ! [X0,X1] : (sP46(X0,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1296])).
fof(f4987,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4826,f4638])).
fof(f4826,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1440,f915,f1294])).
fof(f1440,plain,(
  ( ! [X0,X1] : (sP46('const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1297])).
fof(f4986,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4845,f4638])).
fof(f4845,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1610,f915,f1303])).
fof(f1610,plain,(
  ( ! [X0] : (sP47('F_0',X0,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1439,f1407,f1392])).
fof(f4985,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4849,f4638])).
fof(f4849,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1612,f915,f1303])).
fof(f1612,plain,(
  ( ! [X0] : (sP47('F_0','const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1440,f1407,f1392])).
fof(f4973,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4912,f4638])).
fof(f4912,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1442,f915,f1328])).
fof(f1442,plain,(
  ( ! [X0,X1] : (sP50(X0,X1,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1329])).
fof(f4968,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f4922,f4638])).
fof(f4922,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(unit_resulting_resolution,[],[f1407,f1443,f915,f1328])).
fof(f1443,plain,(
  ( ! [X0,X1] : (sP50('const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1404,f1331])).
fof(f4087,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f4087_D])).
fof(f4087_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4088,plain,(
  ( ! [X3] : (~p(s(bool,sK122(X3,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f4088_D])).
fof(f4088_D,plain,(
  ( ! [X3] : (~p(s(bool,sK122(X3,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f4078,plain,(
  ( ! [X3] : (~p(s(bool,sK122(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f883,f869])).
fof(f883,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK122(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK122(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK122(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK122])],[f360])).
fof(f360,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X45] : ! [X47,X17] : (! [X21] : (s(X45,X17) = s(X45,X21) => p(s(bool,happ(s(fun(X45,bool),X47),s(X45,X21))))) => p(s(bool,'|dollar|exists_1'(s(fun(X45,bool),X47)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.bool.JRH_INDUCT_UTIL')).
fof(f1403,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl0),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1405,plain,(
  s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) | $spl2),
  inference(cnf_transformation,[],[f1405_D])).
fof(f1405_D,plain,(
  s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f771,plain,(
  s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))))),
  inference(cnf_transformation,[],[f475])).
fof(f475,plain,(
  (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | s('type.exhSem.v',sK52) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | s('type.exhSem.v',sK52) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52,sK53,sK54])],[f474])).
fof(f474,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))),s('type.exhSem.v',X0)))) | s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X2)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))),s('type.exhSem.v',X0)))) | s('type.exhSem.v',X0) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f339])).
fof(f339,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))),s('type.exhSem.v',X0)))) <~> s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X2))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X2))),s('type.exhSem.v',X0)))) <=> s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X8))),s('type.exhSem.v',X6)))) <=> s('type.exhSem.v',X6) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X8))),s('type.exhSem.v',X6)))) <=> s('type.exhSem.v',X6) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X8))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2','thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|2')).
fof(f1404,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) | $spl1),
  inference(cnf_transformation,[],[f1404_D])).
fof(f1404_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1406,plain,(
  s('type.exhSem.v',sK52) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) | $spl3),
  inference(cnf_transformation,[],[f1406_D])).
fof(f1406_D,plain,(
  s('type.exhSem.v',sK52) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f772,plain,(
  s('type.exhSem.v',sK52) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK54))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK53),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK54))),s('type.exhSem.v',sK52))))),
  inference(cnf_transformation,[],[f475])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 828
% Passive clauses: 4074
% Generated clauses: 19525
% Final active clauses: 778
% Final passive clauses: 3051
% Input formulas: 133
% Initial clauses: 619
% 
% Duplicate literals: 7
% Trivial inequalities: 7
% Fw subsumption resolutions: 448
% Fw demodulations: 4173
% Bw demodulations: 47
% Condensations: 8
% 
% Simple tautologies: 91
% Equational tautologies: 47
% Forward subsumptions: 7943
% Fw demodulations to eq. taut.: 2170
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 8497
% Factoring: 1
% Forward superposition: 3310
% Backward superposition: 2315
% Self superposition: 30
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 4648
% SAT solver unit clauses: 1800
% SAT solver binary clauses: 1916
% 
% Sat splits: 15
% Sat splitting refutations: 542
% 
% Memory used [KB]: 21620
% Time elapsed: 0.819 s
% ------------------------------
% ------------------------------
% Success in time 2.15 s
