Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exhSem
% (129700)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 158
% Passive clauses: 3158
% Generated clauses: 21594
% Final active clauses: 127
% Final passive clauses: 2752
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 805
% Trivial inequalities: 65
% Fw subsumption resolutions: 1629
% Fw demodulations: 3296
% Bw demodulations: 24
% 
% Simple tautologies: 66
% Equational tautologies: 52
% Forward subsumptions: 11912
% 
% Binary resolution: 454
% Factoring: 2
% Forward superposition: 7072
% Backward superposition: 7743
% Self superposition: 226
% Equality factoring: 106
% Equality resolution: 1
% 
% Split clauses: 248
% Split components: 496
% SAT solver clauses: 698
% SAT solver unit clauses: 36
% SAT solver binary clauses: 124
% 
% Sat splits: 248
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 698
% TWLsolver calls for satisfiability: 140
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exhSem
% (129766)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 131
% Passive clauses: 1106
% Generated clauses: 16328
% Final active clauses: 93
% Final passive clauses: 708
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 605
% Trivial inequalities: 116
% Fw subsumption resolutions: 1020
% Fw demodulations: 1986
% Bw demodulations: 293
% Condensations: 1520
% 
% Simple tautologies: 121
% Equational tautologies: 44
% Forward subsumptions: 8089
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 427
% Factoring: 1
% Forward superposition: 4880
% Backward superposition: 4055
% Self superposition: 177
% Equality factoring: 98
% Equality resolution: 9
% 
% Split clauses: 97
% Split components: 194
% SAT solver clauses: 1722
% SAT solver unit clauses: 131
% SAT solver binary clauses: 44
% 
% Sat splits: 97
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1722
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 15351
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exhSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exhSem
% SZS output start Proof for thm.exhSem
fof(f20588,plain,(
  $false),
  inference(subsumption_resolution,[],[f20587,f216])).
fof(f216,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exhSem.state_fupdfupds_comp|split|5','thm.bool.TRUTH')).
fof(f20587,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f20585,f3908])).
fof(f3908,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2566,f300])).
fof(f300,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f219])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.exhSem.state_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f218,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.exhSem.state_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f2566,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2487,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X16,X8] : ! [X18,X20] : s(X8,X18) = s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X16),'const.pair.,_2'(s(X8,X18),s(X16,X20)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|5','thm.pair.FST')).
fof(f2487,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f259,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X43,X44,X45] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X44))),s(X43,X45)))) <=> s(X43,X44) = s(X43,X45))),
  file('thm.exhSem.state_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f20585,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.exhSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.exhSem.state'(sK0)),fun(fun(sK2,'type.exhSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))))),s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))))))),
  inference(backward_demodulation,[],[f20528,f11592])).
fof(f11592,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.exhSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.exhSem.state'(sK0)),fun(fun(sK2,'type.exhSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))))),s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))))))))),
  inference(unit_resulting_resolution,[],[f215,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f215,plain,(
  s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))) != s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4))) != s(fun(sK2,'type.exhSem.state'(sK0)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK3),'type.exhSem.state'(sK1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.exhSem.state'(sK3)),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.exhSem.state'(X3)),X4))) != s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.exhSem.state'(X3)),X4)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.exhSem.state'(X3)),X4))) = s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.exhSem.state'(X3)),X4)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.exhSem.state'(X3)),X4))) = s(fun(X2,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X3),'type.exhSem.state'(X1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.exhSem.state'(X3)),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.exhSem.state'(X6)),'const.combin.o_2'(s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.exhSem.state'(X9)),X10))) = s(fun(X8,'type.exhSem.state'(X6)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X7)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.exhSem.state'(X9)),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.exhSem.state'(X6)),'const.combin.o_2'(s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.exhSem.state'(X9)),X10))) = s(fun(X8,'type.exhSem.state'(X6)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X7)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.exhSem.state'(X9)),X10)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|5','thm.exhSem.state_fupdfupds_comp|split|5')).
fof(f20528,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X5,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))))),s(fun(X5,'type.exhSem.state'(X0)),X6))) = s(fun(X5,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X2),'type.exhSem.state'(X1)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3))),s(fun(X5,'type.exhSem.state'(X2)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X2)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))),s(fun(X5,'type.exhSem.state'(X0)),X6)))))) )),
  inference(superposition,[],[f277,f272])).
fof(f272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X2),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X2)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X2),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X2)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X2),'type.exhSem.state'(X0)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X2)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X9,X7] : ! [X11,X12] : s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))) = s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X6)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X6)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun('type.exhSem.state'(X9),'type.exhSem.state'(X7)),'const.exhSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|5','thm.exhSem.state_fupdfupds_comp|split|4')).
fof(f277,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X16,X8,X22,X17] : ! [X12,X11,X10] : s(fun(X22,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X22,X8),'const.combin.o_2'(s(fun(X17,X8),X11),s(fun(X22,X17),X10))))) = s(fun(X22,X16),'const.combin.o_2'(s(fun(X17,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X17,X8),X11))),s(fun(X22,X17),X10)))),
  file('thm.exhSem.state_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.exhSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 173
% Passive clauses: 1688
% Generated clauses: 21755
% Final active clauses: 107
% Final passive clauses: 1002
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 661
% Trivial inequalities: 54
% Fw subsumption resolutions: 721
% Fw demodulations: 3591
% Bw demodulations: 497
% Condensations: 937
% Global subsumptions: 2
% 
% Simple tautologies: 58
% Equational tautologies: 141
% Forward subsumptions: 13761
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 739
% Factoring: 1
% Forward superposition: 6712
% Backward superposition: 5935
% Self superposition: 201
% Equality factoring: 78
% Equality resolution: 2
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 1804
% SAT solver unit clauses: 410
% SAT solver binary clauses: 689
% 
% Sat splits: 26
% Sat splitting refutations: 60
% 
% Memory used [KB]: 23539
% Time elapsed: 0.842 s
% ------------------------------
% ------------------------------
% Success in time 2.16 s
