Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (129982)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 3947
% Generated clauses: 19951
% Final active clauses: 144
% Final passive clauses: 3523
% Input formulas: 68
% Initial clauses: 100
% 
% Duplicate literals: 519
% Trivial inequalities: 303
% Fw subsumption resolutions: 1303
% Fw demodulations: 2881
% Bw demodulations: 16
% 
% Simple tautologies: 239
% Equational tautologies: 45
% Forward subsumptions: 10358
% 
% Binary resolution: 329
% Factoring: 1
% Forward superposition: 6710
% Backward superposition: 7439
% Self superposition: 189
% Equality factoring: 70
% Equality resolution: 34
% 
% Split clauses: 204
% Split components: 410
% SAT solver clauses: 525
% SAT solver unit clauses: 9
% SAT solver binary clauses: 113
% 
% Sat splits: 204
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 524
% TWLsolver calls for satisfiability: 143
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f8597,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8409,f212,f317,f322,f324,f8090,f7986,f7987,f8019,f4572,f7989,f513,f519,f275,f360,f233,f337,f270,f354,f7260,f7528,f211,f315,f274,f358])).
fof(f358,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) ) | $spl62),
  inference(cnf_transformation,[],[f358_D])).
fof(f358_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f274,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X16] : ! [X6,X18,X31,X32,X33,X34] : s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X6))) = s(X16,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6))),s(fun('type.num.num',X16),X18),s(X16,X31),s(fun('type.num.num',X16),X32),s(fun(bool,X16),X33),s(X16,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bvp_to_wordProps.tag_case_def|split|0')).
fof(f315,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0))) | $spl0),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f211,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f168])).
fof(f168,plain,(
  ? [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bvp_to_wordProps.tag_distinct|split|3')).
fof(f7528,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) ) | $spl1140),
  inference(cnf_transformation,[],[f7528_D])).
fof(f7528_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) ) <=> ~$spl1140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1140])])).
fof(f7260,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) ) | ($spl0 | $spl62)),
  inference(superposition,[],[f358,f315])).
fof(f354,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) ) | $spl54),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f270,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X16] : ! [X18,X31,X32,X33,X34] : s(X16,X34) = s(X16,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X16),X18),s(X16,X31),s(fun('type.num.num',X16),X32),s(fun(bool,X16),X33),s(X16,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bvp_to_wordProps.tag_case_def|split|4')).
fof(f337,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) <=> p(s(bool,X28)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f360_D])).
fof(f360_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f275,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f519,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl156),
  inference(cnf_transformation,[],[f519_D])).
fof(f519_D,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f513,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl28 | $spl66)),
  inference(resolution,[],[f360,f337])).
fof(f7989,plain,(
  ( ! [X275] : (p(s(bool,happ(s(fun('type.num.num',bool),X275),s('type.num.num',sK0))))) ) | $spl1186),
  inference(cnf_transformation,[],[f7989_D])).
fof(f7989_D,plain,(
  ( ! [X275] : (p(s(bool,happ(s(fun('type.num.num',bool),X275),s('type.num.num',sK0))))) ) <=> ~$spl1186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1186])])).
fof(f4572,plain,(
  ( ! [X361] : (s(bool,'T_0') = s(bool,X361)) ) | $spl786),
  inference(cnf_transformation,[],[f4572_D])).
fof(f4572_D,plain,(
  ( ! [X361] : (s(bool,'T_0') = s(bool,X361)) ) <=> ~$spl786),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl786])])).
fof(f8019,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK0))))) ) | ($spl54 | $spl156 | $spl1140)),
  inference(condensation,[],[f8018])).
fof(f8018,plain,(
  ( ! [X397,X401,X402] : (s(bool,X401) = s(bool,X402) | p(s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X402)) ) | ($spl54 | $spl156 | $spl1140)),
  inference(forward_demodulation,[],[f7878,f354])).
fof(f7878,plain,(
  ( ! [X397,X399,X401,X398,X400,X402] : (p(s(bool,happ(s(fun('type.num.num',bool),X397),s('type.num.num',sK0)))) | s(bool,X402) = s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',bool),X397),s(bool,X398),s(fun('type.num.num',bool),X399),s(fun(bool,bool),X400),s(bool,X401))) | s(bool,'T_0') = s(bool,X402)) ) | ($spl156 | $spl1140)),
  inference(superposition,[],[f519,f7528])).
fof(f7987,plain,(
  ( ! [X270] : (~p(s(bool,happ(s(fun('type.num.num',bool),X270),s('type.num.num',sK0))))) ) | $spl1184),
  inference(cnf_transformation,[],[f7987_D])).
fof(f7987_D,plain,(
  ( ! [X270] : (~p(s(bool,happ(s(fun('type.num.num',bool),X270),s('type.num.num',sK0))))) ) <=> ~$spl1184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1184])])).
fof(f7986,plain,(
  ( ! [X270,X274] : (s(bool,'T_0') = s(bool,X274) | ~p(s(bool,happ(s(fun('type.num.num',bool),X270),s('type.num.num',sK0))))) ) | ($spl28 | $spl54 | $spl1140)),
  inference(forward_demodulation,[],[f7857,f354])).
fof(f7857,plain,(
  ( ! [X271,X273,X270,X272,X274] : (~p(s(bool,happ(s(fun('type.num.num',bool),X270),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',bool),X270),s(bool,X271),s(fun('type.num.num',bool),X272),s(fun(bool,bool),X273),s(bool,X274)))) ) | ($spl28 | $spl1140)),
  inference(superposition,[],[f337,f7528])).
fof(f8090,plain,(
  $false | ($spl1184 | $spl1186)),
  inference(subsumption_resolution,[],[f7987,f7989])).
fof(f324,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f322,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f217,f212])).
fof(f217,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f43])).
fof(f43,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3','thm.bool.NOT_CLAUSES|split|2')).
fof(f317,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f212,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_distinct|split|3',aHL_TRUTH)).
fof(f8409,plain,(
  $false | ($spl2 | $spl9 | $spl786)),
  inference(subsumption_resolution,[],[f8098,f317])).
fof(f8098,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl786)),
  inference(backward_demodulation,[],[f4572,f324])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 118
% Passive clauses: 604
% Generated clauses: 8346
% Final active clauses: 83
% Final passive clauses: 172
% Input formulas: 68
% Initial clauses: 100
% 
% Duplicate literals: 16
% Trivial inequalities: 159
% Fw subsumption resolutions: 375
% Fw demodulations: 1973
% Bw demodulations: 323
% Condensations: 40
% Global subsumptions: 8
% 
% Simple tautologies: 39
% Equational tautologies: 12
% Forward subsumptions: 4025
% Bw demodulations to eq. taut.: 22
% 
% Binary resolution: 109
% Factoring: 1
% Forward superposition: 2509
% Backward superposition: 2004
% Self superposition: 71
% Equality factoring: 4
% Equality resolution: 23
% 
% Split clauses: 99
% Split components: 202
% SAT solver clauses: 1160
% SAT solver unit clauses: 151
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 2
% 
% Sat splits: 99
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1121
% TWLsolver calls for satisfiability: 330
% 
% Memory used [KB]: 10106
% Time elapsed: 0.465 s
% ------------------------------
% ------------------------------
% Success in time 1.12 s
