Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modSem
% (116540)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 186
% Passive clauses: 3033
% Generated clauses: 23687
% Final active clauses: 131
% Final passive clauses: 2620
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 865
% Trivial inequalities: 168
% Fw subsumption resolutions: 2345
% Fw demodulations: 4084
% Bw demodulations: 48
% 
% Simple tautologies: 189
% Equational tautologies: 49
% Forward subsumptions: 12241
% 
% Binary resolution: 631
% Factoring: 3
% Forward superposition: 8788
% Backward superposition: 6210
% Self superposition: 220
% Equality factoring: 96
% Equality resolution: 16
% 
% Split clauses: 291
% Split components: 582
% SAT solver clauses: 889
% SAT solver unit clauses: 43
% SAT solver binary clauses: 153
% 
% Sat splits: 291
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 889
% TWLsolver calls for satisfiability: 167
% 
% Memory used [KB]: 14583
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modSem
% (116560)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 136
% Passive clauses: 1164
% Generated clauses: 14839
% Final active clauses: 101
% Final passive clauses: 796
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 460
% Trivial inequalities: 90
% Fw subsumption resolutions: 883
% Fw demodulations: 1699
% Bw demodulations: 261
% Condensations: 1099
% 
% Simple tautologies: 142
% Equational tautologies: 38
% Forward subsumptions: 7505
% 
% Binary resolution: 366
% Factoring: 1
% Forward superposition: 4644
% Backward superposition: 3699
% Self superposition: 193
% Equality factoring: 76
% Equality resolution: 9
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 1833
% SAT solver unit clauses: 259
% SAT solver binary clauses: 55
% 
% Sat splits: 43
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1833
% TWLsolver calls for satisfiability: 348
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.modSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modSem
% SZS output start Proof for thm.modSem
fof(f49931,plain,(
  $false),
  inference(subsumption_resolution,[],[f49930,f403])).
fof(f403,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,axiom,(
  p(s(bool,'T_0'))),
  file('thm.modSem.state_fupdfupds_comp|split|7','thm.bool.TRUTH')).
fof(f49930,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f49928,f4397])).
fof(f4397,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f4396])).
fof(f4396,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f4324,f551])).
fof(f551,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f405,f406])).
fof(f406,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.modSem.state_fupdfupds_comp|split|7',aHL_BOOL_CASES)).
fof(f405,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.modSem.state_fupdfupds_comp|split|7',aHL_FALSITY)).
fof(f4324,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f3635,f549])).
fof(f549,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f406,f406])).
fof(f3635,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3535,f1569])).
fof(f1569,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X0,X0),'_i_0_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1512,f491])).
fof(f491,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X6,X14] : ! [X10,X9,X16] : s(X13,'const.combin.o_3'(s(fun(X6,X13),X10),s(fun(X14,X6),X9),s(X14,X16))) = s(X13,happ(s(fun(X6,X13),X10),s(X6,happ(s(fun(X14,X6),X9),s(X14,X16)))))),
  file('thm.modSem.state_fupdfupds_comp|split|7','thm.combin.o_THM')).
fof(f1512,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f407,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f401])).
fof(f401,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X47,X48,X49] : (p(s(bool,happ(s(fun(X47,bool),happ(s(fun(X47,fun(X47,bool)),'=_0'),s(X47,X48))),s(X47,X49)))) <=> s(X47,X48) = s(X47,X49))),
  file('thm.modSem.state_fupdfupds_comp|split|7',arity_0_HL_EQ)).
fof(f407,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X41] : ! [X43] : s(X41,X43) = s(X41,happ(s(fun(X41,X41),'_i_0_0'),s(X41,X43)))),
  file('thm.modSem.state_fupdfupds_comp|split|7','_i0_ax')).
fof(f3535,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,happ(s(fun(X73,X75),X76),s(X73,X74))) = s(X75,'const.combin.o_3'(s(fun(X73,X75),X76),s(fun(X73,X73),'_i_0_0'),s(X73,X74)))) )),
  inference(superposition,[],[f491,f407])).
fof(f49928,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.modSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.modSem.state'(sK1)),fun(fun(sK0,'type.modSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))))),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f49868,f11207])).
fof(f11207,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.modSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.modSem.state'(sK1)),fun(fun(sK0,'type.modSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4))),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f402,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f401])).
fof(f402,plain,(
  s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))) != s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4))),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))),s(fun(sK0,'type.modSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))))),s(fun(sK0,'type.modSem.state'(sK1)),sK2))) != s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK4))),s(fun(sK0,'type.modSem.state'(sK1)),'const.combin.o_2'(s(fun('type.modSem.state'(sK1),'type.modSem.state'(sK1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),sK3))),s(fun(sK0,'type.modSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f381])).
fof(f381,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))))),s(fun(X0,'type.modSem.state'(X1)),X2))) != s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4))),s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))),s(fun(X0,'type.modSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))))),s(fun(X0,'type.modSem.state'(X1)),X2))) = s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4))),s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))),s(fun(X0,'type.modSem.state'(X1)),X2)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))))),s(fun(X0,'type.modSem.state'(X1)),X2))) = s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X4))),s(fun(X0,'type.modSem.state'(X1)),'const.combin.o_2'(s(fun('type.modSem.state'(X1),'type.modSem.state'(X1)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X3))),s(fun(X0,'type.modSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9))))),s(fun(X6,'type.modSem.state'(X7)),X8))) = s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10))),s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9))),s(fun(X6,'type.modSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9))))),s(fun(X6,'type.modSem.state'(X7)),X8))) = s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10))),s(fun(X6,'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9))),s(fun(X6,'type.modSem.state'(X7)),X8)))))),
  file('thm.modSem.state_fupdfupds_comp|split|7','thm.modSem.state_fupdfupds_comp|split|7')).
fof(f49868,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2))))),s(fun(X3,'type.modSem.state'(X0)),X4))) = s(fun(X3,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1))),s(fun(X3,'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2))),s(fun(X3,'type.modSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f514,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1))))) = s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2))),s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1)))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1))))) = s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2))),s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1)))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1))))) = s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.combin.o_2'(s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X2))),s(fun('type.modSem.state'(X0),'type.modSem.state'(X0)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),'const.combin.o_2'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10),s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9))))) = s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.combin.o_2'(s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X10))),s(fun('type.modSem.state'(X7),'type.modSem.state'(X7)),'const.modSem.state_defined_types_fupd_1'(s(fun(fun('type.semanticPrimitives.tid_or_exn',bool),fun('type.semanticPrimitives.tid_or_exn',bool)),X9)))))),
  file('thm.modSem.state_fupdfupds_comp|split|7','thm.modSem.state_fupdfupds_comp|split|6')).
fof(f514,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13,X6,X20,X14] : ! [X10,X9,X8] : s(fun(X20,X13),'const.combin.o_2'(s(fun(X6,X13),X10),s(fun(X20,X6),'const.combin.o_2'(s(fun(X14,X6),X9),s(fun(X20,X14),X8))))) = s(fun(X20,X13),'const.combin.o_2'(s(fun(X14,X13),'const.combin.o_2'(s(fun(X6,X13),X10),s(fun(X14,X6),X9))),s(fun(X20,X14),X8)))),
  file('thm.modSem.state_fupdfupds_comp|split|7','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.modSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 320
% Passive clauses: 4632
% Generated clauses: 53485
% Final active clauses: 143
% Final passive clauses: 1522
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 1548
% Trivial inequalities: 190
% Fw subsumption resolutions: 1982
% Fw demodulations: 8870
% Bw demodulations: 2591
% Condensations: 2279
% Global subsumptions: 343
% 
% Simple tautologies: 253
% Equational tautologies: 250
% Forward subsumptions: 32599
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 77
% 
% Unit resulting resolution: 1583
% Factoring: 1
% Forward superposition: 16071
% Backward superposition: 12968
% Self superposition: 441
% Equality factoring: 156
% Equality resolution: 16
% 
% Split clauses: 57
% Split components: 114
% SAT solver clauses: 5520
% SAT solver unit clauses: 1237
% SAT solver binary clauses: 2387
% 
% Sat splits: 57
% Sat splitting refutations: 429
% 
% Memory used [KB]: 48613
% Time elapsed: 2.218 s
% ------------------------------
% ------------------------------
% Success in time 3.56 s
