Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (9439)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 3660
% Generated clauses: 20504
% Final active clauses: 139
% Final passive clauses: 3243
% Input formulas: 68
% Initial clauses: 97
% 
% Duplicate literals: 571
% Trivial inequalities: 344
% Fw subsumption resolutions: 1628
% Fw demodulations: 2882
% Bw demodulations: 19
% 
% Simple tautologies: 241
% Equational tautologies: 41
% Forward subsumptions: 10302
% 
% Binary resolution: 514
% Factoring: 1
% Forward superposition: 6795
% Backward superposition: 7292
% Self superposition: 176
% Equality factoring: 70
% Equality resolution: 25
% 
% Split clauses: 249
% Split components: 501
% SAT solver clauses: 689
% SAT solver unit clauses: 10
% SAT solver binary clauses: 153
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 249
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 687
% TWLsolver calls for satisfiability: 178
% 
% Memory used [KB]: 19573
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (9456)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 1496
% Generated clauses: 16385
% Final active clauses: 107
% Final passive clauses: 917
% Input formulas: 68
% Initial clauses: 97
% 
% Duplicate literals: 470
% Trivial inequalities: 160
% Fw subsumption resolutions: 831
% Fw demodulations: 1763
% Bw demodulations: 439
% Condensations: 958
% Global subsumptions: 4
% 
% Simple tautologies: 192
% Equational tautologies: 38
% Forward subsumptions: 8572
% Bw demodulations to eq. taut.: 30
% 
% Binary resolution: 419
% Factoring: 1
% Forward superposition: 5282
% Backward superposition: 4218
% Self superposition: 139
% Equality factoring: 62
% Equality resolution: 18
% 
% Split clauses: 78
% Split components: 157
% SAT solver clauses: 1960
% SAT solver unit clauses: 141
% SAT solver binary clauses: 52
% SAT solver learnt clauses: 9
% SAT solver learnt literals: 1
% 
% Sat splits: 78
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1943
% TWLsolver calls for satisfiability: 378
% 
% Memory used [KB]: 16502
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f72089,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f225,f327,f325,f323,f224,f326,f223,f324,f322,f334,f836,f847,f846,f3030,f3031,f3015,f3038,f3033,f5341,f5343,f5326,f5941,f5945,f5943,f5942,f6281,f6283,f6294,f6299,f6301,f6303,f6313,f6316,f6319,f6321,f6323,f6325,f6361,f6371,f6373,f6375,f6377,f6379,f6380,f6381,f6382,f6384,f6386,f6388,f6389,f6390,f14191,f14193,f14195,f14199,f14201,f14203,f14205,f14207,f14209,f14211,f14213,f14215,f14217,f14219,f14221,f14223,f14225,f14227,f14229,f14231,f14233,f14235,f14237,f14241,f14243,f14245,f14247,f14249,f14251,f14253,f14255,f14257,f14259,f14261,f14263,f14265,f14267,f14269,f14271,f14273,f14275,f14277,f14279,f14284,f14359,f14361,f14363,f14365,f14367,f14369,f14371,f14373,f14375,f14377,f14383,f14388,f14392,f14466,f14467,f14468,f14469,f14470,f14471,f14472,f14473,f14474,f14475,f14476,f14477,f14478,f14479,f14480,f14481,f14482,f14483,f14484,f14485,f14486,f14487,f14488,f14489,f14490,f14491,f14492,f14493,f14494,f14495,f14496,f14497,f14498,f14499,f14500,f14501,f14502,f14503,f14504,f14505,f14506,f14507,f14508,f14509,f14510,f14511,f14512,f14513,f14514,f14515,f14516,f9187,f5353,f5345,f64480,f64772,f64770,f64481,f64771,f64769,f66279,f66290,f66320,f69934,f69945,f69967,f72082,f72084])).
fof(f72084,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f72083,f325])).
fof(f72083,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | $spl0),
  inference(forward_demodulation,[],[f71859,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X6] : ! [X84,X49] : s(X6,X84) = s(X6,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(X6,X84),s(X6,X49)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.ast.word_size_case_def|split|0')).
fof(f71859,plain,(
  s('type.num.num',sK3) = s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(superposition,[],[f259,f70082])).
fof(f70082,plain,(
  ( ! [X3] : (s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size',X3),s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size',X3),s('type.num.num',sK3),s('type.num.num',sK2)))) ) | $spl0),
  inference(backward_demodulation,[],[f69966,f68417])).
fof(f68417,plain,(
  ( ! [X3] : (s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size',X3),s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size',X3),s('type.num.num',sK3),s('type.num.num',sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f67923,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size',X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun('type.ast.word_size',fun(X0,fun(X0,X0))),'const.ast.word_size_CASE_0'),s('type.ast.word_size',X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size',X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun('type.ast.word_size',fun(X0,fun(X0,X0))),'const.ast.word_size_CASE_0'),s('type.ast.word_size',X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size',X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun('type.ast.word_size',fun(X0,fun(X0,X0))),'const.ast.word_size_CASE_0'),s('type.ast.word_size',X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6] : ! [X11,X12,X24] : s(X6,'const.ast.word_size_CASE_3'(s('type.ast.word_size',X11),s(X6,X12),s(X6,X24))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),happ(s(fun('type.ast.word_size',fun(X6,fun(X6,X6))),'const.ast.word_size_CASE_0'),s('type.ast.word_size',X11))),s(X6,X12))),s(X6,X24)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_3_const.ast.word_size_CASE')).
fof(f67923,plain,(
  ( ! [X3] : (s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size',X3),s('type.num.num',sK3),s('type.num.num',sK4))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),happ(s(fun('type.ast.word_size',fun('type.num.num',fun('type.num.num','type.num.num'))),'const.ast.word_size_CASE_0'),s('type.ast.word_size',X3))),s('type.num.num',sK1))),s('type.num.num',sK2)))) ) | $spl0),
  inference(superposition,[],[f66640,f270])).
fof(f66640,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK1))),s('type.num.num',sK2))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK3))),s('type.num.num',sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f66417,f305])).
fof(f305,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.num.num',X1),happ(s(fun('type.num.num',fun('type.num.num',X1)),X5),s('type.num.num',X2))),s('type.num.num',X3))) = s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.fcp.cart'(bool,X0),X1),X4),s(fun('type.num.num',fun('type.num.num',X1)),X5)))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun('type.num.num',X1),happ(s(fun('type.num.num',fun('type.num.num',X1)),X5),s('type.num.num',X2))),s('type.num.num',X3))) = s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.fcp.cart'(bool,X0),X1),X4),s(fun('type.num.num',fun('type.num.num',X1)),X5)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X1,happ(s(fun('type.num.num',X1),happ(s(fun('type.num.num',fun('type.num.num',X1)),X5),s('type.num.num',X2))),s('type.num.num',X3))) = s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))),s(fun('type.fcp.cart'(bool,X0),X1),X4),s(fun('type.num.num',fun('type.num.num',X1)),X5)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X14] : ! [X7,X8,X13,X15] : s(X14,happ(s(fun('type.num.num',X14),happ(s(fun('type.num.num',fun('type.num.num',X14)),X15),s('type.num.num',X7))),s('type.num.num',X8))) = s(X14,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X7),s('type.num.num',X8))),s(fun('type.fcp.cart'(bool,X6),X14),X13),s(fun('type.num.num',fun('type.num.num',X14)),X15)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_case_def|split|1')).
fof(f66417,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK3))),s('type.num.num',sK4))) = s(X0,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))),s(fun('type.fcp.cart'(bool,sK0),X0),X2),s(fun('type.num.num',fun('type.num.num',X0)),X1)))) ) | $spl0),
  inference(superposition,[],[f305,f322])).
fof(f69966,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | $spl0),
  inference(forward_demodulation,[],[f69492,f254])).
fof(f254,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X18] : ! [X20] : s(X18,X20) = s(X18,happ(s(fun(X18,X18),'_i_0_0'),s(X18,X20)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_i0_ax')).
fof(f69492,plain,(
  s('type.num.num',sK4) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK2))) | $spl0),
  inference(superposition,[],[f254,f68369])).
fof(f68369,plain,(
  ( ! [X57,X56] : (s(X56,happ(s(fun('type.num.num',X56),X57),s('type.num.num',sK2))) = s(X56,happ(s(fun('type.num.num',X56),X57),s('type.num.num',sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f67921,f298])).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X18,X16] : ! [X19,X21] : s(X16,X19) = s(X16,happ(s(fun(X18,X16),'_dst_x0_1'(s(X16,X19))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x0_ax')).
fof(f67921,plain,(
  ( ! [X57,X56] : (s(X56,happ(s(fun('type.num.num',X56),X57),s('type.num.num',sK4))) = s(X56,happ(s(fun('type.num.num',X56),happ(s(fun('type.num.num',fun('type.num.num',X56)),'_dst_x0_1'(s(fun('type.num.num',X56),X57))),s('type.num.num',sK1))),s('type.num.num',sK2)))) ) | $spl0),
  inference(superposition,[],[f66640,f298])).
fof(f72082,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f72081,f325])).
fof(f72081,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | $spl0),
  inference(forward_demodulation,[],[f71840,f259])).
fof(f71840,plain,(
  s('type.num.num',sK3) = s('type.num.num','const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(superposition,[],[f70082,f259])).
fof(f69967,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f69966,f327])).
fof(f69945,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f69944,f327])).
fof(f69944,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | $spl0),
  inference(forward_demodulation,[],[f69456,f254])).
fof(f69456,plain,(
  s('type.num.num',sK4) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK2))) | $spl0),
  inference(superposition,[],[f68369,f254])).
fof(f69934,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f69406,f4597])).
fof(f4597,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f4596,f310])).
fof(f310,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X16,X18,X17] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),X19),s(X18,X21))),s(X17,X20))) = s(X16,happ(s(fun(X18,X16),'_dst_x1x0_2'(s(fun(X18,fun(X17,X16)),X19),s(X17,X20))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x1x0_ax')).
fof(f4596,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f4595,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : ! [X84,X49] : s(X6,X49) = s(X6,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X6,X84),s(X6,X49)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.ast.word_size_case_def|split|1')).
fof(f4595,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X0,X1),s(X0,X2))))))) )),
  inference(forward_demodulation,[],[f4457,f310])).
fof(f4457,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(X0,X1),s(X0,X2))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f260,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X42,X43,X44] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),'=_0'),s(X42,X43))),s(X42,X44)))) <=> s(X42,X43) = s(X42,X44))),
  file('thm.wordSem.word_loc_|49||49||split|1',arity_0_HL_EQ)).
fof(f69406,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK2)))) | ($spl0 | $spl5)),
  inference(unit_resulting_resolution,[],[f4385,f68369,f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f234,f229])).
fof(f229,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.wordSem.word_loc_|49||49||split|1',aHL_BOOL_CASES)).
fof(f234,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X32] : (s(bool,'F_0') = s(bool,X32) <=> ~p(s(bool,X32)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f4385,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4))) | $spl5),
  inference(unit_resulting_resolution,[],[f4214,f230])).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X32] : (s(bool,'T_0') = s(bool,X32) <=> p(s(bool,X32)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f4214,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4)))) | $spl5),
  inference(unit_resulting_resolution,[],[f327,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f66320,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f66319,f327])).
fof(f66319,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f66060,f254])).
fof(f66060,plain,(
  s('type.num.num',sK4) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f254,f64223])).
fof(f64223,plain,(
  ( ! [X57,X56] : (s(X56,happ(s(fun('type.num.num',X56),X57),s('type.num.num',sK2))) = s(X56,happ(s(fun('type.num.num',X56),X57),s('type.num.num',sK4)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f16758,f298])).
fof(f16758,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK1))),s('type.num.num',sK2))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK1))),s('type.num.num',sK4)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f16593,f305])).
fof(f16593,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.num.num',fun('type.num.num',X0)),X1),s('type.num.num',sK1))),s('type.num.num',sK4))) = s(X0,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))),s(fun('type.fcp.cart'(bool,sK0),X0),X2),s(fun('type.num.num',fun('type.num.num',X0)),X1)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f305,f335])).
fof(f335,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f322,f324])).
fof(f66290,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f66289,f327])).
fof(f66289,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f66017,f254])).
fof(f66017,plain,(
  s('type.num.num',sK4) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f64223,f254])).
fof(f66279,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f65968,f4597])).
fof(f65968,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK2)))) | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f4385,f64223,f357])).
fof(f64769,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK4) | $spl26),
  inference(cnf_transformation,[],[f64769_D])).
fof(f64769_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK4) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f64771,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) | $spl29),
  inference(cnf_transformation,[],[f64771_D])).
fof(f64771_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f64481,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) | s('type.num.num',sK1) = s('type.num.num',sK4) | ($spl0 | $spl2)),
  inference(superposition,[],[f306,f16758])).
fof(f64770,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK4) | $spl27),
  inference(cnf_transformation,[],[f64770_D])).
fof(f64770_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK4) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f64772,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) | $spl28),
  inference(cnf_transformation,[],[f64772_D])).
fof(f64772_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f64480,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK2)))) | s('type.num.num',sK1) != s('type.num.num',sK4) | ($spl0 | $spl2)),
  inference(superposition,[],[f307,f16758])).
fof(f5345,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl18),
  inference(cnf_transformation,[],[f5345_D])).
fof(f5345_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5353,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) | $spl21),
  inference(cnf_transformation,[],[f5353_D])).
fof(f5353_D,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f9187,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f287,f254])).
fof(f287,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X42,X45] : (p(s(bool,happ(s(fun(fun(X42,bool),bool),'!_1'),s(fun(X42,bool),X45)))) <=> ! [X43] : p(s(bool,happ(s(fun(X42,bool),X45),s(X42,X43)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','!_HL_QUANT')).
fof(f14516,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14194,f14382])).
fof(f14382,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f14186,f306])).
fof(f14186,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl18),
  inference(forward_demodulation,[],[f14118,f254])).
fof(f14118,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f5345,f286])).
fof(f286,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f14194,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f14186,f234])).
fof(f14515,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14198,f226])).
fof(f226,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f28])).
fof(f28,axiom,(
  p(s(bool,'T_0'))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.TRUTH')).
fof(f14198,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f339,f14186,f285])).
fof(f285,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X22,X23] : ((p(s(bool,X22)) => p(s(bool,X23))) => ((p(s(bool,X23)) => p(s(bool,X22))) => s(bool,X22) = s(bool,X23)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f339,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f228,f230])).
fof(f228,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.wordSem.word_loc_|49||49||split|1',aHL_FALSITY)).
fof(f14514,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f14200,f226])).
fof(f14200,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f853,f14186,f285])).
fof(f853,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK5(X0,'_i_0_0',bool))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f847,f230])).
fof(f14513,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14202,f226])).
fof(f14202,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4385,f14186,f285])).
fof(f14512,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14204,f226])).
fof(f14204,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4413,f14186,f285])).
fof(f4413,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2))) | $spl5),
  inference(unit_resulting_resolution,[],[f4217,f230])).
fof(f4217,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f327,f306])).
fof(f14511,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14206,f226])).
fof(f14206,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5399,f14186,f285])).
fof(f5399,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f5355,f230])).
fof(f5355,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1')))) | $spl5),
  inference(unit_resulting_resolution,[],[f4376,f286])).
fof(f4376,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4214,f286])).
fof(f14510,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14208,f226])).
fof(f14208,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6965,f14186,f285])).
fof(f6965,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f5388,f230])).
fof(f5388,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1')))) | $spl5),
  inference(unit_resulting_resolution,[],[f5355,f286])).
fof(f14509,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14210,f226])).
fof(f14210,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5802,f14186,f285])).
fof(f5802,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5387,f230])).
fof(f5387,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5355,f287])).
fof(f14508,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14212,f226])).
fof(f14212,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5366,f14186,f285])).
fof(f5366,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4376,f230])).
fof(f14507,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14214,f226])).
fof(f14214,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7518,f14186,f285])).
fof(f7518,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4404,f230])).
fof(f4404,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4217,f286])).
fof(f14506,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14216,f226])).
fof(f14216,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5847,f14186,f285])).
fof(f5847,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))) | $spl5),
  inference(unit_resulting_resolution,[],[f5800,f230])).
fof(f5800,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5387,f287])).
fof(f14505,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f14220,f14186])).
fof(f14220,plain,(
  ( ! [X0] : (~p(s(bool,sK5(X0,'_i_0_0',bool)))) ) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f853,f14186,f285])).
fof(f14504,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14222,f14186])).
fof(f14222,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4)))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4385,f14186,f285])).
fof(f14503,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14224,f14186])).
fof(f14224,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2)))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4413,f14186,f285])).
fof(f14502,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14226,f14186])).
fof(f14226,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1')))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5399,f14186,f285])).
fof(f14501,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14228,f14186])).
fof(f14228,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1')))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6965,f14186,f285])).
fof(f14500,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14230,f14186])).
fof(f14230,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5802,f14186,f285])).
fof(f14499,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14232,f14186])).
fof(f14232,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5366,f14186,f285])).
fof(f14498,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14234,f14186])).
fof(f14234,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7518,f14186,f285])).
fof(f14497,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14236,f14186])).
fof(f14236,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5847,f14186,f285])).
fof(f14496,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14240,f226])).
fof(f14240,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f339,f14186,f285])).
fof(f14495,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f14242,f226])).
fof(f14242,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f853,f14186,f285])).
fof(f14494,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14244,f226])).
fof(f14244,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4385,f14186,f285])).
fof(f14493,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14246,f226])).
fof(f14246,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4413,f14186,f285])).
fof(f14492,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14248,f226])).
fof(f14248,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5399,f14186,f285])).
fof(f14491,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14250,f226])).
fof(f14250,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6965,f14186,f285])).
fof(f14490,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14252,f226])).
fof(f14252,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5802,f14186,f285])).
fof(f14489,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14254,f226])).
fof(f14254,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5366,f14186,f285])).
fof(f14488,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14256,f226])).
fof(f14256,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7518,f14186,f285])).
fof(f14487,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14258,f226])).
fof(f14258,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5847,f14186,f285])).
fof(f14486,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f14262,f14186])).
fof(f14262,plain,(
  ( ! [X0] : (~p(s(bool,sK5(X0,'_i_0_0',bool)))) ) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f853,f14186,f285])).
fof(f14485,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14264,f14186])).
fof(f14264,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4)))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4385,f14186,f285])).
fof(f14484,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14266,f14186])).
fof(f14266,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2)))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4413,f14186,f285])).
fof(f14483,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14268,f14186])).
fof(f14268,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1')))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5399,f14186,f285])).
fof(f14482,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14270,f14186])).
fof(f14270,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1')))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6965,f14186,f285])).
fof(f14481,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14272,f14186])).
fof(f14272,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5802,f14186,f285])).
fof(f14480,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14274,f14186])).
fof(f14274,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5366,f14186,f285])).
fof(f14479,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14276,f14186])).
fof(f14276,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4)))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7518,f14186,f285])).
fof(f14478,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14278,f14186])).
fof(f14278,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5847,f14186,f285])).
fof(f14477,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14282,f14186])).
fof(f14282,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0'))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f228,f14186,f293])).
fof(f293,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X43))),s(bool,X44)))) <=> (p(s(bool,X43)) => p(s(bool,X44))))),
  file('thm.wordSem.word_loc_|49||49||split|1','==>_HL_0')).
fof(f14476,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14358,f14382])).
fof(f14358,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f339,f14186,f357])).
fof(f14475,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f14360,f14382])).
fof(f14360,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK5(X1,'_i_0_0',bool))) ) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f853,f14186,f357])).
fof(f14474,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14362,f14382])).
fof(f14362,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4)))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4385,f14186,f357])).
fof(f14473,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14364,f14382])).
fof(f14364,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2)))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4413,f14186,f357])).
fof(f14472,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14366,f14382])).
fof(f14366,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1')))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5399,f14186,f357])).
fof(f14471,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14368,f14382])).
fof(f14368,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1')))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6965,f14186,f357])).
fof(f14470,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14370,f14382])).
fof(f14370,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5802,f14186,f357])).
fof(f14469,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14372,f14382])).
fof(f14372,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2)))))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5366,f14186,f357])).
fof(f14468,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14374,f14382])).
fof(f14374,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4)))))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7518,f14186,f357])).
fof(f14467,plain,(
  $false | ($spl5 | $spl18)),
  inference(subsumption_resolution,[],[f14376,f14382])).
fof(f14376,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) ) | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5847,f14186,f357])).
fof(f14466,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f14391,f14186])).
fof(f14391,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f228,f14186,f293])).
fof(f14392,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f226,f228,f14186,f293])).
fof(f14388,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f228,f228,f14186,f290])).
fof(f290,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X43))),s(bool,X44)))) <=> (p(s(bool,X43)) | p(s(bool,X44))))),
  file('thm.wordSem.word_loc_|49||49||split|1','|bslash||slash|_HL_0')).
fof(f14383,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f327,f14186,f306])).
fof(f14377,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5874,f5847,f14186,f357])).
fof(f5874,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))) | $spl5),
  inference(forward_demodulation,[],[f5861,f4412])).
fof(f4412,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2))) | $spl5),
  inference(unit_resulting_resolution,[],[f4217,f235])).
fof(f235,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f5861,plain,(
  s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2))) = s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))) | $spl5),
  inference(unit_resulting_resolution,[],[f4413,f5800,f413])).
fof(f413,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f407,f228])).
fof(f407,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f282,f229])).
fof(f282,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f14375,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f7542,f7518,f14186,f357])).
fof(f7542,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))))) | $spl5),
  inference(forward_demodulation,[],[f7531,f5874])).
fof(f7531,plain,(
  s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))) = s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5847,f4404,f413])).
fof(f14373,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5365,f5366,f14186,f357])).
fof(f5365,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4376,f235])).
fof(f14371,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5829,f5802,f14186,f357])).
fof(f5829,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))) | $spl5),
  inference(forward_demodulation,[],[f5816,f4412])).
fof(f5816,plain,(
  s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2))) | $spl5),
  inference(unit_resulting_resolution,[],[f4413,f5387,f413])).
fof(f14369,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f6987,f6965,f14186,f357])).
fof(f6987,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'))) | $spl5),
  inference(forward_demodulation,[],[f6976,f5829])).
fof(f6976,plain,(
  s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))) = s(bool,happ(s(fun(fun(fun(fun('type.num.num',bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f5802,f5388,f413])).
fof(f14367,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f5398,f5399,f14186,f357])).
fof(f5398,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(fun('type.num.num',bool),bool),bool),'!_1'),s(fun(fun('type.num.num',bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f5355,f235])).
fof(f14365,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4412,f4413,f14186,f357])).
fof(f14363,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f4384,f4385,f14186,f357])).
fof(f4384,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4))) | $spl5),
  inference(unit_resulting_resolution,[],[f4214,f235])).
fof(f14361,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f852,f853,f14186,f357])).
fof(f852,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f847,f235])).
fof(f14359,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f852,f339,f14186,f357])).
fof(f14284,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f228,f14186,f14186,f293])).
fof(f14279,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5847,f14186,f285])).
fof(f14277,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f7518,f14186,f285])).
fof(f14275,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5366,f14186,f285])).
fof(f14273,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5802,f14186,f285])).
fof(f14271,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f6965,f14186,f285])).
fof(f14269,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5399,f14186,f285])).
fof(f14267,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f4413,f14186,f285])).
fof(f14265,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f4385,f14186,f285])).
fof(f14263,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f853,f14186,f285])).
fof(f14261,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f14186,f339,f14186,f285])).
fof(f14259,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5847,f14186,f285])).
fof(f14257,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f7518,f14186,f285])).
fof(f14255,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5366,f14186,f285])).
fof(f14253,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5802,f14186,f285])).
fof(f14251,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f6965,f14186,f285])).
fof(f14249,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5399,f14186,f285])).
fof(f14247,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f4413,f14186,f285])).
fof(f14245,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f4385,f14186,f285])).
fof(f14243,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f853,f14186,f285])).
fof(f14241,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f226,f339,f14186,f285])).
fof(f14237,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5847,f14186,f285])).
fof(f14235,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f7518,f14186,f285])).
fof(f14233,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5366,f14186,f285])).
fof(f14231,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5802,f14186,f285])).
fof(f14229,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f6965,f14186,f285])).
fof(f14227,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f5399,f14186,f285])).
fof(f14225,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f4413,f14186,f285])).
fof(f14223,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f4385,f14186,f285])).
fof(f14221,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f14186,f853,f14186,f285])).
fof(f14219,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f14186,f339,f14186,f285])).
fof(f14217,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5847,f14186,f285])).
fof(f14215,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f7518,f14186,f285])).
fof(f14213,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5366,f14186,f285])).
fof(f14211,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5802,f14186,f285])).
fof(f14209,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f6965,f14186,f285])).
fof(f14207,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f5399,f14186,f285])).
fof(f14205,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f4413,f14186,f285])).
fof(f14203,plain,(
  $false | ($spl5 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f4385,f14186,f285])).
fof(f14201,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f226,f853,f14186,f285])).
fof(f14199,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f226,f339,f14186,f285])).
fof(f14195,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f852,f14186,f234])).
fof(f14193,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f339,f14186,f231])).
fof(f231,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f14191,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f228,f14186])).
fof(f6390,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6229,f226])).
fof(f6229,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5163])).
fof(f5163,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))))),s('type.num.num',sK2)))) | $spl5),
  inference(forward_demodulation,[],[f5137,f312])).
fof(f312,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X16,X17,X38,X18] : ! [X19,X20,X39,X21] : s(X16,happ(s(fun(X18,X16),'_dst_x0x0x1_3'(s(fun(X17,fun(X38,X16)),X19),s(X17,X20),s(fun(X18,X38),X39))),s(X18,X21))) = s(X16,happ(s(fun(X38,X16),happ(s(fun(X17,fun(X38,X16)),X19),s(X17,X20))),s(X38,happ(s(fun(X18,X38),X39),s(X18,X21)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x0x0x1_ax')).
fof(f5137,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK4))),s('type.num.num',sK2)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4413,f306])).
fof(f6389,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6228,f226])).
fof(f6228,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f4928])).
fof(f4928,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))))),s('type.num.num',sK4)))) | $spl5),
  inference(forward_demodulation,[],[f4902,f312])).
fof(f4902,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK4)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4385,f306])).
fof(f6388,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374,f6376,f6383,f6387])).
fof(f6387,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl22),
  inference(subsumption_resolution,[],[f6223,f226])).
fof(f6223,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f4348])).
fof(f6383,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl22),
  inference(subsumption_resolution,[],[f6216,f226])).
fof(f6216,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f4321])).
fof(f6376,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f6254,f226])).
fof(f6254,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl22),
  inference(condensation,[],[f6185])).
fof(f6185,plain,(
  ( ! [X74,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | p(s(bool,X74))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f5966])).
fof(f6374,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f6268,f226])).
fof(f6268,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl22),
  inference(condensation,[],[f6167])).
fof(f6167,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f293])).
fof(f6372,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f6273,f226])).
fof(f6273,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | $spl22),
  inference(condensation,[],[f6159])).
fof(f6159,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f290])).
fof(f6360,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl22),
  inference(subsumption_resolution,[],[f6141,f226])).
fof(f6141,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f306])).
fof(f6318,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f6317,f226])).
fof(f6317,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl22),
  inference(forward_demodulation,[],[f6099,f5942])).
fof(f6099,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f5747])).
fof(f6315,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f6314,f226])).
fof(f6314,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl22),
  inference(forward_demodulation,[],[f6097,f5942])).
fof(f6097,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f5699])).
fof(f6298,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl22),
  inference(subsumption_resolution,[],[f6297,f226])).
fof(f6297,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl22),
  inference(condensation,[],[f6296])).
fof(f6296,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl22),
  inference(condensation,[],[f6295])).
fof(f6295,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl22),
  inference(forward_demodulation,[],[f6059,f5942])).
fof(f6059,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X63))),s(bool,'T_0')))) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f4277])).
fof(f6293,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl22),
  inference(subsumption_resolution,[],[f6292,f226])).
fof(f6292,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl22),
  inference(condensation,[],[f6291])).
fof(f6291,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl22),
  inference(condensation,[],[f6290])).
fof(f6290,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl22),
  inference(forward_demodulation,[],[f6054,f5942])).
fof(f6054,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X63)))) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f4237])).
fof(f6282,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f5976,f5942])).
fof(f5976,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl22),
  inference(backward_demodulation,[],[f5942,f234])).
fof(f5966,plain,(
  ( ! [X74,X73] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X74))),s(bool,X73)))) | ~p(s(bool,X73)) | p(s(bool,X74))) )),
  inference(forward_demodulation,[],[f5922,f313])).
fof(f5922,plain,(
  ( ! [X74,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X73))))),s(bool,X74))))) )),
  inference(superposition,[],[f293,f254])).
fof(f5947,plain,(
  ( ! [X50,X49] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X50))),s(bool,X49)))) | p(s(bool,X50)) | s(bool,'F_0') = s(bool,X49)) )),
  inference(forward_demodulation,[],[f5946,f313])).
fof(f5946,plain,(
  ( ! [X50,X49] : (p(s(bool,X50)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X49))))),s(bool,X50)))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f5909,f226])).
fof(f5909,plain,(
  ( ! [X50,X49] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X49))))),s(bool,X50)))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f293,f388])).
fof(f5931,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5893,f226])).
fof(f5893,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f293,f354])).
fof(f5891,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f226,f228,f293])).
fof(f5794,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5687,f226])).
fof(f5687,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f290,f354])).
fof(f5789,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,happ(s(fun('type.num.num',bool),X86),s('type.num.num',X87)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X88),s(fun('type.num.num',bool),X86))),s('type.num.num',X87)))) | p(s(bool,X88))) )),
  inference(forward_demodulation,[],[f5788,f268])).
fof(f5788,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X88),s(fun('type.num.num',bool),X86))),s('type.num.num',X87)))) | p(s(bool,X88)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.SUC_1'(s('type.num.num',X87))))))) )),
  inference(forward_demodulation,[],[f5681,f312])).
fof(f5681,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X88))),s(bool,happ(s(fun('type.num.num',bool),X86),s('type.num.num',X87)))))) | p(s(bool,X88)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.SUC_1'(s('type.num.num',X87))))))) )),
  inference(superposition,[],[f290,f268])).
fof(f5759,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X42))),s(bool,X41)))) | p(s(bool,X42)) | s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f5661,f228])).
fof(f5661,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X42))),s(bool,X41)))) | p(s(bool,X42)) | p(s(bool,'F_0')) | s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f290,f2467])).
fof(f5749,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f5748])).
fof(f5748,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5649,f230])).
fof(f5649,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f290,f354])).
fof(f5747,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(condensation,[],[f5746])).
fof(f5746,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f5647,f230])).
fof(f5647,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f290,f229])).
fof(f5741,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,happ(s(fun('type.num.num',bool),X86),s('type.num.num',X87)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X86),s(bool,X88))),s('type.num.num',X87)))) | p(s(bool,X88))) )),
  inference(forward_demodulation,[],[f5740,f268])).
fof(f5740,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X86),s(bool,X88))),s('type.num.num',X87)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.SUC_1'(s('type.num.num',X87)))))) | p(s(bool,X88))) )),
  inference(forward_demodulation,[],[f5642,f313])).
fof(f5642,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),X86),s('type.num.num',X87))))),s(bool,X88)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.SUC_1'(s('type.num.num',X87)))))) | p(s(bool,X88))) )),
  inference(superposition,[],[f290,f268])).
fof(f5711,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X41))),s(bool,X42)))) | p(s(bool,X42)) | s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f5622,f228])).
fof(f5622,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X41))),s(bool,X42)))) | p(s(bool,'F_0')) | p(s(bool,X42)) | s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f290,f2467])).
fof(f5701,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f5700])).
fof(f5700,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5610,f230])).
fof(f5610,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f290,f354])).
fof(f5699,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(condensation,[],[f5698])).
fof(f5698,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f5608,f230])).
fof(f5608,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f290,f229])).
fof(f5607,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f228,f228,f290])).
fof(f5453,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK8(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5431,f226])).
fof(f5431,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK8(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f288,f354])).
fof(f5450,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun(X27,bool),X29),s(X27,sK8('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29)),X27))))) | ~p(s(bool,happ(s(fun(fun(X27,bool),bool),'?_1'),s(fun(X27,bool),X29))))) )),
  inference(forward_demodulation,[],[f5428,f260])).
fof(f5428,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun(fun(X27,bool),bool),'?_1'),s(fun(X27,bool),X29)))) | p(s(bool,happ(s(fun(X27,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29))),s(X27,sK8('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29)),X27)))))) )),
  inference(superposition,[],[f288,f260])).
fof(f5449,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK8('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26)),X24))))) | ~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X25))))) )),
  inference(forward_demodulation,[],[f5427,f259])).
fof(f5427,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X25)))) | p(s(bool,happ(s(fun(X24,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26))),s(X24,sK8('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26)),X24)))))) )),
  inference(superposition,[],[f288,f259])).
fof(f5448,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK8('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | ~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23))))) )),
  inference(forward_demodulation,[],[f5426,f299])).
fof(f5426,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23)))) | p(s(bool,happ(s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))),s(X20,sK8('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(superposition,[],[f288,f299])).
fof(f5447,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | ~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X18))))) )),
  inference(forward_demodulation,[],[f5425,f300])).
fof(f5425,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X18)))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))),s(X16,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(superposition,[],[f288,f300])).
fof(f5446,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK8('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))))))) )),
  inference(forward_demodulation,[],[f5424,f268])).
fof(f5424,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | p(s(bool,happ(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s(X12,sK8('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(superposition,[],[f288,f268])).
fof(f5445,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK8('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10))))) )),
  inference(forward_demodulation,[],[f5423,f261])).
fof(f5423,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10)))) | p(s(bool,happ(s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))),s(X9,sK8('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(superposition,[],[f288,f261])).
fof(f5444,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK8(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(forward_demodulation,[],[f5422,f296])).
fof(f5422,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK8(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(superposition,[],[f288,f296])).
fof(f5443,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK8(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | ~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4))))) )),
  inference(forward_demodulation,[],[f5421,f298])).
fof(f5421,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4)))) | p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))),s(X2,sK8(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(superposition,[],[f288,f298])).
fof(f5442,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f5420,f254])).
fof(f5420,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))),s(X0,sK8(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(superposition,[],[f288,f254])).
fof(f5328,plain,(
  ( ! [X37] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),sK7('const.pair.FST_0','type.pair.prod'(bool,X37)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X37),bool),bool),'!_1'),s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f287,f296])).
fof(f5327,plain,(
  ( ! [X35,X36] : (~p(s(bool,X36)) | p(s(bool,happ(s(fun(fun(X35,bool),bool),'!_1'),s(fun(X35,bool),'_dst_x0_1'(s(bool,X36))))))) )),
  inference(superposition,[],[f287,f298])).
fof(f5346,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5315,f226])).
fof(f5315,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f287,f354])).
fof(f5311,plain,(
  ( ! [X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X19)))) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK7(X19,bool))) )),
  inference(superposition,[],[f287,f2467])).
fof(f5309,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X15)))) | s(bool,X16) = s(bool,sK7(X15,bool)) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f287,f2467])).
fof(f5303,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK7(X3,bool))) )),
  inference(superposition,[],[f287,f354])).
fof(f5302,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK7(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f287,f354])).
fof(f5301,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK7(X0,bool))) )),
  inference(superposition,[],[f287,f229])).
fof(f5337,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun(fun(X27,bool),bool),'!_1'),s(fun(X27,bool),X29)))) | ~p(s(bool,happ(s(fun(X27,bool),X29),s(X27,sK7('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29)),X27)))))) )),
  inference(forward_demodulation,[],[f5300,f260])).
fof(f5300,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun(X27,bool),X29),s(X27,sK7('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29)),X27))))) | p(s(bool,happ(s(fun(fun(X27,bool),bool),'!_1'),s(fun(X27,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X27,bool),X28),s(fun(X27,bool),X29))))))) )),
  inference(superposition,[],[f287,f260])).
fof(f5336,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),X25)))) | ~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK7('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26)),X24)))))) )),
  inference(forward_demodulation,[],[f5299,f259])).
fof(f5299,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK7('const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26)),X24))))) | p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X24,bool),X25),s(fun(X24,bool),X26))))))) )),
  inference(superposition,[],[f287,f259])).
fof(f5335,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),X23)))) | ~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK7('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(forward_demodulation,[],[f5298,f299])).
fof(f5298,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK7('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))))))) )),
  inference(superposition,[],[f287,f299])).
fof(f5334,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f5297,f300])).
fof(f5297,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f287,f300])).
fof(f5333,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f5296,f268])).
fof(f5296,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))))) )),
  inference(superposition,[],[f287,f268])).
fof(f5332,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f5295,f261])).
fof(f5295,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f287,f261])).
fof(f5331,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | ~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f5294,f296])).
fof(f5294,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(superposition,[],[f287,f296])).
fof(f5330,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(forward_demodulation,[],[f5293,f298])).
fof(f5293,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))))))) )),
  inference(superposition,[],[f287,f298])).
fof(f5329,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f5292,f254])).
fof(f5292,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f287,f254])).
fof(f4862,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.wordSem.word_loc'(X35),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X35),X37))) != s('type.wordSem.word_loc'(X35),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X35),X38))) | s('type.fcp.cart'(bool,X35),X38) = s('type.fcp.cart'(bool,X35),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s('type.fcp.cart'(bool,X35),X36),s('type.fcp.cart'(bool,X35),X37)))) )),
  inference(superposition,[],[f266,f260])).
fof(f4861,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.wordSem.word_loc'(X31),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X31),X32))) != s('type.wordSem.word_loc'(X31),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X31),X34))) | s('type.fcp.cart'(bool,X31),X34) = s('type.fcp.cart'(bool,X31),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s('type.fcp.cart'(bool,X31),X32),s('type.fcp.cart'(bool,X31),X33)))) )),
  inference(superposition,[],[f266,f259])).
fof(f4860,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.wordSem.word_loc'(X26),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X26),X29))) != s('type.wordSem.word_loc'(X26),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X26),X30))) | s('type.fcp.cart'(bool,X26),X30) = s('type.fcp.cart'(bool,X26),'const.pair.SND_1'(s('type.pair.prod'(X27,'type.fcp.cart'(bool,X26)),'const.pair.,_2'(s(X27,X28),s('type.fcp.cart'(bool,X26),X29)))))) )),
  inference(superposition,[],[f266,f299])).
fof(f4859,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.wordSem.word_loc'(X21),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X21),X23))) != s('type.wordSem.word_loc'(X21),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X21),X25))) | s('type.fcp.cart'(bool,X21),X25) = s('type.fcp.cart'(bool,X21),'const.pair.FST_1'(s('type.pair.prod'('type.fcp.cart'(bool,X21),X22),'const.pair.,_2'(s('type.fcp.cart'(bool,X21),X23),s(X22,X24)))))) )),
  inference(superposition,[],[f266,f300])).
fof(f4858,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.wordSem.word_loc'(X16),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X16),X20))) != s('type.wordSem.word_loc'(X16),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X16),happ(s(fun('type.num.num','type.fcp.cart'(bool,X16)),X18),s('type.num.num',X19))))) | s('type.fcp.cart'(bool,X16),X20) = s('type.fcp.cart'(bool,X16),'const.ind_type.FCONS_3'(s('type.fcp.cart'(bool,X16),X17),s(fun('type.num.num','type.fcp.cart'(bool,X16)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f266,f268])).
fof(f4857,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.wordSem.word_loc'(X12),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X12),X13))) != s('type.wordSem.word_loc'(X12),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X12),X15))) | s('type.fcp.cart'(bool,X12),X15) = s('type.fcp.cart'(bool,X12),'const.ind_type.FCONS_3'(s('type.fcp.cart'(bool,X12),X13),s(fun('type.num.num','type.fcp.cart'(bool,X12)),X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f266,f261])).
fof(f4856,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.wordSem.word_loc'(X8),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X8),X11))) != s('type.wordSem.word_loc'(X8),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X8),'const.pair.FST_1'(s('type.pair.prod'('type.fcp.cart'(bool,X8),X9),X10))))) | s('type.fcp.cart'(bool,X8),X11) = s('type.fcp.cart'(bool,X8),happ(s(fun('type.pair.prod'('type.fcp.cart'(bool,X8),X9),'type.fcp.cart'(bool,X8)),'const.pair.FST_0'),s('type.pair.prod'('type.fcp.cart'(bool,X8),X9),X10)))) )),
  inference(superposition,[],[f266,f296])).
fof(f4855,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.wordSem.word_loc'(X3),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X3),X5))) != s('type.wordSem.word_loc'(X3),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X3),X7))) | s('type.fcp.cart'(bool,X3),X7) = s('type.fcp.cart'(bool,X3),happ(s(fun(X4,'type.fcp.cart'(bool,X3)),'_dst_x0_1'(s('type.fcp.cart'(bool,X3),X5))),s(X4,X6)))) )),
  inference(superposition,[],[f266,f298])).
fof(f4854,plain,(
  ( ! [X2,X0,X1] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) != s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))) | s('type.fcp.cart'(bool,X0),X2) = s('type.fcp.cart'(bool,X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.fcp.cart'(bool,X0)),'_i_0_0'),s('type.fcp.cart'(bool,X0),X1)))) )),
  inference(superposition,[],[f266,f254])).
fof(f4697,plain,(
  ( ! [X33,X31,X34,X32] : (s(X31,X32) != s(X31,X33) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X32))),s(X31,X33))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4558,f228])).
fof(f4558,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'F_0')) | s(X31,X32) != s(X31,X33) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X32))),s(X31,X33))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4690,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(subsumption_resolution,[],[f4550,f228])).
fof(f4550,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(superposition,[],[f307,f229])).
fof(f4535,plain,(
  ( ! [X87,X86] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X87))),s(bool,X86)))) | s(bool,'F_0') != s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f307,f229])).
fof(f4533,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X81))),s(bool,X80)))) | s(bool,'F_0') != s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4531,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X75))),s(bool,X74)))) | s(bool,'T_0') != s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f307,f354])).
fof(f4529,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X70))),s(bool,X69)))) | s(bool,X68) != s(bool,X70) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4527,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X63))),s(bool,'F_0')))) | s(bool,X62) != s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4521,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X46))),s(bool,X45)))) | s(bool,X44) != s(bool,X46) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f307,f354])).
fof(f4520,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,X41) != s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f307,f354])).
fof(f4519,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,'F_0')))) | s(bool,X39) != s(bool,X40) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f307,f229])).
fof(f4668,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f4667,f268])).
fof(f4667,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f4514,f312])).
fof(f4514,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f307,f268])).
fof(f4495,plain,(
  ( ! [X87,X86] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X86))),s(bool,X87)))) | s(bool,'F_0') != s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f307,f229])).
fof(f4493,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X80))),s(bool,X81)))) | s(bool,'F_0') != s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4491,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X74))),s(bool,X75)))) | s(bool,'T_0') != s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f307,f354])).
fof(f4489,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X70)))) | s(bool,X68) != s(bool,X70) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4487,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X63)))) | s(bool,X62) != s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f307,f2467])).
fof(f4481,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,X46)))) | s(bool,X44) != s(bool,X46) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f307,f354])).
fof(f4480,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X41) != s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f307,f354])).
fof(f4479,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X40)))) | s(bool,X39) != s(bool,X40) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f307,f229])).
fof(f4641,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f4640,f268])).
fof(f4640,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f4474,f313])).
fof(f4474,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f307,f268])).
fof(f4635,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f4432,f312])).
fof(f4432,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f254,f307])).
fof(f4627,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f4626,f268])).
fof(f4626,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f4440,f313])).
fof(f4440,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) )),
  inference(unit_resulting_resolution,[],[f268,f307])).
fof(f4619,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.fcp.cart'(bool,X0),bool),'_dst_x0x0x1_3'(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'))),s('type.fcp.cart'(bool,X0),X1))))) )),
  inference(forward_demodulation,[],[f4447,f312])).
fof(f4447,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(X0),bool),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))),s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f255,f307])).
fof(f4594,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f4593,f310])).
fof(f4593,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f4592,f296])).
fof(f4592,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f4458,f310])).
fof(f4458,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(unit_resulting_resolution,[],[f296,f307])).
fof(f4591,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f4590,f312])).
fof(f4590,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f4589,f310])).
fof(f4589,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f4588,f268])).
fof(f4588,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) )),
  inference(forward_demodulation,[],[f4459,f310])).
fof(f4459,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(unit_resulting_resolution,[],[f268,f307])).
fof(f4571,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(X0),bool),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))))) )),
  inference(forward_demodulation,[],[f4570,f310])).
fof(f4570,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(X0),bool),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))))) )),
  inference(forward_demodulation,[],[f4569,f255])).
fof(f4569,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(X0),bool),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))),s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1))))))) )),
  inference(forward_demodulation,[],[f4466,f310])).
fof(f4466,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(X0),bool),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),bool)),'=_0'),s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1))))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f255,f307])).
fof(f4567,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f4566,f310])).
fof(f4566,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f4565,f280])).
fof(f4565,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.bvp_space.op_space_req_2'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X0))),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f4468,f310])).
fof(f4468,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvp_space.op_space_req_2'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X0))),s('type.num.num',X1))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f280,f307])).
fof(f4564,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f4563,f310])).
fof(f4563,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f4562,f262])).
fof(f4562,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))))))))) )),
  inference(forward_demodulation,[],[f4469,f310])).
fof(f4469,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f262,f307])).
fof(f4370,plain,(
  ( ! [X6,X4,X5,X3] : (s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f4301,f226])).
fof(f4301,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f306,f354])).
fof(f4277,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X63))),s(bool,'F_0')))) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f306,f2467])).
fof(f4270,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'T_0')))) | s(bool,X41) = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f306,f354])).
fof(f4269,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,'F_0')))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f306,f229])).
fof(f4348,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f4347,f268])).
fof(f4347,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f4264,f312])).
fof(f4264,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f306,f268])).
fof(f4245,plain,(
  ( ! [X87,X86] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X86))),s(bool,X87)))) | s(bool,'F_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f306,f229])).
fof(f4243,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X80))),s(bool,X81)))) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f306,f2467])).
fof(f4241,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X74))),s(bool,X75)))) | s(bool,'T_0') = s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f306,f354])).
fof(f4239,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X70)))) | s(bool,X68) = s(bool,X70) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f306,f2467])).
fof(f4237,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X63)))) | s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f306,f2467])).
fof(f4231,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X45))),s(bool,X46)))) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f306,f354])).
fof(f4230,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X41) = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f306,f354])).
fof(f4229,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X40)))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f306,f229])).
fof(f4321,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f4320,f268])).
fof(f4320,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f4224,f313])).
fof(f4224,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f306,f268])).
fof(f3810,plain,(
  ( ! [X225,X226,X224] : (~p(s(bool,happ(s(fun(bool,bool),X225),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X225)))) | s(bool,X226) = s(bool,sK5(X224,X225,bool)) | s(bool,'F_0') = s(bool,X226)) )),
  inference(superposition,[],[f265,f2467])).
fof(f4205,plain,(
  ( ! [X111,X112,X113] : (~p(s(bool,X112)) | s(bool,X113) = s(bool,'|dollar|exists_1'(s(fun(X111,bool),'_dst_x0_1'(s(bool,X112))))) | s(bool,'F_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f3774,f228])).
fof(f3774,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'F_0')) | ~p(s(bool,X112)) | s(bool,X113) = s(bool,'|dollar|exists_1'(s(fun(X111,bool),'_dst_x0_1'(s(bool,X112))))) | s(bool,'F_0') = s(bool,X113)) )),
  inference(superposition,[],[f837,f2467])).
fof(f4204,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') = s(bool,X109) | s(bool,X110) = s(bool,'|dollar|exists_1'(s(fun(X108,bool),'_dst_x0_1'(s(bool,X109))))) | s(bool,'F_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f3773,f228])).
fof(f3773,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X109) | s(bool,X110) = s(bool,'|dollar|exists_1'(s(fun(X108,bool),'_dst_x0_1'(s(bool,X109))))) | s(bool,'F_0') = s(bool,X110)) )),
  inference(superposition,[],[f2645,f2467])).
fof(f4203,plain,(
  ( ! [X107,X105,X106,X104] : (s(X104,X106) = s(X104,sK5(X106,X105,X104)) | s(bool,X107) = s(bool,'|dollar|exists_1'(s(fun(X104,bool),X105))) | s(bool,'F_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f3772,f228])).
fof(f3772,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,'F_0')) | s(X104,X106) = s(X104,sK5(X106,X105,X104)) | s(bool,X107) = s(bool,'|dollar|exists_1'(s(fun(X104,bool),X105))) | s(bool,'F_0') = s(bool,X107)) )),
  inference(superposition,[],[f264,f2467])).
fof(f4202,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,X94)) | s(bool,X96) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X94))),s(bool,X95))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f3768,f228])).
fof(f3768,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,'F_0')) | p(s(bool,X94)) | s(bool,X96) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X94))),s(bool,X95))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f294,f2467])).
fof(f4197,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f3738,f228])).
fof(f3738,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f282,f2467])).
fof(f4196,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X33))),s(bool,X34)))) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f3703,f228])).
fof(f3703,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X33))),s(bool,X34)))) | p(s(bool,'F_0')) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f294,f2467])).
fof(f3692,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f339,f2467])).
fof(f3691,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f234,f2467])).
fof(f3606,plain,(
  ( ! [X239,X240,X238] : (~p(s(bool,happ(s(fun(bool,bool),X239),s(bool,X240)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X239)))) | s(bool,'F_0') = s(bool,X240) | s(bool,'F_0') = s(bool,sK5(X238,X239,bool))) )),
  inference(superposition,[],[f265,f2467])).
fof(f3485,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') = s(bool,X82) | s(bool,X83) = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X81),s(bool,X82))) | s(bool,'F_0') = s(bool,X83)) )),
  inference(superposition,[],[f2467,f260])).
fof(f3484,plain,(
  ( ! [X80,X78,X79] : (s(bool,'F_0') = s(bool,X78) | s(bool,X80) = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X78),s(bool,X79))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f2467,f259])).
fof(f3483,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'F_0') = s(bool,X76) | s(bool,X77) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76))))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f2467,f299])).
fof(f3482,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X71) | s(bool,X73) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72))))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f2467,f300])).
fof(f3481,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X67),s('type.num.num',X68))) | s(bool,X69) = s(bool,'const.ind_type.FCONS_3'(s(bool,X66),s(fun('type.num.num',bool),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f2467,f268])).
fof(f3480,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X63) | s(bool,X65) = s(bool,'const.ind_type.FCONS_3'(s(bool,X63),s(fun('type.num.num',bool),X64),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f2467,f261])).
fof(f3479,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),X61))) | s(bool,X62) = s(bool,happ(s(fun('type.pair.prod'(bool,X60),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X60),X61))) | s(bool,'F_0') = s(bool,X62)) )),
  inference(superposition,[],[f2467,f296])).
fof(f3478,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') = s(bool,X57) | s(bool,X59) = s(bool,happ(s(fun(X56,bool),'_dst_x0_1'(s(bool,X57))),s(X56,X58))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f2467,f298])).
fof(f3448,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f2467,f2467])).
fof(f3412,plain,(
  ( ! [X94,X95,X96] : (s(bool,X95) = s(bool,X96) | s(bool,'F_0') = s(bool,X96) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X94),s(bool,X95)))) )),
  inference(superposition,[],[f2467,f260])).
fof(f3411,plain,(
  ( ! [X92,X93,X91] : (s(bool,X91) = s(bool,X93) | s(bool,'F_0') = s(bool,X93) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X91),s(bool,X92)))) )),
  inference(superposition,[],[f2467,f259])).
fof(f3410,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X89) = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X87,bool),'const.pair.,_2'(s(X87,X88),s(bool,X89)))))) )),
  inference(superposition,[],[f2467,f299])).
fof(f3409,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X83),'const.pair.,_2'(s(bool,X84),s(X83,X85)))))) )),
  inference(superposition,[],[f2467,f300])).
fof(f3408,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X82) = s(bool,happ(s(fun('type.num.num',bool),X80),s('type.num.num',X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X79),s(fun('type.num.num',bool),X80),s('type.num.num','const.num.SUC_1'(s('type.num.num',X81)))))) )),
  inference(superposition,[],[f2467,f268])).
fof(f3407,plain,(
  ( ! [X78,X76,X77] : (s(bool,X76) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2467,f261])).
fof(f3406,plain,(
  ( ! [X74,X75,X73] : (s(bool,X75) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X73),X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X73),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X73),X74)))) )),
  inference(superposition,[],[f2467,f296])).
fof(f3405,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X70) = s(bool,X72) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,X70))),s(X69,X71)))) )),
  inference(superposition,[],[f2467,f298])).
fof(f3404,plain,(
  ( ! [X68,X67] : (s(bool,X67) = s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X67)))) )),
  inference(superposition,[],[f2467,f254])).
fof(f3280,plain,(
  ( ! [X218,X216,X215,X217,X214] : (p(s(bool,'|dollar|exists_1'(s(fun(X214,bool),happ(s(fun('type.num.num',fun(X214,bool)),X216),s('type.num.num',X217)))))) | s(X214,X218) = s(X214,sK5(X218,'const.ind_type.FCONS_3'(s(fun(X214,bool),X215),s(fun('type.num.num',fun(X214,bool)),X216),s('type.num.num','const.num.SUC_1'(s('type.num.num',X217)))),X214))) )),
  inference(superposition,[],[f264,f268])).
fof(f3319,plain,(
  ( ! [X212,X210,X213,X211,X209] : (p(s(bool,'|dollar|exists_1'(s(fun(X209,bool),happ(s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num',X212)))))) | ~p(s(bool,happ(s(fun(X209,bool),happ(s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num',X212))),s(X209,sK5(X213,'const.ind_type.FCONS_3'(s(fun(X209,bool),X210),s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num','const.num.SUC_1'(s('type.num.num',X212)))),X209)))))) )),
  inference(forward_demodulation,[],[f3279,f268])).
fof(f3279,plain,(
  ( ! [X212,X210,X213,X211,X209] : (~p(s(bool,happ(s(fun(X209,bool),happ(s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num',X212))),s(X209,sK5(X213,'const.ind_type.FCONS_3'(s(fun(X209,bool),X210),s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num','const.num.SUC_1'(s('type.num.num',X212)))),X209))))) | p(s(bool,'|dollar|exists_1'(s(fun(X209,bool),'const.ind_type.FCONS_3'(s(fun(X209,bool),X210),s(fun('type.num.num',fun(X209,bool)),X211),s('type.num.num','const.num.SUC_1'(s('type.num.num',X212))))))))) )),
  inference(superposition,[],[f265,f268])).
fof(f3308,plain,(
  ( ! [X158,X159,X160] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X159),s('type.num.num',X160))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X158),s(fun('type.num.num',bool),X159),s('type.num.num','const.num.SUC_1'(s('type.num.num',X160)))))) )),
  inference(forward_demodulation,[],[f3267,f253])).
fof(f3267,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,'const.bool.DATATYPE_1'(s(X161,X162))) != s(bool,happ(s(fun('type.num.num',bool),X159),s('type.num.num',X160))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X158),s(fun('type.num.num',bool),X159),s('type.num.num','const.num.SUC_1'(s('type.num.num',X160)))))) )),
  inference(superposition,[],[f772,f268])).
fof(f3265,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X152),s('type.num.num',X153))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X151),s(fun('type.num.num',bool),X152),s('type.num.num','const.num.SUC_1'(s('type.num.num',X153)))))) )),
  inference(superposition,[],[f388,f268])).
fof(f3304,plain,(
  ( ! [X121,X122,X120] : (p(s(bool,happ(s(fun('type.num.num',bool),X120),s('type.num.num',X121)))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X120),s(bool,X122))),s('type.num.num',X121))))) )),
  inference(forward_demodulation,[],[f3303,f268])).
fof(f3303,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X120),s(bool,X122))),s('type.num.num',X121)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.SUC_1'(s('type.num.num',X121))))))) )),
  inference(forward_demodulation,[],[f3257,f313])).
fof(f3257,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),X120),s('type.num.num',X121))))),s(bool,X122)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.SUC_1'(s('type.num.num',X121))))))) )),
  inference(superposition,[],[f294,f268])).
fof(f3225,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f268,f354])).
fof(f3224,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,X10) = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f268,f354])).
fof(f3223,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f268,f229])).
fof(f3017,plain,(
  ( ! [X37,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X36),X37)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X36),bool),bool),'?_1'),s(fun('type.pair.prod'(bool,X36),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f289,f296])).
fof(f3016,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,happ(s(fun(fun(X33,bool),bool),'?_1'),s(fun(X33,bool),'_dst_x0_1'(s(bool,X34))))))) )),
  inference(superposition,[],[f289,f298])).
fof(f3034,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3008,f226])).
fof(f3008,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f289,f354])).
fof(f2676,plain,(
  ( ! [X24,X23,X27] : (p(s(bool,happ(s(fun(X23,bool),X24),s(X23,X27)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X23,bool),bool),'!_1'),s(fun(X23,bool),X24)))) )),
  inference(subsumption_resolution,[],[f2675,f226])).
fof(f2675,plain,(
  ( ! [X24,X23,X27] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X23,bool),X24),s(X23,X27)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X23,bool),bool),'!_1'),s(fun(X23,bool),X24)))) )),
  inference(forward_demodulation,[],[f2664,f253])).
fof(f2664,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X25,X26)))) | p(s(bool,happ(s(fun(X23,bool),X24),s(X23,X27)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X23,bool),bool),'!_1'),s(fun(X23,bool),X24)))) )),
  inference(superposition,[],[f286,f2405])).
fof(f2673,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2659,f226])).
fof(f2659,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f286,f354])).
fof(f2645,plain,(
  ( ! [X88,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),'_dst_x0_1'(s(bool,X88)))))) | s(bool,'F_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f2644,f226])).
fof(f2644,plain,(
  ( ! [X88,X89] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),'_dst_x0_1'(s(bool,X88)))))) | s(bool,'F_0') = s(bool,X88)) )),
  inference(forward_demodulation,[],[f2561,f253])).
fof(f2561,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),'_dst_x0_1'(s(bool,X88)))))) | ~p(s(bool,'const.bool.DATATYPE_1'(s(X86,X87)))) | s(bool,'F_0') = s(bool,X88)) )),
  inference(superposition,[],[f837,f2405])).
fof(f2627,plain,(
  ( ! [X261,X260] : (~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X261)))) | s(bool,'F_0') = s(bool,sK5(X260,X261,bool))) )),
  inference(forward_demodulation,[],[f2537,f253])).
fof(f2537,plain,(
  ( ! [X261,X263,X260,X262] : (~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,'const.bool.DATATYPE_1'(s(X262,X263)))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X261)))) | s(bool,'F_0') = s(bool,sK5(X260,X261,bool))) )),
  inference(superposition,[],[f265,f2405])).
fof(f2621,plain,(
  ( ! [X118,X120,X119] : (p(s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119)))) | s(bool,'F_0') = s(bool,happ(s(fun(X118,bool),X119),s(X118,sK5(X120,X119,X118))))) )),
  inference(subsumption_resolution,[],[f2620,f226])).
fof(f2620,plain,(
  ( ! [X118,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119)))) | s(bool,'F_0') = s(bool,happ(s(fun(X118,bool),X119),s(X118,sK5(X120,X119,X118))))) )),
  inference(forward_demodulation,[],[f2502,f253])).
fof(f2502,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X121,X122)))) | p(s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119)))) | s(bool,'F_0') = s(bool,happ(s(fun(X118,bool),X119),s(X118,sK5(X120,X119,X118))))) )),
  inference(superposition,[],[f265,f2405])).
fof(f2613,plain,(
  ( ! [X66,X63] : (s(bool,X63) != s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2612,f226])).
fof(f2612,plain,(
  ( ! [X66,X63] : (~p(s(bool,'T_0')) | s(bool,X63) != s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X63)) )),
  inference(forward_demodulation,[],[f2493,f253])).
fof(f2493,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X64,X65)))) | s(bool,X63) != s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f357,f2405])).
fof(f2611,plain,(
  ( ! [X43,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X46))),s(bool,X43)))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2610,f226])).
fof(f2610,plain,(
  ( ! [X43,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X46))),s(bool,X43)))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(forward_demodulation,[],[f2488,f253])).
fof(f2488,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X44,X45)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X46))),s(bool,X43)))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f295,f2405])).
fof(f2609,plain,(
  ( ! [X31,X34] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X34))),s(bool,X31)))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f2608,f226])).
fof(f2608,plain,(
  ( ! [X31,X34] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X34))),s(bool,X31)))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(forward_demodulation,[],[f2485,f253])).
fof(f2485,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X32,X33)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X34))),s(bool,X31)))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f292,f2405])).
fof(f2607,plain,(
  ( ! [X30,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X27))),s(bool,X30)))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2606,f226])).
fof(f2606,plain,(
  ( ! [X30,X27] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X27))),s(bool,X30)))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f2484,f253])).
fof(f2484,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X28,X29)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X27))),s(bool,X30)))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f291,f2405])).
fof(f2467,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f2405,f2405])).
fof(f2436,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,X92) = s(bool,'const.bool.DATATYPE_1'(s(X93,X94))) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X91),s(bool,X92)))) )),
  inference(superposition,[],[f2405,f260])).
fof(f2435,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) = s(bool,'const.bool.DATATYPE_1'(s(X89,X90))) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X87),s(bool,X88)))) )),
  inference(superposition,[],[f2405,f259])).
fof(f2434,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,X84) = s(bool,'const.bool.DATATYPE_1'(s(X85,X86))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X82,bool),'const.pair.,_2'(s(X82,X83),s(bool,X84)))))) )),
  inference(superposition,[],[f2405,f299])).
fof(f2433,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) = s(bool,'const.bool.DATATYPE_1'(s(X80,X81))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X77),'const.pair.,_2'(s(bool,X78),s(X77,X79)))))) )),
  inference(superposition,[],[f2405,f300])).
fof(f2432,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X73) = s(bool,'const.bool.DATATYPE_1'(s(X75,X76))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X73),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2405,f261])).
fof(f2594,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(forward_demodulation,[],[f2431,f253])).
fof(f2431,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'const.bool.DATATYPE_1'(s(X71,X72))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f2405,f296])).
fof(f2430,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,X65) = s(bool,'const.bool.DATATYPE_1'(s(X67,X68))) | s(bool,'F_0') = s(bool,happ(s(fun(X64,bool),'_dst_x0_1'(s(bool,X65))),s(X64,X66)))) )),
  inference(superposition,[],[f2405,f298])).
fof(f2429,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,'const.bool.DATATYPE_1'(s(X62,X63))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X61)))) )),
  inference(superposition,[],[f2405,f254])).
fof(f2405,plain,(
  ( ! [X30,X31,X32] : (s(bool,X32) = s(bool,'const.bool.DATATYPE_1'(s(X30,X31))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2383,f226])).
fof(f2383,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X32) = s(bool,'const.bool.DATATYPE_1'(s(X30,X31))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f2322,f253])).
fof(f2404,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X28,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f2382,f226])).
fof(f2382,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X28,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f2322,f1119])).
fof(f2402,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2380,f226])).
fof(f2380,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2322,f388])).
fof(f2396,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f2371,f226])).
fof(f2371,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2322,f354])).
fof(f2333,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),'_dst_x0_1'(s(bool,X38))))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f2332,f254])).
fof(f2332,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))))))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2289,f226])).
fof(f2289,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))))))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f1088,f388])).
fof(f2330,plain,(
  ( ! [X28,X29] : (s(bool,X29) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2284,f226])).
fof(f2284,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | s(bool,X29) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f357,f388])).
fof(f2329,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X21),s(fun(bool,bool),'_i_0_0'))),s(bool,X20)))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(forward_demodulation,[],[f2328,f312])).
fof(f2328,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X21))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20)))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2281,f226])).
fof(f2281,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X21))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20)))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f295,f388])).
fof(f2326,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15),s(fun(bool,bool),'_i_0_0'))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f2325,f312])).
fof(f2325,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X14)))))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f2278,f226])).
fof(f2278,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X14)))))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f292,f388])).
fof(f2324,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X13))),s(bool,X12)))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2323,f313])).
fof(f2323,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))))),s(bool,X13)))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2277,f226])).
fof(f2277,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))))),s(bool,X13)))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f291,f388])).
fof(f2322,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2321,f254])).
fof(f2321,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2276,f226])).
fof(f2276,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f285,f388])).
fof(f2320,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X39),s(bool,X40)))) )),
  inference(forward_demodulation,[],[f2262,f254])).
fof(f2262,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X40))) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X39),s(bool,X40)))) )),
  inference(superposition,[],[f388,f260])).
fof(f2319,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X37),s(bool,X38)))) )),
  inference(forward_demodulation,[],[f2261,f254])).
fof(f2261,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X37))) | s(bool,'F_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X37),s(bool,X38)))) )),
  inference(superposition,[],[f388,f259])).
fof(f2318,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(forward_demodulation,[],[f2260,f254])).
fof(f2260,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(superposition,[],[f388,f299])).
fof(f2317,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2259,f254])).
fof(f2259,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f388,f300])).
fof(f2316,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2258,f254])).
fof(f2258,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f388,f261])).
fof(f2257,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X27),X28)))) )),
  inference(superposition,[],[f388,f296])).
fof(f2315,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),'_dst_x0_1'(s(bool,X25))),s(X24,X26)))) )),
  inference(forward_demodulation,[],[f2256,f254])).
fof(f2256,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),'_dst_x0_1'(s(bool,X25))),s(X24,X26)))) )),
  inference(superposition,[],[f388,f298])).
fof(f1860,plain,(
  ( ! [X167,X169,X168] : (~p(s(bool,happ(s(fun(bool,bool),X168),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) | s(bool,X169) = s(bool,sK5(X167,X168,bool)) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f265,f354])).
fof(f2229,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66)))) | s(bool,X68) = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK5(X67,X66,X65)))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f1829,f226])).
fof(f1829,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66)))) | s(bool,X68) = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK5(X67,X66,X65)))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f265,f354])).
fof(f2227,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0_1'(s(bool,X49))))) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1821,f226])).
fof(f1821,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0_1'(s(bool,X49))))) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f1088,f354])).
fof(f2224,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f1816,f226])).
fof(f1816,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f357,f354])).
fof(f2223,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X32))),s(bool,X31)))) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1815,f226])).
fof(f1815,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X32))),s(bool,X31)))) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f295,f354])).
fof(f2222,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X23))),s(bool,X22)))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1812,f226])).
fof(f1812,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X23))),s(bool,X22)))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f292,f354])).
fof(f2221,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X19))),s(bool,X20)))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1811,f226])).
fof(f1811,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X19))),s(bool,X20)))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f291,f354])).
fof(f2220,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1810,f226])).
fof(f1810,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f285,f354])).
fof(f1789,plain,(
  ( ! [X57,X58] : (s(bool,X57) = s(bool,'const.bool.DATATYPE_1'(s(bool,X57))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f253,f354])).
fof(f2217,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),'_dst_x0_1'(s(bool,X48)))))) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1787,f226])).
fof(f1787,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),'_dst_x0_1'(s(bool,X48)))))) | ~p(s(bool,'T_0')) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f837,f354])).
fof(f1772,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1119,f354])).
fof(f1710,plain,(
  ( ! [X173,X174,X172] : (~p(s(bool,happ(s(fun(bool,bool),X173),s(bool,X174)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173)))) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,sK5(X172,X173,bool))) )),
  inference(superposition,[],[f265,f354])).
fof(f1628,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X48),s(bool,X49))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f354,f260])).
fof(f1627,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X45),s(bool,X46))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f354,f259])).
fof(f1626,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43))))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f354,f299])).
fof(f1625,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') = s(bool,X38) | s(bool,X40) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,X38),s(X37,X39))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f354,f300])).
fof(f1624,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,X36) = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f354,f261])).
fof(f1623,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32))) | s(bool,X33) = s(bool,happ(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X31),X32))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f354,f296])).
fof(f1622,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),'_dst_x0_1'(s(bool,X28))),s(X27,X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f354,f298])).
fof(f1621,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f354,f254])).
fof(f1608,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f354,f354])).
fof(f1585,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X57),s(bool,X58)))) )),
  inference(superposition,[],[f354,f260])).
fof(f1584,plain,(
  ( ! [X54,X56,X55] : (s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X54),s(bool,X55)))) )),
  inference(superposition,[],[f354,f259])).
fof(f1583,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X51),s(bool,X52)))))) )),
  inference(superposition,[],[f354,f299])).
fof(f1582,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,X47),s(X46,X48)))))) )),
  inference(superposition,[],[f354,f300])).
fof(f1581,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X43),s(fun('type.num.num',bool),X44),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f354,f261])).
fof(f1580,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X40),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X40),X41)))) )),
  inference(superposition,[],[f354,f296])).
fof(f1579,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,X37))),s(X36,X38)))) )),
  inference(superposition,[],[f354,f298])).
fof(f1578,plain,(
  ( ! [X35,X34] : (s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X34)))) )),
  inference(superposition,[],[f354,f254])).
fof(f1437,plain,(
  ( ! [X19,X18] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X19),s(fun(bool,bool),'_i_0_0'))),s(bool,X18)))) | ~p(s(bool,X18))) )),
  inference(forward_demodulation,[],[f1422,f312])).
fof(f1422,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X19))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18))))))) )),
  inference(superposition,[],[f295,f254])).
fof(f1413,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f226,f295])).
fof(f1236,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(subsumption_resolution,[],[f1221,f228])).
fof(f1221,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(superposition,[],[f294,f229])).
fof(f1224,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1194,f228])).
fof(f1194,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f294,f229])).
fof(f1192,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f228,f294])).
fof(f1190,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f1166,f226])).
fof(f1166,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f357,f1119])).
fof(f1119,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f226,f1064,f285])).
fof(f1120,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f1064,f1064,f285])).
fof(f1121,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f1064,f291])).
fof(f1122,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f1064,f292])).
fof(f1125,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))))) )),
  inference(unit_resulting_resolution,[],[f254,f1064,f357])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f298,f1064,f357])).
fof(f1127,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f261,f1064,f357])).
fof(f1128,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(X0,X2)))))) )),
  inference(unit_resulting_resolution,[],[f300,f1064,f357])).
fof(f1129,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f299,f1064,f357])).
fof(f1130,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f259,f1064,f357])).
fof(f1131,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X0),s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0')))))))) )),
  inference(unit_resulting_resolution,[],[f260,f1064,f357])).
fof(f1132,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f1064,f837])).
fof(f1107,plain,(
  ( ! [X10,X11] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X11),s(fun(bool,bool),'_i_0_0'))),s(bool,X10)))) | ~p(s(bool,X10))) )),
  inference(forward_demodulation,[],[f1095,f312])).
fof(f1095,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X11))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))))))) )),
  inference(superposition,[],[f292,f254])).
fof(f1089,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f226,f292])).
fof(f1088,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) )),
  inference(subsumption_resolution,[],[f1078,f228])).
fof(f1078,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) )),
  inference(superposition,[],[f837,f229])).
fof(f1064,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f226,f837])).
fof(f1021,plain,(
  ( ! [X10,X9] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X10))),s(bool,X9)))) | ~p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f1011,f313])).
fof(f1011,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))))),s(bool,X10))))) )),
  inference(superposition,[],[f291,f254])).
fof(f1006,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f226,f291])).
fof(f931,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,'|dollar|exists_1'(s(fun(X96,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X96,bool),X97),X98)))))) | s(X96,X99) = s(X96,sK5(X99,happ(s(fun('type.pair.prod'(fun(X96,bool),X97),fun(X96,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X96,bool),X97),X98)),X96))) )),
  inference(superposition,[],[f264,f296])).
fof(f943,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'|dollar|exists_1'(s(fun(X92,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X92,bool),X93),X94)))))) | ~p(s(bool,happ(s(fun(X92,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X92,bool),X93),X94))),s(X92,sK5(X95,happ(s(fun('type.pair.prod'(fun(X92,bool),X93),fun(X92,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X92,bool),X93),X94)),X92)))))) )),
  inference(forward_demodulation,[],[f930,f296])).
fof(f930,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,happ(s(fun(X92,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X92,bool),X93),X94))),s(X92,sK5(X95,happ(s(fun('type.pair.prod'(fun(X92,bool),X93),fun(X92,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X92,bool),X93),X94)),X92))))) | p(s(bool,'|dollar|exists_1'(s(fun(X92,bool),happ(s(fun('type.pair.prod'(fun(X92,bool),X93),fun(X92,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X92,bool),X93),X94))))))) )),
  inference(superposition,[],[f265,f296])).
fof(f942,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X84),X85))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X84),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X84),X85)))) )),
  inference(forward_demodulation,[],[f928,f253])).
fof(f928,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'const.bool.DATATYPE_1'(s(X86,X87))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X84),X85))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X84),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X84),X85)))) )),
  inference(superposition,[],[f772,f296])).
fof(f918,plain,(
  ( ! [X61,X60] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),sK5(X61,'const.pair.FST_0','type.pair.prod'(bool,X60)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X60),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f265,f296])).
fof(f903,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4)))) )),
  inference(superposition,[],[f296,f229])).
fof(f837,plain,(
  ( ! [X4,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f265,f298])).
fof(f834,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f265,f229])).
fof(f844,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK5(X29,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X26,bool),X27),s(fun(X26,bool),X28)),X26)))))) )),
  inference(forward_demodulation,[],[f833,f260])).
fof(f833,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK5(X29,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X26,bool),X27),s(fun(X26,bool),X28)),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X26,bool),X27),s(fun(X26,bool),X28))))))) )),
  inference(superposition,[],[f265,f260])).
fof(f843,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) | ~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK5(X25,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X22,bool),X23),s(fun(X22,bool),X24)),X22)))))) )),
  inference(forward_demodulation,[],[f832,f259])).
fof(f832,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK5(X25,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X22,bool),X23),s(fun(X22,bool),X24)),X22))))) | p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X22,bool),X23),s(fun(X22,bool),X24))))))) )),
  inference(superposition,[],[f265,f259])).
fof(f842,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | ~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK5(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f831,f299])).
fof(f831,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK5(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20))))))))) )),
  inference(superposition,[],[f265,f299])).
fof(f841,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK5(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f830,f300])).
fof(f830,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK5(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f265,f300])).
fof(f840,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f829,f261])).
fof(f829,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f265,f261])).
fof(f839,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | ~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3)))))) )),
  inference(forward_demodulation,[],[f828,f298])).
fof(f828,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6))))))) )),
  inference(superposition,[],[f265,f298])).
fof(f838,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f827,f254])).
fof(f827,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f265,f254])).
fof(f797,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,'const.bool.DATATYPE_1'(s(X38,X39))) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X36),s(bool,X37)))) )),
  inference(superposition,[],[f772,f260])).
fof(f796,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) != s(bool,'const.bool.DATATYPE_1'(s(X34,X35))) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f772,f259])).
fof(f795,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X29) != s(bool,'const.bool.DATATYPE_1'(s(X30,X31))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29)))))) )),
  inference(superposition,[],[f772,f299])).
fof(f794,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X23) != s(bool,'const.bool.DATATYPE_1'(s(X25,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X23),s(X22,X24)))))) )),
  inference(superposition,[],[f772,f300])).
fof(f793,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X18) != s(bool,'const.bool.DATATYPE_1'(s(X20,X21))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X18),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f772,f261])).
fof(f792,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X14) != s(bool,'const.bool.DATATYPE_1'(s(X16,X17))) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),'_dst_x0_1'(s(bool,X14))),s(X13,X15)))) )),
  inference(superposition,[],[f772,f298])).
fof(f791,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,'const.bool.DATATYPE_1'(s(X11,X12))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10)))) )),
  inference(superposition,[],[f772,f254])).
fof(f772,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f755,f226])).
fof(f755,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f357,f253])).
fof(f738,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f737,f228])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f264,f229])).
fof(f736,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | s(X26,X29) = s(X26,sK5(X29,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(fun(X26,bool),X27),s(fun(X26,bool),X28)),X26))) )),
  inference(superposition,[],[f264,f260])).
fof(f735,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) | s(X22,X25) = s(X22,sK5(X25,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(fun(X22,bool),X23),s(fun(X22,bool),X24)),X22))) )),
  inference(superposition,[],[f264,f259])).
fof(f734,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | s(X17,X21) = s(X17,sK5(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) )),
  inference(superposition,[],[f264,f299])).
fof(f733,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | s(X12,X16) = s(X12,sK5(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) )),
  inference(superposition,[],[f264,f300])).
fof(f732,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | s(X8,X11) = s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) )),
  inference(superposition,[],[f264,f261])).
fof(f731,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X7) = s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f264,f298])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f264,f254])).
fof(f698,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6)))))) )),
  inference(superposition,[],[f300,f229])).
fof(f640,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) )),
  inference(superposition,[],[f299,f229])).
fof(f586,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)))) )),
  inference(superposition,[],[f298,f229])).
fof(f536,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f261,f229])).
fof(f482,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W64_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f260,f229])).
fof(f444,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ast.word_size_CASE_3'(s('type.ast.word_size','const.ast.W8_0'),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f259,f229])).
fof(f388,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f254,f229])).
fof(f354,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f229,f229])).
fof(f319,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X16,X40,X18,X17,X38,X41] : ! [X19,X20,X39,X21] : s(fun(X17,fun(X38,X16)),'_dst_x00x00x11_3'(s(fun(X40,fun(X41,X16)),X19),s(X40,happ(s(fun(X18,X40),X20),s(X18,X21))),s(fun(X17,fun(X38,X41)),X39))) = s(fun(X17,fun(X38,X16)),happ(s(fun(X18,fun(X17,fun(X38,X16))),'_dst_x000x100x011_3'(s(fun(X40,fun(X41,X16)),X19),s(fun(X18,X40),X20),s(fun(X17,fun(X38,X41)),X39))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x000x100x011_ax')).
fof(f318,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X16,X38,X17,X40,X18] : ! [X19,X20,X39,X21] : s(fun(X17,X16),'_dst_x0x1x1_3'(s(fun(X38,fun(X40,X16)),X19),s(fun(X17,X38),X20),s(fun(X17,X40),happ(s(fun(X18,fun(X17,X40)),X39),s(X18,X21))))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x01x11_3'(s(fun(X38,fun(X40,X16)),X19),s(fun(X17,X38),X20),s(fun(X18,fun(X17,X40)),X39))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x00x01x11_ax')).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X16,X38,X17,X40,X18] : ! [X19,X20,X39,X21] : s(fun(X17,X16),'_dst_x0x0x1_3'(s(fun(X38,fun(X40,X16)),X19),s(X38,X20),s(fun(X17,X40),happ(s(fun(X18,fun(X17,X40)),X39),s(X18,X21))))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x00x11_3'(s(fun(X38,fun(X40,X16)),X19),s(X38,X20),s(fun(X18,fun(X17,X40)),X39))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x00x00x11_ax')).
fof(f316,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X16,X17,X38,X18,X40] : ! [X19,X20,X39,X21] : s(fun(X17,X16),'_dst_x0x1x0_3'(s(fun(X38,fun(X40,X16)),X19),s(fun(X17,X38),happ(s(fun(X18,fun(X17,X38)),X20),s(X18,X21))),s(X40,X39))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x11x00_3'(s(fun(X38,fun(X40,X16)),X19),s(fun(X18,fun(X17,X38)),X20),s(X40,X39))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x00x11x00_ax')).
fof(f315,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X16,X38,X18,X17,X40] : ! [X19,X20,X39,X21] : s(fun(X17,X16),'_dst_x0x0x1_3'(s(fun(X38,fun(X40,X16)),X19),s(X38,happ(s(fun(X18,X38),X20),s(X18,X21))),s(fun(X17,X40),X39))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x10x01_3'(s(fun(X38,fun(X40,X16)),X19),s(fun(X18,X38),X20),s(fun(X17,X40),X39))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x00x10x01_ax')).
fof(f314,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X16,X17,X38,X18] : ! [X19,X20,X39,X21] : s(X16,happ(s(fun(X18,X16),'_dst_x0x1x1_3'(s(fun(X17,fun(X38,X16)),X19),s(fun(X18,X17),X20),s(fun(X18,X38),X39))),s(X18,X21))) = s(X16,happ(s(fun(X38,X16),happ(s(fun(X17,fun(X38,X16)),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))),s(X38,happ(s(fun(X18,X38),X39),s(X18,X21)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x0x1x1_ax')).
fof(f313,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X16,X17,X18,X38] : ! [X19,X20,X39,X21] : s(X16,happ(s(fun(X18,X16),'_dst_x0x1x0_3'(s(fun(X17,fun(X38,X16)),X19),s(fun(X18,X17),X20),s(X38,X39))),s(X18,X21))) = s(X16,happ(s(fun(X38,X16),happ(s(fun(X17,fun(X38,X16)),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))),s(X38,X39)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x0x1x0_ax')).
fof(f311,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X16,X17,X38,X18] : ! [X19,X20,X21] : s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X38,X16),X19),s(fun(X17,X38),happ(s(fun(X18,fun(X17,X38)),X20),s(X18,X21))))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x11_2'(s(fun(X38,X16),X19),s(fun(X18,fun(X17,X38)),X20))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x00x11_ax')).
fof(f331,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),'_dst_x1x0_2'(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(fun('type.num.num',X0),X2))),s(X0,X1))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f310,f272])).
fof(f330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),X2),s(fun('type.fcp.cart'(bool,X0),X1),X3),s(fun('type.num.num',fun('type.num.num',X1)),X4))) = s(X1,happ(s(fun(fun('type.num.num',fun('type.num.num',X1)),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num',fun('type.num.num',X1)),X1)),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1))),'const.wordSem.word_loc_CASE_0'),s(fun('type.fcp.cart'(bool,X0),X1),X3))),s('type.wordSem.word_loc'(X0),X2))),s(fun('type.num.num',fun('type.num.num',X1)),X4)))) )),
  inference(backward_demodulation,[],[f310,f304])).
fof(f329,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(X1,fun(fun('type.num.num',X0),X0)),'_dst_x1x0_2'(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(X1,X5))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(backward_demodulation,[],[f310,f297])).
fof(f309,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))) = s(X16,happ(s(fun(X18,X16),'_dst_x0x1_2'(s(fun(X17,X16),X19),s(fun(X18,X17),X20))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dst_x0x1_ax')).
fof(f308,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X16,X17,X18] : ! [X20,X21] : s(fun(fun(X17,X16),X16),'_dstsel_0_x0_1'(s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))) = s(fun(fun(X17,X16),X16),happ(s(fun(X18,fun(fun(X17,X16),X16)),'_dstsel_1_x10_1'(s(fun(X18,X17),X20))),s(X18,X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dstsel_1_x10_ax')).
fof(f304,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),X2),s(fun('type.fcp.cart'(bool,X0),X1),X3),s(fun('type.num.num',fun('type.num.num',X1)),X4))) = s(X1,happ(s(fun(fun('type.num.num',fun('type.num.num',X1)),X1),happ(s(fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1)),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1))),'const.wordSem.word_loc_CASE_0'),s('type.wordSem.word_loc'(X0),X2))),s(fun('type.fcp.cart'(bool,X0),X1),X3))),s(fun('type.num.num',fun('type.num.num',X1)),X4)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),X2),s(fun('type.fcp.cart'(bool,X0),X1),X3),s(fun('type.num.num',fun('type.num.num',X1)),X4))) = s(X1,happ(s(fun(fun('type.num.num',fun('type.num.num',X1)),X1),happ(s(fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1)),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1))),'const.wordSem.word_loc_CASE_0'),s('type.wordSem.word_loc'(X0),X2))),s(fun('type.fcp.cart'(bool,X0),X1),X3))),s(fun('type.num.num',fun('type.num.num',X1)),X4)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X0),X2),s(fun('type.fcp.cart'(bool,X0),X1),X3),s(fun('type.num.num',fun('type.num.num',X1)),X4))) = s(X1,happ(s(fun(fun('type.num.num',fun('type.num.num',X1)),X1),happ(s(fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1)),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.fcp.cart'(bool,X0),X1),fun(fun('type.num.num',fun('type.num.num',X1)),X1))),'const.wordSem.word_loc_CASE_0'),s('type.wordSem.word_loc'(X0),X2))),s(fun('type.fcp.cart'(bool,X0),X1),X3))),s(fun('type.num.num',fun('type.num.num',X1)),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X14] : ! [X11,X12,X24] : s(X14,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X6),X11),s(fun('type.fcp.cart'(bool,X6),X14),X12),s(fun('type.num.num',fun('type.num.num',X14)),X24))) = s(X14,happ(s(fun(fun('type.num.num',fun('type.num.num',X14)),X14),happ(s(fun(fun('type.fcp.cart'(bool,X6),X14),fun(fun('type.num.num',fun('type.num.num',X14)),X14)),happ(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.fcp.cart'(bool,X6),X14),fun(fun('type.num.num',fun('type.num.num',X14)),X14))),'const.wordSem.word_loc_CASE_0'),s('type.wordSem.word_loc'(X6),X11))),s(fun('type.fcp.cart'(bool,X6),X14),X12))),s(fun('type.num.num',fun('type.num.num',X14)),X24)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_3_const.wordSem.word_loc_CASE')).
fof(f303,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.fcp.cart'(bool,X1),X0),X3),s('type.fcp.cart'(bool,X1),X2))) = s(X0,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X1),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X1),X2))),s(fun('type.fcp.cart'(bool,X1),X0),X3),s(fun('type.num.num',fun('type.num.num',X0)),X4)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.fcp.cart'(bool,X1),X0),X3),s('type.fcp.cart'(bool,X1),X2))) = s(X0,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X1),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X1),X2))),s(fun('type.fcp.cart'(bool,X1),X0),X3),s(fun('type.num.num',fun('type.num.num',X0)),X4)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun('type.fcp.cart'(bool,X1),X0),X3),s('type.fcp.cart'(bool,X1),X2))) = s(X0,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X1),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X1),X2))),s(fun('type.fcp.cart'(bool,X1),X0),X3),s(fun('type.num.num',fun('type.num.num',X0)),X4)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14,X6] : ! [X25,X13,X15] : s(X14,happ(s(fun('type.fcp.cart'(bool,X6),X14),X13),s('type.fcp.cart'(bool,X6),X25))) = s(X14,'const.wordSem.word_loc_CASE_3'(s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X25))),s(fun('type.fcp.cart'(bool,X6),X14),X13),s(fun('type.num.num',fun('type.num.num',X14)),X15)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_case_def|split|0')).
fof(f302,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.wordSem.word_loc_|49||49||split|1',aHL_EXT)).
fof(f301,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X16,X18] : ! [X20,X21] : s(X16,happ(s(fun(X18,X16),X21),s(X18,X20))) = s(X16,happ(s(fun(fun(X18,X16),X16),'_dstsel_0_x0_1'(s(X18,X20))),s(fun(X18,X16),X21)))),
  file('thm.wordSem.word_loc_|49||49||split|1','_dstsel_0_x0_ax')).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X14,X6] : ! [X33,X34] : s(X6,X33) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X33),s(X14,X34)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.pair.FST')).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X14] : ! [X33,X34] : s(X14,X34) = s(X14,'const.pair.SND_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X33),s(X14,X34)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.pair.SND')).
fof(f297,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14,X6] : ! [X27] : ? [X13] : ! [X28,X29,X30] : s(X14,happ(s(fun('type.ind_type.recspace'(X6),X14),X13),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X28),s(X6,X29),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30))))) = s(X14,happ(s(fun(fun('type.num.num',X14),X14),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14)))),X27),s('type.num.num',X28))),s(X6,X29))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30))),s(fun('type.num.num',X14),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X14),X13),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.ind_type.CONSTR_REC')).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X14] : ! [X11] : s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),X11))) = s(X6,happ(s(fun('type.pair.prod'(X6,X14),X6),'const.pair.FST_0'),s('type.pair.prod'(X6,X14),X11)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_1_const.pair.FST')).
fof(f294,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f295,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f291,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f292,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f288,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0)))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X42,X45] : (p(s(bool,happ(s(fun(fun(X42,bool),bool),'?_1'),s(fun(X42,bool),X45)))) <=> ? [X43] : p(s(bool,happ(s(fun(X42,bool),X45),s(X42,X43)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','?_HL_QUANT')).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f281,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.bit.DIV_2EXP_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.num.num','const.bit.DIV_2EXP_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X33,X26] : s('type.num.num','const.bit.DIV_2EXP_2'(s('type.num.num',X33),s('type.num.num',X26))) = s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X33)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bit.DIV_2EXP_def')).
fof(f280,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bvp_space.op_space_req_2'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bvp_space.op_space_req_2'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X46,X47] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bvp_space.op_space_req_2'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X47))),s('type.num.num',X46)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bvp_space.op_space_req_def|split|5')).
fof(f279,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.fcp.cart'(bool,X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.fcp.cart'(bool,X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.fcp.cart'(bool,X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.fcp.cart'(bool,X0),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x1x0_2'(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.bool.ARB_0'))),s('type.num.num','const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',bool)),'_dst_x00x11_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',fun('type.num.num',bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.fcp.cart'(bool,X0),'const.bool.ARB_0'))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'))))))),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),sK6(X0)))))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.fcp.cart'(bool,X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.fcp.cart'(bool,X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.fcp.cart'(bool,X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.fcp.cart'(bool,X0),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x1x0_2'(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.bool.ARB_0'))),s('type.num.num','const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',bool)),'_dst_x00x11_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',fun('type.num.num',bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.fcp.cart'(bool,X0),'const.bool.ARB_0'))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'))))))),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),sK6(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f130])).
fof(f130,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.fcp.cart'(bool,X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.fcp.cart'(bool,X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.fcp.cart'(bool,X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.fcp.cart'(bool,X0),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x1x0_2'(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.bool.ARB_0'))),s('type.num.num','const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',bool)),'_dst_x00x11_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',fun('type.num.num',bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.fcp.cart'(bool,X0),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.fcp.cart'(bool,X0),'const.bool.ARB_0'))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'))))))),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X0),'type.pair.prod'('type.num.num','type.num.num')))),X1))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ? [X36] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.fcp.cart'(bool,X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.fcp.cart'(bool,X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.fcp.cart'(bool,X6),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.fcp.cart'(bool,X6),'type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x1x0_2'(s(fun('type.fcp.cart'(bool,X6),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.num.num','type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'),s('type.num.num','const.bool.ARB_0'))),s('type.num.num','const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),'_dst_x0x1_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',bool)),'_dst_x00x11_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.num.num',fun('type.num.num',bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'),s(fun('type.num.num',fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x00x11_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),happ(s(fun('type.fcp.cart'(bool,X6),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'const.pair.,_0'),s('type.fcp.cart'(bool,X6),'const.bool.ARB_0'))),s(fun('type.num.num',fun('type.num.num','type.pair.prod'('type.num.num','type.num.num'))),'const.pair.,_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'const.ind_type.BOTTOM_0'))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num'))),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),'_i_0_0'))))))),s(fun('type.wordSem.word_loc'(X6),'type.ind_type.recspace'('type.pair.prod'('type.fcp.cart'(bool,X6),'type.pair.prod'('type.num.num','type.num.num')))),X36))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_TY_DEF')).
fof(f278,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X41,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : ! [X25,X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,happ(s(fun('type.num.num',X6),X13),s('type.num.num',X25))) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Global_1'(s('type.num.num',X25))),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|0')).
fof(f277,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X41,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Cons_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Cons_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Cons_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X25,X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,happ(s(fun('type.num.num',X6),X51),s('type.num.num',X25))) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.Cons_1'(s('type.num.num',X25))),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|5')).
fof(f276,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X41,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobal_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobal_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41] : s(X0,happ(s(fun('type.num.num',X0),X3),s('type.num.num',X1))) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobal_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(fun('type.num.num',X0),X3),s(X0,X4),s(X0,X5),s(X0,X6),s(fun('type.num.num',X0),X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(X0,X15),s(fun('type.num.num',X0),X16),s(X0,X17),s(fun('type.num.num',fun('type.num.num',X0)),X18),s(fun('type.num.num',X0),X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(X0,X24),s(fun('type.num.num',X0),X25),s(fun('type.num.num',X0),X26),s(X0,X27),s(fun('type.integer.int',X0),X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(X0,X37),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X38),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X39),s(X0,X40),s(X0,X41)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X25,X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,happ(s(fun('type.num.num',X6),X15),s('type.num.num',X25))) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobal_1'(s('type.num.num',X25))),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|1')).
fof(f275,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,X3) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.AllocGlobal_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X3) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.AllocGlobal_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X3) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.AllocGlobal_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,X48) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.AllocGlobal_0'),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|2')).
fof(f274,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,X5) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X5) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X5) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,X50) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.SetGlobalsPtr_0'),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|4')).
fof(f273,plain,(
  ( ! [X28,X24,X37,X4,X33,X0,X12,X8,X21,X17,X29,X25,X38,X5,X34,X1,X13,X9,X22,X18,X30,X26,X39,X6,X35,X2,X14,X10,X23,X19,X31,X27,X7,X36,X3,X32,X15,X11,X40,X20,X16] : (s(X0,X4) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.GlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X4) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.GlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : s(X0,X4) = s(X0,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.GlobalsPtr_0'),s(fun('type.num.num',X0),X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(X0,X4),s(X0,X5),s(fun('type.num.num',X0),X6),s(X0,X7),s(X0,X8),s(X0,X9),s(X0,X10),s(X0,X11),s(X0,X12),s(X0,X13),s(X0,X14),s(fun('type.num.num',X0),X15),s(X0,X16),s(fun('type.num.num',fun('type.num.num',X0)),X17),s(fun('type.num.num',X0),X18),s(X0,X19),s(X0,X20),s(X0,X21),s(X0,X22),s(X0,X23),s(fun('type.num.num',X0),X24),s(fun('type.num.num',X0),X25),s(X0,X26),s(fun('type.integer.int',X0),X27),s(X0,X28),s(X0,X29),s(X0,X30),s(X0,X31),s(X0,X32),s(X0,X33),s(X0,X34),s(X0,X35),s(X0,X36),s(fun('type.ast.word_size',fun('type.ast.opw',X0)),X37),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X0))),X38),s(X0,X39),s(X0,X40)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X13,X15,X48,X49,X50,X51,X52,X53,X54,X55,X56,X46,X57,X47,X58,X59,X60,X61,X62,X63,X64,X65,X66,X67,X68,X69,X70,X71,X72,X73,X74,X75,X76,X77,X78,X79,X80,X81,X82,X83] : s(X6,X49) = s(X6,'const.closLang.op_CASE_41'(s('type.closLang.op','const.closLang.GlobalsPtr_0'),s(fun('type.num.num',X6),X13),s(fun('type.num.num',X6),X15),s(X6,X48),s(X6,X49),s(X6,X50),s(fun('type.num.num',X6),X51),s(X6,X52),s(X6,X53),s(X6,X54),s(X6,X55),s(X6,X56),s(X6,X46),s(X6,X57),s(X6,X47),s(fun('type.num.num',X6),X58),s(X6,X59),s(fun('type.num.num',fun('type.num.num',X6)),X60),s(fun('type.num.num',X6),X61),s(X6,X62),s(X6,X63),s(X6,X64),s(X6,X65),s(X6,X66),s(fun('type.num.num',X6),X67),s(fun('type.num.num',X6),X68),s(X6,X69),s(fun('type.integer.int',X6),X70),s(X6,X71),s(X6,X72),s(X6,X73),s(X6,X74),s(X6,X75),s(X6,X76),s(X6,X77),s(X6,X78),s(X6,X79),s(fun('type.ast.word_size',fun('type.ast.opw',X6)),X80),s(fun('type.ast.word_size',fun('type.ast.shift',fun('type.num.num',X6))),X81),s(X6,X82),s(X6,X83)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.closLang.op_case_def|split|3')).
fof(f272,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X11,X12,X24] : s(X6,'const.ind_type.FCONS_3'(s(X6,X11),s(fun('type.num.num',X6),X12),s('type.num.num',X24))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X11))),s(fun('type.num.num',X6),X12))),s('type.num.num',X24)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_3_const.ind_type.FCONS')).
fof(f271,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X11,X12,X24] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X11),s(X6,X12),s(fun('type.num.num','type.ind_type.recspace'(X6)),X24))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X11))),s(X6,X12))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X24)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_1_3_const.ind_type.CONSTR')).
fof(f269,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X13,X7,X8] : s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_size_def|split|1')).
fof(f268,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X25,X13,X26] : s(X6,happ(s(fun('type.num.num',X6),X13),s('type.num.num',X26))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X25),s(fun('type.num.num',X6),X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.ind_type.FCONS|split|1')).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) != s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) != s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)) & (s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2) | s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2)))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))) <=> s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2))) <=> s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X25,X35] : (s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X25))) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X35))) <=> s('type.fcp.cart'(bool,X6),X25) = s('type.fcp.cart'(bool,X6),X35))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_|49||49||split|0')).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2) | s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2)))) )),
  inference(cnf_transformation,[],[f208])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X31,X32] : (! [X33] : (s(X6,X32) = s(X6,X33) => p(s(bool,happ(s(fun(X6,bool),X31),s(X6,X33))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X31)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.JRH_INDUCT_UTIL')).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.num.num','type.wordSem.word_loc'(X0)),happ(s(fun('type.num.num',fun('type.num.num','type.wordSem.word_loc'(X0))),'const.wordSem.Loc_0'),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.num.num','type.wordSem.word_loc'(X0)),happ(s(fun('type.num.num',fun('type.num.num','type.wordSem.word_loc'(X0))),'const.wordSem.Loc_0'),s('type.num.num',X1))),s('type.num.num',X2)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.num.num','type.wordSem.word_loc'(X0)),happ(s(fun('type.num.num',fun('type.num.num','type.wordSem.word_loc'(X0))),'const.wordSem.Loc_0'),s('type.num.num',X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X11,X12] : s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X11),s('type.num.num',X12))) = s('type.wordSem.word_loc'(X6),happ(s(fun('type.num.num','type.wordSem.word_loc'(X6)),happ(s(fun('type.num.num',fun('type.num.num','type.wordSem.word_loc'(X6))),'const.wordSem.Loc_0'),s('type.num.num',X11))),s('type.num.num',X12)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_2_const.wordSem.Loc')).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2)))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X2)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X13,X25] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X25)))))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_size_def|split|0')).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X25,X13] : s(X6,X25) = s(X6,'const.ind_type.FCONS_3'(s(X6,X25),s(fun('type.num.num',X6),X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.ind_type.FCONS|split|0')).
fof(f255,plain,(
  ( ! [X0,X1] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X1))) = s('type.wordSem.word_loc'(X0),happ(s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X0),X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X11] : s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X11))) = s('type.wordSem.word_loc'(X6),happ(s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),'const.wordSem.Word_0'),s('type.fcp.cart'(bool,X6),X11)))),
  file('thm.wordSem.word_loc_|49||49||split|1','ARITY_0_1_const.wordSem.Word')).
fof(f253,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X33] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X6,X33)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.bool.DATATYPE_TAG_THM')).
fof(f6386,plain,(
  $false | ($spl2 | $spl22)),
  inference(global_subsumption,[],[f6385,f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374,f6376,f6383])).
fof(f6385,plain,(
  ( ! [X54] : (p(s(bool,X54))) ) | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f6218,f226])).
fof(f6218,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54))) ) | ($spl2 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5951])).
fof(f5951,plain,(
  ( ! [X54] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s(bool,X54))),s('type.num.num',sK1)))) | p(s(bool,X54))) ) | $spl2),
  inference(forward_demodulation,[],[f5950,f313])).
fof(f5950,plain,(
  ( ! [X54] : (p(s(bool,X54)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK1))))),s(bool,X54))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5913,f226])).
fof(f5913,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK1))))),s(bool,X54))))) ) | $spl2),
  inference(superposition,[],[f293,f4721])).
fof(f4721,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f339,f4568,f2322])).
fof(f4568,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK1))),s('type.num.num',sK1)))) | $spl2),
  inference(forward_demodulation,[],[f4467,f324])).
fof(f4467,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK3))),s('type.num.num',sK1)))) | $spl2),
  inference(unit_resulting_resolution,[],[f324,f307])).
fof(f6384,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374,f6376,f6383])).
fof(f6382,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6213,f226])).
fof(f6213,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5354])).
fof(f5354,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2))),s('type.num.num',sK7(happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK2)),'type.num.num'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f4376,f287])).
fof(f6381,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6198,f226])).
fof(f6198,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f4404])).
fof(f6380,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6197,f226])).
fof(f6197,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5388])).
fof(f6379,plain,(
  $false | ($spl12 | $spl22)),
  inference(global_subsumption,[],[f6378,f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374,f6376])).
fof(f6378,plain,(
  ( ! [X51] : (p(s(bool,X51))) ) | ($spl12 | $spl22)),
  inference(subsumption_resolution,[],[f6191,f226])).
fof(f6191,plain,(
  ( ! [X51] : (~p(s(bool,'T_0')) | p(s(bool,X51))) ) | ($spl12 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5949])).
fof(f5949,plain,(
  ( ! [X51] : (~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(bool,bool),bool),'?_1'),s(bool,X51))),s(fun(bool,bool),'_i_0_0')))) | p(s(bool,X51))) ) | $spl12),
  inference(forward_demodulation,[],[f5948,f313])).
fof(f5948,plain,(
  ( ! [X51] : (p(s(bool,X51)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X51))))) ) | $spl12),
  inference(subsumption_resolution,[],[f5910,f226])).
fof(f5910,plain,(
  ( ! [X51] : (~p(s(bool,'T_0')) | p(s(bool,X51)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X51))))) ) | $spl12),
  inference(superposition,[],[f293,f3060])).
fof(f3060,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | $spl12),
  inference(unit_resulting_resolution,[],[f339,f3033,f2322])).
fof(f6377,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374,f6376])).
fof(f6375,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372,f6374])).
fof(f6373,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360,f6372])).
fof(f6371,plain,(
  $false | ($spl8 | $spl22)),
  inference(global_subsumption,[],[f6370,f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360])).
fof(f6370,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f6156,f226])).
fof(f6156,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | ($spl8 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5954])).
fof(f5954,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl8),
  inference(forward_demodulation,[],[f5953,f993])).
fof(f993,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f957,f231])).
fof(f957,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(unit_resulting_resolution,[],[f853,f264])).
fof(f5953,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl8),
  inference(condensation,[],[f5952])).
fof(f5952,plain,(
  ( ! [X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))),s(bool,X56))))) ) | $spl8),
  inference(subsumption_resolution,[],[f5914,f234])).
fof(f5914,plain,(
  ( ! [X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))),s(bool,X56)))) | s(bool,'F_0') = s(bool,X55)) ) | $spl8),
  inference(superposition,[],[f293,f2403])).
fof(f2403,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X27)) ) | $spl8),
  inference(subsumption_resolution,[],[f2381,f226])).
fof(f2381,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X27)) ) | $spl8),
  inference(superposition,[],[f2322,f993])).
fof(f6361,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318,f6360])).
fof(f6325,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f6324,f226])).
fof(f6324,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(forward_demodulation,[],[f6106,f5942])).
fof(f6106,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl22),
  inference(backward_demodulation,[],[f5942,f5891])).
fof(f6323,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6322,f226])).
fof(f6322,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f6105,f5942])).
fof(f6105,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5880])).
fof(f5880,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | $spl5),
  inference(forward_demodulation,[],[f5879,f5874])).
fof(f5879,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))) | $spl5),
  inference(forward_demodulation,[],[f5856,f312])).
fof(f5856,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num')))))),s(bool,happ(s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))),s('type.num.num',sK7(sK7('!_1',fun('type.num.num',bool)),'type.num.num'))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5800,f5800,f290])).
fof(f6321,plain,(
  $false | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f6320,f226])).
fof(f6320,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22)),
  inference(forward_demodulation,[],[f6103,f5942])).
fof(f6103,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'),s(fun(fun('type.num.num',bool),bool),'!_1'))),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | ($spl5 | $spl22)),
  inference(backward_demodulation,[],[f5942,f5835])).
fof(f5835,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun(fun('type.num.num',bool),bool),'!_1'))),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | $spl5),
  inference(forward_demodulation,[],[f5834,f5829])).
fof(f5834,plain,(
  ~p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))),s(fun(fun('type.num.num',bool),bool),'!_1'))),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))) | $spl5),
  inference(forward_demodulation,[],[f5811,f312])).
fof(f5811,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool))))))),s(bool,happ(s(fun(fun('type.num.num',bool),bool),'!_1'),s(fun('type.num.num',bool),sK7('!_1',fun('type.num.num',bool)))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f5387,f5387,f290])).
fof(f6319,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315,f6318])).
fof(f6316,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298,f6315])).
fof(f6313,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f6312,f226])).
fof(f6312,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(forward_demodulation,[],[f6096,f5942])).
fof(f6096,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl22),
  inference(backward_demodulation,[],[f5942,f5607])).
fof(f6303,plain,(
  $false | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f6302,f226])).
fof(f6302,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl22)),
  inference(forward_demodulation,[],[f6061,f5942])).
fof(f6061,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl8 | $spl22)),
  inference(backward_demodulation,[],[f5942,f4315])).
fof(f4315,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl8),
  inference(forward_demodulation,[],[f4216,f852])).
fof(f4216,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,sK5(X0,'_i_0_0',bool)))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f853,f306])).
fof(f6301,plain,(
  $false | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f6300,f226])).
fof(f6300,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl22)),
  inference(forward_demodulation,[],[f6060,f5942])).
fof(f6060,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl8 | $spl22)),
  inference(backward_demodulation,[],[f5942,f4314])).
fof(f4314,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl8),
  inference(forward_demodulation,[],[f4313,f310])).
fof(f4313,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl8),
  inference(forward_demodulation,[],[f4312,f852])).
fof(f4312,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,sK5(X0,'_i_0_0',bool)))))) ) | $spl8),
  inference(forward_demodulation,[],[f4219,f310])).
fof(f4219,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK5(X0,'_i_0_0',bool)))),s(bool,'T_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f853,f306])).
fof(f6299,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293,f6298])).
fof(f6294,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282,f6293])).
fof(f6283,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f225,f224,f223,f226,f228,f229,f231,f230,f235,f234,f253,f254,f255,f259,f260,f261,f262,f263,f265,f264,f267,f266,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f285,f282,f287,f286,f289,f288,f292,f291,f290,f295,f294,f293,f296,f297,f298,f299,f300,f301,f302,f303,f304,f305,f307,f306,f308,f309,f310,f329,f330,f331,f311,f312,f313,f314,f315,f316,f317,f318,f319,f339,f354,f357,f388,f413,f444,f482,f536,f586,f640,f698,f730,f731,f732,f733,f734,f735,f736,f738,f772,f791,f792,f793,f794,f795,f796,f797,f838,f839,f840,f841,f842,f843,f844,f834,f836,f837,f903,f918,f942,f943,f931,f1006,f1021,f1064,f1088,f1089,f1107,f1132,f1131,f1130,f1129,f1128,f1127,f1126,f1125,f1122,f1121,f1120,f1119,f1190,f1192,f1224,f1236,f1413,f1437,f1578,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1608,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1710,f1772,f2217,f1789,f2220,f2221,f2222,f2223,f2224,f2227,f2229,f1860,f2315,f2257,f2316,f2317,f2318,f2319,f2320,f2322,f2324,f2326,f2329,f2330,f2333,f2396,f2402,f2404,f2405,f2429,f2430,f2594,f2432,f2433,f2434,f2435,f2436,f2467,f2607,f2609,f2611,f2613,f2621,f2627,f2645,f2673,f2676,f3030,f3034,f3015,f3016,f3017,f3223,f3224,f3225,f3304,f3265,f3308,f3319,f3280,f3404,f3405,f3406,f3407,f3408,f3409,f3410,f3411,f3412,f3448,f3478,f3479,f3480,f3481,f3482,f3483,f3484,f3485,f3606,f3691,f3692,f4196,f4197,f4202,f4203,f4204,f4205,f3810,f4321,f4229,f4230,f4231,f4237,f4239,f4241,f4243,f4245,f4348,f4269,f4270,f4277,f4370,f4564,f4567,f4571,f4591,f4594,f4597,f4619,f4627,f4635,f4641,f4479,f4480,f4481,f4487,f4489,f4491,f4493,f4495,f4668,f4519,f4520,f4521,f4527,f4529,f4531,f4533,f4535,f4690,f4697,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f5329,f5330,f5331,f5332,f5333,f5334,f5335,f5336,f5337,f5301,f5302,f5303,f5309,f5311,f5341,f5346,f5326,f5327,f5328,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5450,f5453,f5607,f5699,f5701,f5711,f5741,f5747,f5749,f5759,f5789,f5794,f5891,f5931,f5941,f5945,f5947,f5966,f5942,f6282])).
fof(f6281,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f5975,f226])).
fof(f5975,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(backward_demodulation,[],[f5942,f228])).
fof(f5942,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(cnf_transformation,[],[f5942_D])).
fof(f5942_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5943,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl24),
  inference(cnf_transformation,[],[f5943_D])).
fof(f5943_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5945,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f5944])).
fof(f5944,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X39)))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5905,f357])).
fof(f5905,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X39)))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f293,f354])).
fof(f5941,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f5940])).
fof(f5940,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X36)))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f5904,f357])).
fof(f5904,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X36)))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f293,f354])).
fof(f5326,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f287,f254])).
fof(f5343,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | $spl16),
  inference(cnf_transformation,[],[f5343_D])).
fof(f5343_D,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5341,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7('_i_0_0',bool))),
  inference(subsumption_resolution,[],[f5313,f226])).
fof(f5313,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7('_i_0_0',bool))),
  inference(superposition,[],[f287,f388])).
fof(f3033,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl12),
  inference(cnf_transformation,[],[f3033_D])).
fof(f3033_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3038,plain,(
  ( ! [X32] : (~p(s(bool,X32))) ) | $spl14),
  inference(cnf_transformation,[],[f3038_D])).
fof(f3038_D,plain,(
  ( ! [X32] : (~p(s(bool,X32))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3015,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f289,f254])).
fof(f3031,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f3031_D])).
fof(f3031_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3030,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3006,f226])).
fof(f3006,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f289,f388])).
fof(f846,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(cnf_transformation,[],[f846_D])).
fof(f846_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f847,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool)))) ) | $spl8),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f836,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f265,f254])).
fof(f334,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f333])).
fof(f333,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f332,f324])).
fof(f332,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK2))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f323,f326])).
fof(f322,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) | $spl0),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f324,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | $spl2),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f223,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4)))),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  (s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) | (s('type.num.num',sK1) = s('type.num.num',sK3) & s('type.num.num',sK2) = s('type.num.num',sK4))) & (s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) | s('type.num.num',sK1) != s('type.num.num',sK3) | s('type.num.num',sK2) != s('type.num.num',sK4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f190])).
fof(f190,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X2) = s('type.num.num',X4))) & (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) != s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) | s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X2) != s('type.num.num',X4)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X2) = s('type.num.num',X4))) & (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) != s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) | (s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X2) != s('type.num.num',X4))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) <~> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X2) = s('type.num.num',X4)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) <=> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X2) = s('type.num.num',X4)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.wordSem.word_loc'(X0),'const.wordSem.Loc_2'(s('type.num.num',X3),s('type.num.num',X4))) <=> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X2) = s('type.num.num',X4)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X7),s('type.num.num',X8))) = s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X9),s('type.num.num',X10))) <=> (s('type.num.num',X7) = s('type.num.num',X9) & s('type.num.num',X8) = s('type.num.num',X10)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X7),s('type.num.num',X8))) = s('type.wordSem.word_loc'(X6),'const.wordSem.Loc_2'(s('type.num.num',X9),s('type.num.num',X10))) <=> (s('type.num.num',X7) = s('type.num.num',X9) & s('type.num.num',X8) = s('type.num.num',X10)))),
  file('thm.wordSem.word_loc_|49||49||split|1','thm.wordSem.word_loc_|49||49||split|1')).
fof(f326,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | $spl4),
  inference(cnf_transformation,[],[f326_D])).
fof(f326_D,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f224,plain,(
  s('type.num.num',sK2) = s('type.num.num',sK4) | s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4)))),
  inference(cnf_transformation,[],[f191])).
fof(f323,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) | $spl1),
  inference(cnf_transformation,[],[f323_D])).
fof(f323_D,plain,(
  s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f325,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK3) | $spl3),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f327,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK4) | $spl5),
  inference(cnf_transformation,[],[f327_D])).
fof(f327_D,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK4) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f225,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK4) | s('type.num.num',sK1) != s('type.num.num',sK3) | s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s('type.wordSem.word_loc'(sK0),'const.wordSem.Loc_2'(s('type.num.num',sK3),s('type.num.num',sK4)))),
  inference(cnf_transformation,[],[f191])).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 261
% Passive clauses: 3295
% Generated clauses: 72192
% Final active clauses: 102
% Final passive clauses: 1033
% Input formulas: 68
% Initial clauses: 97
% Discarded non-redundant clauses: 8095
% 
% Duplicate literals: 840
% Trivial inequalities: 91
% Fw subsumption resolutions: 1468
% Fw demodulations: 6983
% Bw demodulations: 949
% Condensations: 2859
% Global subsumptions: 14
% 
% Simple tautologies: 143
% Equational tautologies: 270
% Forward subsumptions: 50028
% Fw demodulations to eq. taut.: 58
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 2963
% Factoring: 1
% Forward superposition: 36916
% Backward superposition: 17938
% Self superposition: 501
% Equality factoring: 88
% Equality resolution: 15
% 
% Split clauses: 13
% Split components: 27
% SAT solver clauses: 3480
% SAT solver unit clauses: 672
% SAT solver binary clauses: 1667
% 
% Sat splits: 13
% Sat splitting refutations: 143
% 
% Memory used [KB]: 104262
% Time elapsed: 4.951 s
% ------------------------------
% ------------------------------
% Success in time 6.3 s
