Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ml_translator
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_translator
% SZS output start Proof for thm.ml_translator
fof(f26801,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f26800])).
fof(f26800,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num',sK12)))))),
  inference(forward_demodulation,[],[f26452,f516])).
fof(f516,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X7)))),
  file('thm.ml_translator.ml_translator_|49|95_prove_ltb|split|2','thm.arithmetic.ADD_SYM')).
fof(f26452,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK11)))))),
  inference(superposition,[],[f25288,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8,X7,X6] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X8))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.ml_translator.ml_translator_|49|95_prove_ltb|split|2','thm.arithmetic.LEFT_ADD_DISTRIB')).
fof(f25288,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))),
  inference(forward_demodulation,[],[f25287,f516])).
fof(f25287,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK11))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))),
  inference(forward_demodulation,[],[f394,f515])).
fof(f515,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X7] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X7)))),
  file('thm.ml_translator.ml_translator_|49|95_prove_ltb|split|2','thm.arithmetic.MULT_SYM')).
fof(f394,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK11))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK12),s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK11))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK12),s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f245])).
fof(f245,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.ml_translator.ml_translator_|49|95_prove_ltb|split|2','thm.ml_translator.ml_translator_|49|95_prove_ltb|split|2')).
% SZS output end Proof for thm.ml_translator
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 829
% Passive clauses: 4262
% Generated clauses: 26412
% Final active clauses: 802
% Final passive clauses: 3141
% Input formulas: 114
% Initial clauses: 277
% 
% Duplicate literals: 221
% Trivial inequalities: 52
% Fw subsumption resolutions: 1501
% Fw demodulations: 2088
% Bw demodulations: 8
% 
% Simple tautologies: 188
% Equational tautologies: 28
% Forward subsumptions: 17558
% Fw demodulations to eq. taut.: 19
% 
% Binary resolution: 5033
% Factoring: 1318
% Forward superposition: 7623
% Backward superposition: 8090
% Self superposition: 55
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 199
% Split components: 399
% SAT solver clauses: 376
% SAT solver unit clauses: 13
% SAT solver binary clauses: 66
% 
% Sat splits: 199
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 376
% TWLsolver calls for satisfiability: 332
% 
% Memory used [KB]: 14711
% Time elapsed: 0.470 s
% ------------------------------
% ------------------------------
% Success in time 0.48 s
