Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.evalProps
% (13305)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 4158
% Generated clauses: 15151
% Final active clauses: 92
% Final passive clauses: 3833
% Input formulas: 14
% Initial clauses: 25
% 
% Duplicate literals: 779
% Trivial inequalities: 97
% Fw subsumption resolutions: 1563
% Fw demodulations: 788
% Bw demodulations: 4
% 
% Simple tautologies: 244
% Equational tautologies: 82
% Forward subsumptions: 6414
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 187
% Factoring: 2
% Forward superposition: 6096
% Backward superposition: 5029
% Self superposition: 249
% Equality factoring: 178
% Equality resolution: 22
% 
% Split clauses: 402
% Split components: 826
% SAT solver clauses: 1081
% SAT solver unit clauses: 32
% SAT solver binary clauses: 219
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 2
% 
% Sat splits: 402
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1076
% TWLsolver calls for satisfiability: 147
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (13330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 165
% Passive clauses: 1471
% Generated clauses: 9749
% Final active clauses: 105
% Final passive clauses: 1003
% Input formulas: 14
% Initial clauses: 25
% 
% Duplicate literals: 226
% Trivial inequalities: 70
% Fw subsumption resolutions: 221
% Fw demodulations: 1250
% Bw demodulations: 161
% Condensations: 317
% 
% Simple tautologies: 52
% Equational tautologies: 26
% Forward subsumptions: 3893
% Fw demodulations to eq. taut.: 8
% Bw demodulations to eq. taut.: 161
% 
% Binary resolution: 275
% Factoring: 2
% Forward superposition: 2669
% Backward superposition: 2638
% Self superposition: 62
% Equality factoring: 38
% Equality resolution: 30
% 
% Split clauses: 72
% Split components: 149
% SAT solver clauses: 1997
% SAT solver unit clauses: 24
% SAT solver binary clauses: 44
% 
% Sat splits: 72
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1997
% TWLsolver calls for satisfiability: 264
% 
% Memory used [KB]: 16247
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f1385,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f54,f80,f79,f53,f82,f52,f84,f78,f90,f93,f91,f106,f110,f108,f115,f444,f445,f439,f433,f450,f452,f451,f453,f446,f107,f471,f473,f472,f589,f598,f599,f600,f601,f602,f1069,f1072,f1075,f1078,f1081,f1084,f1087,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1112,f1120,f1129,f1138,f1148,f1156,f1165,f1174,f1184,f1192,f1201,f1210,f1219,f1222,f1225,f1228,f1231,f1234,f1237,f1240,f1243,f1246,f1249,f1252,f1255,f1258,f1261,f1270,f1278,f1288,f1297,f1306,f1314,f1324,f1333,f1342,f1350,f1360])).
fof(f1360,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1359,f616])).
fof(f616,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(X0,X2)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f79,f72])).
fof(f72,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X6] : ! [X9,X16,X18,X19] : (s('type.pair.prod'(X8,X6),'const.pair.,_2'(s(X8,X9),s(X6,X16))) = s('type.pair.prod'(X8,X6),'const.pair.,_2'(s(X8,X18),s(X6,X19))) <=> (s(X8,X9) = s(X8,X18) & s(X6,X16) = s(X6,X19)))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f1359,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1358,f84])).
fof(f1358,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1357,f76])).
fof(f76,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8,X6,X7] : ! [X16,X9,X11] : s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X16))),s(fun(X6,fun(X7,X8)),X11))) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X11),s(X6,X9))),s(X7,X16)))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.pair.pair_case_thm')).
fof(f1357,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f844,f70])).
fof(f70,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X6] : ! [X9] : ? [X20,X21] : s('type.pair.prod'(X8,X6),X9) = s('type.pair.prod'(X8,X6),'const.pair.,_2'(s(X8,X20),s(X6,X21)))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f70,f74])).
fof(f74,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f1350,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1349,f616])).
fof(f1349,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1348,f84])).
fof(f1348,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f851,f76])).
fof(f851,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f70,f74])).
fof(f1342,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1341,f616])).
fof(f1341,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1340,f84])).
fof(f1340,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1339,f76])).
fof(f1339,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f858,f84])).
fof(f858,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f84,f74])).
fof(f1333,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1332,f616])).
fof(f1332,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1331,f84])).
fof(f1331,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f865,f76])).
fof(f865,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f84,f74])).
fof(f1324,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1323,f616])).
fof(f1323,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1322,f84])).
fof(f1322,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f1321,f76])).
fof(f1321,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f872,f472])).
fof(f872,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f472,f74])).
fof(f1314,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1313,f616])).
fof(f1313,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1312,f84])).
fof(f1312,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f879,f76])).
fof(f879,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f472,f74])).
fof(f1306,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1305,f616])).
fof(f1305,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1304,f84])).
fof(f1304,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1303,f76])).
fof(f1303,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f886,f626])).
fof(f626,plain,(
  s(sK0,sK6) = s(sK0,sK9(sK3,sK1,sK0)) | $spl6),
  inference(unit_resulting_resolution,[],[f70,f605])).
fof(f605,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X0),s(sK1,X1))) | s(sK0,sK6) = s(sK0,X0)) ) | $spl6),
  inference(superposition,[],[f72,f84])).
fof(f886,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK9(sK3,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f626,f74])).
fof(f1297,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1296,f616])).
fof(f1296,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1295,f84])).
fof(f1295,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f893,f76])).
fof(f893,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f626,f74])).
fof(f1288,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1287,f616])).
fof(f1287,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1286,f84])).
fof(f1286,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1285,f76])).
fof(f1285,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f900,f770])).
fof(f770,plain,(
  ( ! [X0] : (s(sK0,sK6) = s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0))) ) | ($spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f472,f628])).
fof(f628,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),X0) | s(sK0,sK6) = s(sK0,sK9(X0,sK1,sK0))) ) | $spl6),
  inference(superposition,[],[f605,f70])).
fof(f900,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f770,f74])).
fof(f1278,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1277,f616])).
fof(f1277,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1276,f84])).
fof(f1276,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f907,f76])).
fof(f907,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f770,f74])).
fof(f1270,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1269,f616])).
fof(f1269,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1268,f84])).
fof(f1268,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1267,f76])).
fof(f1267,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f914,f809])).
fof(f809,plain,(
  s(sK1,sK7) = s(sK1,sK10(sK3,sK1,sK0)) | $spl6),
  inference(unit_resulting_resolution,[],[f70,f662])).
fof(f662,plain,(
  ( ! [X6,X5] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X5),s(sK1,X6))) | s(sK1,sK7) = s(sK1,X6)) ) | $spl6),
  inference(superposition,[],[f73,f84])).
fof(f73,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f51])).
fof(f914,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK10(sK3,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f809,f74])).
fof(f1261,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1260,f616])).
fof(f1260,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1259,f84])).
fof(f1259,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f921,f76])).
fof(f921,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f809,f74])).
fof(f1258,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1257,f632])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,sK2),'const.pair.,_2'(s(X0,X1),s(sK2,sK4))) != s('type.pair.prod'(X0,sK2),'const.pair.,_2'(s(X0,X2),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f79,f73])).
fof(f1257,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1256,f84])).
fof(f1256,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | $spl4),
  inference(forward_demodulation,[],[f922,f76])).
fof(f922,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f70,f82,f74])).
fof(f1255,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1254,f632])).
fof(f1254,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1253,f84])).
fof(f1253,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f923,f76])).
fof(f923,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f84,f82,f74])).
fof(f1252,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1251,f632])).
fof(f1251,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1250,f84])).
fof(f1250,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f924,f76])).
fof(f924,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f472,f82,f74])).
fof(f1249,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1248,f632])).
fof(f1248,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1247,f84])).
fof(f1247,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f925,f76])).
fof(f925,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f626,f82,f74])).
fof(f1246,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1245,f632])).
fof(f1245,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1244,f84])).
fof(f1244,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f926,f76])).
fof(f926,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f770,f82,f74])).
fof(f1243,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1242,f632])).
fof(f1242,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1241,f84])).
fof(f1241,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f927,f76])).
fof(f927,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f809,f82,f74])).
fof(f1240,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1239,f616])).
fof(f1239,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1238,f84])).
fof(f1238,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | $spl4),
  inference(forward_demodulation,[],[f928,f76])).
fof(f928,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f82,f74])).
fof(f1237,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1236,f632])).
fof(f1236,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1235,f84])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | $spl4),
  inference(forward_demodulation,[],[f929,f76])).
fof(f929,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f70,f82,f74])).
fof(f1234,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1233,f632])).
fof(f1233,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1232,f84])).
fof(f1232,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f930,f76])).
fof(f930,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f84,f82,f74])).
fof(f1231,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1230,f632])).
fof(f1230,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1229,f84])).
fof(f1229,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f931,f76])).
fof(f931,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f472,f82,f74])).
fof(f1228,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1227,f632])).
fof(f1227,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1226,f84])).
fof(f1226,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f932,f76])).
fof(f932,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f626,f82,f74])).
fof(f1225,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1224,f632])).
fof(f1224,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1223,f84])).
fof(f1223,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f933,f76])).
fof(f933,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f770,f82,f74])).
fof(f1222,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1221,f632])).
fof(f1221,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1220,f84])).
fof(f1220,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f934,f76])).
fof(f934,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f809,f82,f74])).
fof(f1219,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1218,f632])).
fof(f1218,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1217,f84])).
fof(f1217,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | $spl4),
  inference(forward_demodulation,[],[f935,f76])).
fof(f935,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f82,f74])).
fof(f1210,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1209,f616])).
fof(f1209,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1208,f84])).
fof(f1208,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f942,f76])).
fof(f942,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f70,f74])).
fof(f1201,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1200,f616])).
fof(f1200,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1199,f84])).
fof(f1199,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1198,f76])).
fof(f1198,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f949,f70])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK2,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f70,f74])).
fof(f1192,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1191,f616])).
fof(f1191,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1190,f84])).
fof(f1190,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f956,f76])).
fof(f956,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f84,f74])).
fof(f1184,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1183,f616])).
fof(f1183,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1182,f84])).
fof(f1182,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1181,f76])).
fof(f1181,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f963,f84])).
fof(f963,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f84,f74])).
fof(f1174,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1173,f616])).
fof(f1173,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1172,f84])).
fof(f1172,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f970,f76])).
fof(f970,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3)))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f472,f74])).
fof(f1165,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1164,f616])).
fof(f1164,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1163,f84])).
fof(f1163,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f1162,f76])).
fof(f1162,plain,(
  s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK3))) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f977,f472])).
fof(f977,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,sK4),s('type.pair.prod'(sK0,sK1),sK3))) = s('type.pair.prod'(sK2,'type.pair.prod'(sK0,sK1)),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f472,f74])).
fof(f1156,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1155,f616])).
fof(f1155,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1154,f84])).
fof(f1154,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f984,f76])).
fof(f984,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f626,f74])).
fof(f1148,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1147,f616])).
fof(f1147,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1146,f84])).
fof(f1146,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1145,f76])).
fof(f1145,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f991,f626])).
fof(f991,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK9(sK3,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f626,f74])).
fof(f1138,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1137,f616])).
fof(f1137,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1136,f84])).
fof(f1136,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f998,f76])).
fof(f998,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f770,f74])).
fof(f1129,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1128,f616])).
fof(f1128,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1127,f84])).
fof(f1127,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1126,f76])).
fof(f1126,plain,(
  s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK6))) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1005,f770])).
fof(f1005,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,sK4),s(sK0,sK6))) = s('type.pair.prod'(sK2,sK0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f82,f770,f74])).
fof(f1120,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1119,f616])).
fof(f1119,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1118,f84])).
fof(f1118,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1012,f76])).
fof(f1012,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK10(sK3,sK1,sK0)))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f809,f74])).
fof(f1112,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1111,f616])).
fof(f1111,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1110,f84])).
fof(f1110,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1109,f76])).
fof(f1109,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK7))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1019,f809])).
fof(f1019,plain,(
  s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,sK4),s(sK1,sK7))) = s('type.pair.prod'(sK2,sK1),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK1,sK10(sK3,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f82,f809,f74])).
fof(f1108,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1107,f632])).
fof(f1107,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1106,f84])).
fof(f1106,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1020,f76])).
fof(f1020,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f70,f82,f74])).
fof(f1105,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1104,f632])).
fof(f1104,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1103,f84])).
fof(f1103,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1021,f76])).
fof(f1021,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f84,f82,f74])).
fof(f1102,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1101,f632])).
fof(f1101,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1100,f84])).
fof(f1100,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f1022,f76])).
fof(f1022,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f472,f82,f74])).
fof(f1099,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1098,f632])).
fof(f1098,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1097,f84])).
fof(f1097,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1023,f76])).
fof(f1023,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f626,f82,f74])).
fof(f1096,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1095,f632])).
fof(f1095,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4)))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1094,f84])).
fof(f1094,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1024,f76])).
fof(f1024,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f770,f82,f74])).
fof(f1093,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1092,f632])).
fof(f1092,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1091,f84])).
fof(f1091,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1025,f76])).
fof(f1025,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f809,f82,f74])).
fof(f1090,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1089,f632])).
fof(f1089,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1088,f84])).
fof(f1088,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | $spl4),
  inference(forward_demodulation,[],[f1026,f76])).
fof(f1026,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f82,f74])).
fof(f1087,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1086,f632])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1085,f84])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1027,f76])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK2),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f70,f82,f74])).
fof(f1084,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1083,f632])).
fof(f1083,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1082,f84])).
fof(f1082,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1028,f76])).
fof(f1028,plain,(
  s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f84,f82,f74])).
fof(f1081,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1080,f632])).
fof(f1080,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1079,f84])).
fof(f1079,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl22)),
  inference(forward_demodulation,[],[f1029,f76])).
fof(f1029,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK3),s(sK2,sK4))) = s('type.pair.prod'('type.pair.prod'(sK0,sK1),sK2),'const.pair.,_2'(s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl22)),
  inference(unit_resulting_resolution,[],[f472,f82,f74])).
fof(f1078,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1077,f632])).
fof(f1077,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1076,f84])).
fof(f1076,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1030,f76])).
fof(f1030,plain,(
  s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f626,f82,f74])).
fof(f1075,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl22)),
  inference(subsumption_resolution,[],[f1074,f632])).
fof(f1074,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1073,f84])).
fof(f1073,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(forward_demodulation,[],[f1031,f76])).
fof(f1031,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK6),s(sK2,sK4))) = s('type.pair.prod'(sK0,sK2),'const.pair.,_2'(s(sK0,sK9(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))))) ) | ($spl4 | $spl6 | $spl22)),
  inference(unit_resulting_resolution,[],[f770,f82,f74])).
fof(f1072,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1071,f632])).
fof(f1071,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1070,f84])).
fof(f1070,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1032,f76])).
fof(f1032,plain,(
  s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK7),s(sK2,sK4))) = s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK10(sK3,sK1,sK0)),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f809,f82,f74])).
fof(f1069,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1068,f616])).
fof(f1068,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1067,f84])).
fof(f1067,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))))) | $spl4),
  inference(forward_demodulation,[],[f1033,f76])).
fof(f1033,plain,(
  s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,sK4),s(sK2,sK4))) = s('type.pair.prod'(sK2,sK2),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))),s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) | $spl4),
  inference(unit_resulting_resolution,[],[f82,f82,f74])).
fof(f602,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f79,f54,f53,f52,f55,f58,f69,f66,f70,f71,f74,f73,f72,f75,f76,f82,f85,f87,f90,f106,f109,f110,f64,f63,f65,f57,f150,f152,f147,f148,f173,f206,f228,f361,f362,f363,f242,f364,f366,f368,f369,f371,f373,f375,f444,f439,f390,f433,f404,f405,f406,f450,f452,f420,f421,f459,f460,f462,f464,f471,f475,f151,f80,f480,f497,f498,f500,f502,f504,f506,f509,f522,f561,f483,f481,f84])).
fof(f481,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f82,f80])).
fof(f483,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | ($spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f482])).
fof(f482,plain,(
  s(sK2,sK4) != s(sK2,sK4) | s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | ($spl2 | $spl4)),
  inference(superposition,[],[f80,f82])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f560,f70])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK1,sK10(sK3,sK1,sK0)))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f492,f76])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK9(sK3,sK1,sK0)))),s(sK1,sK10(sK3,sK1,sK0)))),s(X0,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f480,f72])).
fof(f522,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4))) | s(X0,X3) = s(X0,sK9(X2,X1,X0))) )),
  inference(superposition,[],[f72,f70])).
fof(f509,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s('type.pair.prod'(bool,X62),'const.pair.,_2'(s(bool,X61),s(X62,X63))) != s('type.pair.prod'(bool,X62),'const.pair.,_2'(s(bool,X64),s(X62,X65))) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f72,f58])).
fof(f506,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s('type.pair.prod'(bool,X45),'const.pair.,_2'(s(bool,'F_0'),s(X45,X46))) != s('type.pair.prod'(bool,X45),'const.pair.,_2'(s(bool,X47),s(X45,X48))) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f72,f58])).
fof(f504,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X33),s(X34,X35))) != s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X36),s(X34,X37))) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f72,f85])).
fof(f502,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'T_0'),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X23),s(X21,X24))) | s(bool,X20) = s(bool,X23) | s(bool,X20) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f72,f85])).
fof(f500,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.pair.prod'(sK2,X12),'const.pair.,_2'(s(sK2,sK4),s(X12,X13))) != s('type.pair.prod'(sK2,X12),'const.pair.,_2'(s(sK2,X14),s(X12,X15))) | s(sK2,X14) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl4),
  inference(superposition,[],[f72,f82])).
fof(f498,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),X3),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(X3,X4))) != s('type.pair.prod'('type.pair.prod'(X0,X1),X3),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X5),s(X3,X6))) | s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f72,f70])).
fof(f497,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,sK9('const.pair.,_2'(s(X0,X1),s(X2,X3)),X2,X0))) )),
  inference(unit_resulting_resolution,[],[f70,f72])).
fof(f480,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK9(sK3,sK1,sK0)))),s(sK1,sK10(sK3,sK1,sK0)))) | $spl2),
  inference(unit_resulting_resolution,[],[f70,f80])).
fof(f151,plain,(
  ( ! [X1] : (s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))),s(fun(sK0,fun(sK1,sK2)),sK5))) = s(sK2,sK8(happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)),X1,sK2)) | ~p(s(bool,happ(s(fun(sK2,bool),X1),s(sK2,sK4))))) ) | $spl4),
  inference(forward_demodulation,[],[f146,f76])).
fof(f146,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK2,bool),X1),s(sK2,sK4)))) | s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))) = s(sK2,sK8(happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)),X1,sK2))) ) | $spl4),
  inference(superposition,[],[f64,f82])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(condensation,[],[f474])).
fof(f474,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK8(X14,X13,X12)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f468,f87])).
fof(f468,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK8(X14,X13,X12)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f65,f85])).
fof(f464,plain,(
  ( ! [X19,X18] : (p(s(bool,happ(s(fun(bool,bool),X19),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK8(X18,X19,bool))) )),
  inference(superposition,[],[f65,f58])).
fof(f462,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14)))) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | s(bool,'T_0') = s(bool,sK8(X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f65,f85])).
fof(f460,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,X8) = s(bool,sK8(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f65,f85])).
fof(f459,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X5),s('type.pair.prod'(sK0,sK1),sK3)))))) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X3)))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X5,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,sK8(X3,X4,bool))) ) | $spl6),
  inference(superposition,[],[f65,f363])).
fof(f421,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1))) | s(bool,X0) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3)))) ) | $spl6),
  inference(equality_factoring,[],[f363])).
fof(f420,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X34),s('type.pair.prod'(sK0,sK1),sK3))) = s(bool,sK8(happ(s(fun('type.pair.prod'(sK0,sK1),bool),X34),s('type.pair.prod'(sK0,sK1),sK3)),X36,bool)) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X34,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X35)) ) | $spl6),
  inference(superposition,[],[f64,f363])).
fof(f406,plain,(
  ( ! [X50,X51,X49] : (s('type.pair.prod'(X50,bool),X49) = s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,sK9(X49,bool,X50)),s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X51),s('type.pair.prod'(sK0,sK1),sK3))))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X51,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,sK10(X49,bool,X50))) ) | $spl6),
  inference(superposition,[],[f70,f363])).
fof(f405,plain,(
  ( ! [X47,X48,X46] : (s('type.pair.prod'(bool,X47),X46) = s('type.pair.prod'(bool,X47),'const.pair.,_2'(s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X48),s('type.pair.prod'(sK0,sK1),sK3))),s(X47,sK10(X46,X47,bool)))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X48,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,sK9(X46,X47,bool))) ) | $spl6),
  inference(superposition,[],[f70,f363])).
fof(f404,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X44),s('type.pair.prod'(sK0,sK1),sK3)))))) | s(bool,X43) = s(bool,sK8(X43,X45,bool)) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X44,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f64,f363])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3)))) | s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1))) = s('type.pair.prod'(sK0,sK1),sK8(sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)),X1,'type.pair.prod'(sK0,sK1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X0),s('type.pair.prod'(sK0,sK1),sK3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f64,f363])).
fof(f375,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),sK8('const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))),X3,'type.pair.prod'(X0,X1))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X3),s('type.pair.prod'(X0,X1),X2))))) )),
  inference(forward_demodulation,[],[f374,f70])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X3),s('type.pair.prod'(X0,X1),X2)))) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0)))) = s('type.pair.prod'(X0,X1),sK8('const.pair.,_2'(s(X0,sK9(X2,X1,X0)),s(X1,sK10(X2,X1,X0))),X3,'type.pair.prod'(X0,X1)))) )),
  inference(superposition,[],[f64,f70])).
fof(f373,plain,(
  ( ! [X12,X13] : (s('type.pair.prod'(X13,bool),X12) = s('type.pair.prod'(X13,bool),'const.pair.,_2'(s(X13,sK9(X12,bool,X13)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X12,bool,X13))) )),
  inference(superposition,[],[f70,f58])).
fof(f371,plain,(
  ( ! [X6,X8,X7] : (s('type.pair.prod'(X7,bool),X6) = s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,sK9(X6,bool,X7)),s(bool,X8))) | s(bool,'T_0') = s(bool,sK10(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f70,f85])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X1,bool),X0) = s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK9(X0,bool,X1)),s(bool,'T_0'))) | s(bool,X2) = s(bool,sK10(X0,bool,X1)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f70,f85])).
fof(f368,plain,(
  ( ! [X12,X13] : (s('type.pair.prod'(bool,X13),X12) = s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,'F_0'),s(X13,sK10(X12,X13,bool)))) | s(bool,'T_0') = s(bool,sK9(X12,X13,bool))) )),
  inference(superposition,[],[f70,f58])).
fof(f366,plain,(
  ( ! [X6,X8,X7] : (s('type.pair.prod'(bool,X7),X6) = s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X8),s(X7,sK10(X6,X7,bool)))) | s(bool,'T_0') = s(bool,sK9(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f70,f85])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(bool,X1),X0) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,sK10(X0,X1,bool)))) | s(bool,X2) = s(bool,sK9(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f70,f85])).
fof(f242,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | s(bool,X29) = s(bool,sK8(X29,X30,bool)) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f64,f85])).
fof(f363,plain,(
  ( ! [X28,X27] : (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X27,'type.pair.prod'(sK0,sK1))) | s(bool,X28) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X27),s('type.pair.prod'(sK0,sK1),sK3))) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(subsumption_resolution,[],[f241,f55])).
fof(f241,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X27,'type.pair.prod'(sK0,sK1))) | s(bool,X28) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X27),s('type.pair.prod'(sK0,sK1),sK3))) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f150,f85])).
fof(f362,plain,(
  ( ! [X26,X24,X23,X25] : (s(X23,X25) = s(X23,sK8(X25,X24,X23)) | s(bool,X26) = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f240,f55])).
fof(f240,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | s(X23,X25) = s(X23,sK8(X25,X24,X23)) | s(bool,X26) = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f64,f85])).
fof(f361,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f232,f55])).
fof(f232,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f69,f85])).
fof(f228,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,'T_0') = s(bool,sK8('T_0',X20,bool)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f64,f85])).
fof(f206,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X31)))) | s(bool,X30) = s(bool,sK8(X30,X32,bool)) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f64,f85])).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f85,f85])).
fof(f148,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | s(bool,X4) = s(bool,sK8(X4,X5,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f64,f58])).
fof(f147,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK8('F_0',X3,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f64,f58])).
fof(f152,plain,(
  ( ! [X1] : (s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) = s(sK2,sK8(happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7)),X1,sK2)) | ~p(s(bool,happ(s(fun(sK2,bool),X1),s(sK2,sK4))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f151,f84])).
fof(f150,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X0),s('type.pair.prod'(sK0,sK1),sK3)))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)))) ) | $spl6),
  inference(forward_demodulation,[],[f145,f84])).
fof(f145,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X0),s('type.pair.prod'(sK0,sK1),sK3)))) | s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X0,'type.pair.prod'(sK0,sK1)))) ) | $spl6),
  inference(superposition,[],[f64,f84])).
fof(f57,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.evalProps.evalProps_42_prove_ltb',aHL_FALSITY)).
fof(f65,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK8(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : ! [X17,X18] : (? [X9] : (s(X8,X9) = s(X8,X18) & p(s(bool,happ(s(fun(X8,bool),X17),s(X8,X9))))) <=> p(s(bool,happ(s(fun(X8,bool),X17),s(X8,X18)))))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.bool.UNWIND_THM2')).
fof(f63,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f47])).
fof(f64,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK8(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f47])).
fof(f109,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f103,f57])).
fof(f103,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f66,f58])).
fof(f87,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f57,f58])).
fof(f85,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f58,f58])).
fof(f75,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7,X8] : ! [X22,X23] : s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X22),s(fun(X6,fun(X7,X8)),X23))) = s(X8,happ(s(fun(fun(X6,fun(X7,X8)),X8),happ(s(fun('type.pair.prod'(X6,X7),fun(fun(X6,fun(X7,X8)),X8)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X6,X7),X22))),s(fun(X6,fun(X7,X8)),X23)))),
  file('thm.evalProps.evalProps_42_prove_ltb','ARITY_0_2_const.pair.pair_CASE')).
fof(f71,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.evalProps_42_prove_ltb',aHL_EXT)).
fof(f66,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f69,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f38])).
fof(f58,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.evalProps_42_prove_ltb',aHL_BOOL_CASES)).
fof(f55,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.evalProps_42_prove_ltb',aHL_TRUTH)).
fof(f601,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f54,f53,f52,f55,f58,f69,f66,f70,f71,f74,f73,f72,f75,f76,f82,f85,f87,f90,f106,f109,f110,f64,f63,f65,f57,f150,f152,f147,f148,f173,f206,f228,f361,f362,f363,f242,f364,f366,f368,f369,f371,f373,f375,f444,f439,f390,f433,f404,f405,f406,f450,f452,f420,f421,f459,f460,f462,f464,f471,f475,f151,f80,f480,f497,f498,f500,f502,f504,f506,f509,f522,f561,f483,f481,f84])).
fof(f600,plain,(
  $false | ($spl2 | $spl4 | $spl6 | $spl22)),
  inference(global_subsumption,[],[f472,f54,f53,f52,f55,f58,f69,f66,f70,f71,f74,f73,f72,f75,f76,f84,f82,f85,f87,f90,f106,f109,f110,f64,f63,f65,f57,f150,f152,f147,f148,f173,f206,f228,f361,f362,f363,f242,f364,f366,f368,f369,f371,f373,f375,f444,f439,f390,f433,f404,f405,f406,f450,f452,f420,f421,f459,f460,f462,f464,f471,f475,f151,f80,f480,f497,f498,f500,f502,f504,f506,f509,f522,f561,f483,f481])).
fof(f599,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f54,f53,f52,f55,f58,f69,f66,f70,f71,f74,f73,f72,f75,f76,f84,f82,f85,f87,f90,f106,f109,f110,f64,f63,f65,f57,f150,f152,f147,f148,f173,f206,f228,f361,f362,f363,f242,f364,f366,f368,f369,f371,f373,f375,f444,f439,f390,f433,f404,f405,f406,f450,f452,f420,f421,f459,f460,f462,f464,f471,f475,f151,f80,f480,f497,f498,f500,f502,f504,f506,f509,f522,f561,f483,f481])).
fof(f598,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f54,f53,f52,f55,f58,f69,f66,f70,f71,f74,f73,f72,f75,f76,f84,f82,f85,f87,f90,f106,f109,f110,f64,f63,f65,f57,f150,f152,f147,f148,f173,f206,f228,f361,f362,f363,f242,f364,f366,f368,f369,f371,f373,f375,f444,f439,f390,f433,f404,f405,f406,f450,f452,f420,f421,f459,f460,f462,f464,f471,f475,f151,f80,f480,f497,f498,f500,f502,f504,f506,f509,f522,f561,f483])).
fof(f589,plain,(
  $false | ($spl0 | $spl2)),
  inference(equality_resolution,[],[f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X2)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f558,f78])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f557,f70])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK9(sK3,sK1,sK0)),s(sK1,sK10(sK3,sK1,sK0)))),s(fun(sK0,fun(sK1,sK2)),sK5))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f495,f76])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,sK4),s(X0,X1))) != s('type.pair.prod'(sK2,X0),'const.pair.,_2'(s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK9(sK3,sK1,sK0)))),s(sK1,sK10(sK3,sK1,sK0)))),s(X0,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f480,f72])).
fof(f472,plain,(
  ( ! [X3] : (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X3,'type.pair.prod'(sK0,sK1)))) ) | $spl22),
  inference(cnf_transformation,[],[f472_D])).
fof(f472_D,plain,(
  ( ! [X3] : (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X3,'type.pair.prod'(sK0,sK1)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) ) | $spl24),
  inference(cnf_transformation,[],[f473_D])).
fof(f473_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f471,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X3,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) ) | $spl6),
  inference(subsumption_resolution,[],[f465,f150])).
fof(f465,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X3),s('type.pair.prod'(sK0,sK1),sK3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X3,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) ) | $spl6),
  inference(superposition,[],[f65,f363])).
fof(f107,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f446,plain,(
  ( ! [X16] : (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X16,'type.pair.prod'(sK0,sK1))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X16),s('type.pair.prod'(sK0,sK1),sK3)))) ) | $spl18),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  ( ! [X16] : (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X16,'type.pair.prod'(sK0,sK1))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X16),s('type.pair.prod'(sK0,sK1),sK3)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f453,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X29),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X29,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f418,f87])).
fof(f418,plain,(
  ( ! [X30,X29] : (p(s(bool,X30)) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X29),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X29,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f108,f363])).
fof(f451,plain,(
  ( ! [X6,X8] : (p(s(bool,X8)) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X6,'type.pair.prod'(sK0,sK1))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X6),s('type.pair.prod'(sK0,sK1),sK3)))) ) | $spl20),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  ( ! [X6,X8] : (p(s(bool,X8)) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X6,'type.pair.prod'(sK0,sK1))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X6),s('type.pair.prod'(sK0,sK1),sK3)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f452,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X9),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X9,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(subsumption_resolution,[],[f411,f87])).
fof(f411,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X9),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X9,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f66,f363])).
fof(f450,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X6),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X6,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f410,f87])).
fof(f410,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X6),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X6,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X7)) ) | $spl6),
  inference(superposition,[],[f66,f363])).
fof(f433,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) | s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f391])).
fof(f391,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) | s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f58,f363])).
fof(f439,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X16),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X16,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X15)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f381])).
fof(f381,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X16),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X16,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl6),
  inference(superposition,[],[f363,f58])).
fof(f445,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X2),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X2,'type.pair.prod'(sK0,sK1)))) ) | $spl16),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X2),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X2,'type.pair.prod'(sK0,sK1)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X2),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X2,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X1),s('type.pair.prod'(sK0,sK1),sK3))) = s(bool,happ(s(fun('type.pair.prod'(sK0,sK1),bool),X2),s('type.pair.prod'(sK0,sK1),sK3))) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X2,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),sK8('const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)),X1,'type.pair.prod'(sK0,sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f363,f363])).
fof(f115,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f111,f55])).
fof(f111,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f107,f57])).
fof(f108,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f110,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f104,f87])).
fof(f104,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f66,f58])).
fof(f106,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f102,f87])).
fof(f102,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f66,f58])).
fof(f91,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f93,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f58])).
fof(f78,plain,(
  s(sK2,sK4) = s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) | $spl0),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  s(sK2,sK4) = s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f84,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | $spl6),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f52,plain,(
  s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | s(sK2,sK4) = s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  (s(sK2,sK4) = s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) | (s('type.pair.prod'(sK0,sK1),sK3) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) & s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))))) & (s(sK2,sK4) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) | ! [X6,X7] : (s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,X6))),s(sK1,X7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f41])).
fof(f41,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X2,X4) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) | ? [X8,X9] : (s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X9))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X8))),s(X1,X9))))) & (s(X2,X4) != s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) | ! [X6,X7] : (s('type.pair.prod'(X0,X1),X3) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) | s(X2,X4) != s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7))))))),
  inference(rectify,[],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X2,X4) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) | ? [X6,X7] : (s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7))))) & (s(X2,X4) != s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) | ! [X6,X7] : (s('type.pair.prod'(X0,X1),X3) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) | s(X2,X4) != s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7))))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X2,X4) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) <~> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7)))))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X2,X4) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7)))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : (s(X2,X4) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X3) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X5),s(X0,X6))),s(X1,X7)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : (s(X8,X10) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X9),s(fun(X6,fun(X7,X8)),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X6,X7),X9) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13))) & s(X8,X10) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X11),s(X6,X12))),s(X7,X13)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : (s(X8,X10) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X9),s(fun(X6,fun(X7,X8)),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X6,X7),X9) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13))) & s(X8,X10) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X11),s(X6,X12))),s(X7,X13)))))),
  file('thm.evalProps.evalProps_42_prove_ltb','thm.evalProps.evalProps_42_prove_ltb')).
fof(f82,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f53,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,sK6))),s(sK1,sK7))) | s(sK2,sK4) = s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))),
  inference(cnf_transformation,[],[f42])).
fof(f79,plain,(
  s(sK2,sK4) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) | $spl1),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  s(sK2,sK4) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f80,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f54,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK5),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK3) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7))) | s(sK2,sK4) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK3),s(fun(sK0,fun(sK1,sK2)),sK5)))) )),
  inference(cnf_transformation,[],[f42])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 41
% Passive clauses: 139
% Generated clauses: 1327
% Final active clauses: 29
% Final passive clauses: 76
% Input formulas: 14
% Initial clauses: 25
% 
% Duplicate literals: 53
% Trivial inequalities: 9
% Fw subsumption resolutions: 74
% Fw demodulations: 314
% Bw demodulations: 8
% Condensations: 67
% Global subsumptions: 5
% 
% Simple tautologies: 21
% Equational tautologies: 34
% Forward subsumptions: 397
% Fw demodulations to eq. taut.: 142
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 248
% Factoring: 1
% Forward superposition: 271
% Backward superposition: 166
% Self superposition: 26
% Equality factoring: 28
% Equality resolution: 5
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 231
% SAT solver unit clauses: 9
% SAT solver binary clauses: 52
% 
% Sat splits: 13
% Sat splitting refutations: 59
% 
% Memory used [KB]: 2046
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 1.371 s
