Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.holSemanticsExtra
% (70014)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 163
% Passive clauses: 2580
% Generated clauses: 15583
% Final active clauses: 151
% Final passive clauses: 2270
% Input formulas: 15
% Initial clauses: 33
% 
% Duplicate literals: 520
% Trivial inequalities: 413
% Fw subsumption resolutions: 1218
% Fw demodulations: 2871
% Bw demodulations: 6
% 
% Simple tautologies: 349
% Equational tautologies: 21
% Forward subsumptions: 7163
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 185
% Forward superposition: 6001
% Backward superposition: 4169
% Self superposition: 103
% Equality factoring: 44
% Equality resolution: 42
% 
% Split clauses: 306
% Split components: 619
% SAT solver clauses: 592
% SAT solver unit clauses: 16
% SAT solver binary clauses: 99
% 
% Sat splits: 306
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 591
% TWLsolver calls for satisfiability: 124
% 
% Memory used [KB]: 12537
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.holSemanticsExtra
% (70051)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 693
% Generated clauses: 8586
% Final active clauses: 107
% Final passive clauses: 539
% Input formulas: 15
% Initial clauses: 33
% 
% Duplicate literals: 264
% Trivial inequalities: 197
% Fw subsumption resolutions: 516
% Fw demodulations: 1275
% Bw demodulations: 45
% Condensations: 560
% 
% Simple tautologies: 159
% Equational tautologies: 15
% Forward subsumptions: 3850
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 83
% Forward superposition: 2893
% Backward superposition: 1999
% Self superposition: 66
% Equality factoring: 28
% Equality resolution: 30
% 
% Split clauses: 94
% Split components: 199
% SAT solver clauses: 1056
% SAT solver unit clauses: 33
% SAT solver binary clauses: 18
% 
% Sat splits: 94
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1056
% TWLsolver calls for satisfiability: 237
% 
% Memory used [KB]: 9722
% Time elapsed: 0.719 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.holSemanticsExtra
% (70070)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 626
% Passive clauses: 15453
% Generated clauses: 192665
% Final active clauses: 115
% Final passive clauses: 1120
% Input formulas: 15
% Initial clauses: 33
% 
% Duplicate literals: 7670
% Trivial inequalities: 2198
% Fw subsumption resolutions: 8498
% Fw demodulations: 29147
% Bw demodulations: 11564
% Condensations: 14631
% Global subsumptions: 673
% 
% Simple tautologies: 1686
% Equational tautologies: 3213
% Forward subsumptions: 108466
% Fw demodulations to eq. taut.: 32
% Bw demodulations to eq. taut.: 1019
% 
% Unit resulting resolution: 17509
% Forward superposition: 40647
% Backward superposition: 35674
% Self superposition: 775
% Equality factoring: 349
% Equality resolution: 69
% 
% Split clauses: 258
% Split components: 530
% SAT solver clauses: 17439
% SAT solver unit clauses: 177
% SAT solver binary clauses: 2495
% 
% Sat splits: 258
% Sat splitting refutations: 795
% 
% Memory used [KB]: 87248
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.holSemanticsExtra
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.holSemanticsExtra
% SZS output start Proof for thm.holSemanticsExtra
fof(f8928,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f88,f89,f90,f91,f92,f94,f101,f95,f96,f99,f108,f109,f110,f111,f112,f114,f120,f115,f116,f119,f131,f160,f162,f179,f180,f196,f195,f194,f209,f223,f222,f210,f224,f215,f228,f216,f231,f237,f238,f239,f240,f241,f243,f251,f244,f245,f253,f254,f255,f256,f257,f259,f260,f261,f293,f299,f300,f301,f302,f303,f304,f306,f307,f308,f325,f330,f331,f332,f333,f334,f335,f357,f337,f338,f356,f339,f362,f363,f364,f365,f366,f326,f367,f389,f369,f370,f371,f407,f394,f395,f396,f397,f398,f413,f399,f401,f402,f416,f417,f418,f419,f420,f422,f279,f102,f423,f424,f493,f230,f511,f506,f522,f524,f525,f527,f529,f530,f388,f557,f650,f738,f740,f741,f743,f745,f746,f764,f770,f772,f773,f774,f777,f778,f827,f829,f830,f831,f834,f835,f873,f875,f876,f878,f880,f881,f901,f909,f911,f733,f729,f725,f721,f716,f712,f640,f636,f632,f628,f623,f619,f956,f229,f957,f869,f854,f490,f455,f295,f903,f904,f1100,f1101,f1155,f1157,f1158,f1159,f1162,f1163,f1303,f1388,f1390,f1391,f1392,f1395,f1396,f250,f1421,f1420,f1439,f1440,f1457,f1459,f1460,f1461,f1464,f1465,f1480,f1481,f1484,f1486,f1487,f1489,f1491,f1492,f412,f1551,f1553,f1554,f1555,f1558,f1559,f1564,f1566,f1567,f1569,f1571,f1572,f1577,f1579,f1580,f1582,f1584,f1585,f1834,f1837,f1839,f1840,f1841,f1843,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1835,f1854,f1856,f1857,f1858,f1861,f1862,f1876,f1878,f1879,f1880,f1883,f1884,f1889,f1891,f1892,f1894,f1896,f1897,f1902,f1904,f1905,f1907,f1909,f1910,f1969,f1971,f1972,f1973,f1976,f1977,f1999,f2000,f2220,f2223,f2225,f2226,f2227,f2229,f2230,f2231,f2232,f2233,f2234,f2235,f2236,f2221,f2257,f2259,f2263,f2265,f2266,f2267,f2270,f2271,f2276,f2278,f2279,f2281,f2283,f2284,f2333,f2335,f2336,f2338,f2340,f2341,f2495,f2498,f2501,f2503,f2505,f2507,f2496,f2515,f2517,f2518,f2519,f2522,f2523,f2528,f2530,f2531,f2533,f2535,f2536,f2624,f2594,f2595,f2625,f2736,f2656,f2658,f2659,f2660,f2663,f2664,f2670,f2672,f2673,f2675,f2677,f2678,f2686,f2688,f2689,f2691,f2693,f2694,f2860,f2862,f2863,f2865,f2867,f2868,f2897,f2898,f2909,f2911,f2912,f2913,f2916,f2917,f2955,f2957,f2958,f2959,f2962,f2963,f3217,f3220,f3222,f3223,f3224,f3226,f3227,f3228,f3229,f3230,f3231,f3232,f3233,f161,f3218,f3238,f3240,f3241,f3242,f3245,f3246,f3260,f3262,f3263,f3264,f3267,f3268,f3301,f3459,f3461,f3462,f3464,f3466,f3467,f3612,f3615,f3617,f3618,f3619,f3621,f3622,f3623,f3624,f3625,f3626,f3627,f3628,f163,f3613,f3638,f3640,f3641,f3643,f3645,f3646,f3719,f3862,f3865,f3868,f3870,f3872,f3874,f3863,f4069,f4028,f4074,f4029,f4075,f4030,f2628,f4031,f4070,f2629,f4258,f411,f409,f4300,f4302,f4303,f4305,f4307,f4308,f4347,f4348,f358,f4364,f4366,f4367,f4368,f4371,f4372,f4452,f4467,f4470,f4453,f4549,f4564,f4567,f4550,f4618,f4620,f4621,f4623,f4625,f4626,f4678,f4680,f4681,f4683,f4685,f4686,f4738,f4740,f4741,f4742,f4745,f4746,f4798,f4800,f4801,f4802,f4805,f4806,f4858,f4860,f4861,f4862,f4865,f4866,f4922,f4924,f4925,f4926,f4929,f4930,f4986,f4988,f4989,f4990,f4993,f4994,f5034,f5036,f5037,f5048,f5050,f5051,f5052,f5055,f5056,f5105,f5107,f5108,f5109,f5112,f5113,f5153,f5155,f5156,f5204,f5206,f5207,f5209,f5211,f5212,f5275,f5278,f5281,f5517,f5533,f5536,f5538,f5544,f5546,f5603,f5605,f5606,f5608,f5610,f5611,f225,f5616,f5618,f5619,f5620,f5623,f5624,f278,f5666,f5668,f5669,f5670,f5673,f5674,f1613,f5679,f5681,f5682,f5684,f5686,f5687,f1614,f5690,f5692,f5693,f5695,f5697,f5698,f1615,f5869,f5870,f5822,f5824,f3302,f5871,f6000,f6001,f5952,f5873,f6002,f6048,f6050,f6051,f6052,f6055,f6056,f1942,f6061,f6063,f6064,f6066,f6068,f6069,f1943,f6072,f6074,f6075,f6077,f6079,f6080,f1944,f6134,f6136,f6137,f6138,f6141,f6142,f799,f6170,f6172,f6173,f6174,f6177,f6178,f6189,f6191,f6192,f6193,f6196,f6197,f6208,f6210,f6211,f6212,f6215,f6216,f1997,f6248,f6250,f6251,f6253,f6255,f6256,f6303,f6305,f6306,f6307,f6310,f6311,f2312,f6315,f6317,f6318,f6320,f6322,f6323,f2314,f6371,f6373,f6374,f6375,f6378,f6379,f3299,f6413,f6415,f6416,f6418,f6420,f6421,f6469,f6471,f6472,f6473,f6476,f6477,f2565,f6481,f6483,f6484,f6486,f6488,f6489,f2567,f6542,f6544,f6545,f6546,f6549,f6550,f2946,f6585,f6587,f6588,f6589,f6592,f6593,f4410,f6630,f6632,f6633,f6634,f6637,f6638,f2737,f6644,f6646,f6647,f6649,f6651,f6652,f2738,f6662,f6664,f6665,f6667,f6669,f6670,f2739,f6873,f6875,f6878,f6880,f6882,f6884,f1838,f6988,f6990,f6993,f6995,f6997,f6999,f3221,f7024,f7026,f7027,f7028,f7031,f7032,f840,f7174,f7176,f7179,f7181,f7183,f7185,f2224,f7193,f5282,f7339,f7341,f7344,f7346,f7348,f7350,f3616,f7358,f5279,f7368,f7370,f7371,f7373,f7375,f7376,f7421,f7423,f7424,f7425,f7428,f7429,f7470,f7472,f7473,f7475,f7477,f7478,f7386,f7513,f7499,f7498,f7496,f7494,f7493,f7491,f3481,f7608,f7610,f7611,f7612,f7615,f7616,f7435,f7795,f5276,f390,f7844,f7846,f7847,f7849,f7851,f7852,f2895,f7900,f7902,f7903,f7905,f7907,f7908,f4345,f7957,f7959,f7960,f7961,f7964,f7965,f5066,f7974,f7976,f7977,f7978,f7981,f7982,f1478,f8455,f8457,f8458,f8460,f8462,f8463,f5229,f8521,f8523,f8524,f8526,f8528,f8529,f4664,f8582,f8584,f8585,f8587,f8589,f8590,f4724,f8643,f8645,f8646,f8647,f8650,f8651,f4786,f8703,f8705,f8706,f8707,f8710,f8711,f4846,f8763,f8765,f8766,f8767,f8770,f8771,f4906,f8826,f8828,f8829,f100,f8830,f104,f8833,f103,f8834,f4970,f105,f8920])).
fof(f8920,plain,(
  $false | $spl61),
  inference(subsumption_resolution,[],[f8890,f409])).
fof(f8890,plain,(
  p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f107,f249])).
fof(f249,plain,(
  p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f236,f64])).
fof(f64,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X25] : ! [X14,X15] : s(X25,X15) = s(X25,'const.pair.SND_1'(s('type.pair.prod'(X13,X25),'const.pair.,_2'(s(X13,X14),s(X25,X15)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.pair.SND')).
fof(f236,plain,(
  p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f57,f60])).
fof(f60,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) | (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))))) & (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) | (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))))) & ((~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))) | p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) <=> (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) <=> (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X21,X22] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X22)))) <=> (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X22)))))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X22))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X22))))))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.holSemantics.is_interpretation_def')).
fof(f57,plain,(
  p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) & ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2)))) & ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2)))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2))))) & ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2)))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2))))) => p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2)))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2))))) => p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X8))))) => p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12)))) & p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X8))))) => p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6))),X8)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.holSemanticsExtra.is_interpretation_reduce')).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3))))) )),
  inference(resolution,[],[f56,f62])).
fof(f62,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))) | p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5))))) | p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3)))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5))))) => p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3)))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5))))) => p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X5)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X10,X12,X23,X24] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12)))) & p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X23),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X24))))) => p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X23),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X24)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.holSemanticsExtra.is_term_assignment_reduce')).
fof(f56,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))))),
  inference(cnf_transformation,[],[f46])).
fof(f105,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4970,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK7('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X0)) ) | $spl176),
  inference(cnf_transformation,[],[f4970_D])).
fof(f4970_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK7('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f8834,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,sK7('const.pair.,_2'(s(bool,X57),s(X58,X59)),X58,bool)) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f71])).
fof(f71,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X19,X15,X20] : ((s(bool,X14) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X15) = s(bool,X20))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.bool.IMP_CONG')).
fof(f3876,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,X264)) | s(bool,X267) != s(bool,sK7('const.pair.,_2'(s(bool,X264),s(X265,X266)),X265,bool)) | s(bool,'F_0') = s(bool,X267)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3811,f737])).
fof(f737,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X4))) ) | $spl0),
  inference(forward_demodulation,[],[f99,f106])).
fof(f106,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))),
  inference(subsumption_resolution,[],[f98,f81])).
fof(f81,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.holSemanticsExtra.is_interpretation_reduce',aHL_FALSITY)).
fof(f98,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))),
  inference(superposition,[],[f55,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.holSemanticsExtra.is_interpretation_reduce',aHL_BOOL_CASES)).
fof(f55,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))))),
  inference(cnf_transformation,[],[f46])).
fof(f3811,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,X264)) | s(bool,X267) != s(bool,sK7('const.pair.,_2'(s(bool,X264),s(X265,X266)),X265,bool)) | s(bool,'F_0') = s(bool,X267) | s(bool,'T_0') = s(bool,X264)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f494,f3611])).
fof(f3611,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool)) | s(bool,'T_0') = s(bool,X9)) ) | ($spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3490,f295])).
fof(f3490,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f3129,f85])).
fof(f3129,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X10),s(X9,X11)),X9,bool)) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f3000,f65])).
fof(f65,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X25,X13] : ! [X14,X15] : s(X13,X14) = s(X13,'const.pair.FST_1'(s('type.pair.prod'(X13,X25),'const.pair.,_2'(s(X13,X14),s(X25,X15)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.pair.FST')).
fof(f3000,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'F_0') = s(bool,sK7(X1,X0,bool))) ) | $spl6),
  inference(superposition,[],[f1545,f66])).
fof(f66,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13,X25] : ! [X14] : ? [X26,X27] : s('type.pair.prod'(X13,X25),X14) = s('type.pair.prod'(X13,X25),'const.pair.,_2'(s(X13,X26),s(X25,X27)))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.pair.ABS_PAIR_THM')).
fof(f1545,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f1533])).
fof(f1533,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f856,f85])).
fof(f856,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X8),s(X7,X9)))))) ) | $spl6),
  inference(forward_demodulation,[],[f847,f106])).
fof(f847,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X8),s(X7,X9)))))) ) | $spl6),
  inference(superposition,[],[f766,f65])).
fof(f766,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(subsumption_resolution,[],[f760,f82])).
fof(f82,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.holSemanticsExtra.is_interpretation_reduce',aHL_TRUTH)).
fof(f760,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f488,f106])).
fof(f488,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f475,f81])).
fof(f475,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'F_0')) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f494,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f484,f82])).
fof(f484,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f103,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f8833,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,sK7('const.pair.,_2'(s(bool,X50),s(X51,X52)),X51,bool)) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f74])).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f104,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) | $spl10),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f8830,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,sK7('const.pair.,_2'(s(bool,X29),s(X30,X31)),X30,bool)) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f73])).
fof(f73,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f100,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) | $spl2),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f8829,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,sK7('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool)) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f72])).
fof(f72,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f8828,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,sK7('const.pair.,_2'(s(bool,X15),s(X16,X17)),X16,bool)) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f70])).
fof(f70,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f8826,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK7('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f4906,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,X3)),bool,X1)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X0)) ) | $spl174),
  inference(cnf_transformation,[],[f4906_D])).
fof(f4906_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,X3)),bool,X1)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f8771,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,sK8('const.pair.,_2'(s(X57,X58),s(bool,X59)),bool,X57)) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f86])).
fof(f2509,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,X266)) | s(bool,X267) != s(bool,sK8('const.pair.,_2'(s(X264,X265),s(bool,X266)),bool,X264)) | s(bool,'F_0') = s(bool,X267)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2462,f737])).
fof(f2462,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,X266)) | s(bool,X267) != s(bool,sK8('const.pair.,_2'(s(X264,X265),s(bool,X266)),bool,X264)) | s(bool,'F_0') = s(bool,X267) | s(bool,'T_0') = s(bool,X266)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f494,f2219])).
fof(f2219,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10)) | s(bool,'T_0') = s(bool,X9)) ) | ($spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2139,f295])).
fof(f2139,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f1764,f85])).
fof(f1764,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X9,X10),s(bool,X11)),bool,X9)) | s(bool,'T_0') = s(bool,X11)) ) | $spl6),
  inference(superposition,[],[f1655,f64])).
fof(f1655,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1))) | s(bool,'F_0') = s(bool,sK8(X1,bool,X0))) ) | $spl6),
  inference(superposition,[],[f1526,f66])).
fof(f1526,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f1514])).
fof(f1514,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f855,f85])).
fof(f855,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) ) | $spl6),
  inference(forward_demodulation,[],[f846,f106])).
fof(f846,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) ) | $spl6),
  inference(superposition,[],[f766,f64])).
fof(f8770,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,sK8('const.pair.,_2'(s(X50,X51),s(bool,X52)),bool,X50)) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f74])).
fof(f8767,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,sK8('const.pair.,_2'(s(X29,X30),s(bool,X31)),bool,X29)) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f73])).
fof(f8766,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,sK8('const.pair.,_2'(s(X22,X23),s(bool,X24)),bool,X22)) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f72])).
fof(f8765,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,sK8('const.pair.,_2'(s(X15,X16),s(bool,X17)),bool,X15)) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f70])).
fof(f8763,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,X3)),bool,X1)) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f68])).
fof(f4846,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X0) != s(bool,X3)) ) | $spl172),
  inference(cnf_transformation,[],[f4846_D])).
fof(f4846_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X0) != s(bool,X3)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f8711,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X57,X58)),X57,bool)) | s(bool,X56) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | $spl136),
  inference(resolution,[],[f4546,f86])).
fof(f4546,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X13,X15)),X13,bool)) | s(bool,X14) != s(bool,X16)) ) | $spl136),
  inference(superposition,[],[f3221,f65])).
fof(f8710,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X50,X51)),X50,bool)) | s(bool,X49) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl136),
  inference(resolution,[],[f4546,f74])).
fof(f8707,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X28),s(X29,X30)),X29,bool)) | s(bool,X28) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | $spl136),
  inference(resolution,[],[f4546,f73])).
fof(f8706,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X22,X23)),X22,bool)) | s(bool,X21) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl136),
  inference(resolution,[],[f4546,f72])).
fof(f8705,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X15,X16)),X15,bool)) | s(bool,X14) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl136),
  inference(resolution,[],[f4546,f70])).
fof(f8703,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X0) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | $spl136),
  inference(resolution,[],[f4546,f68])).
fof(f4786,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X2) != s(bool,X3)) ) | $spl170),
  inference(cnf_transformation,[],[f4786_D])).
fof(f4786_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X2) != s(bool,X3)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f8651,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X56,X57),s(bool,X58)),bool,X56)) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | $spl90),
  inference(resolution,[],[f4449,f86])).
fof(f4449,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X13,X14),s(bool,X15)),bool,X13)) | s(bool,X15) != s(bool,X16)) ) | $spl90),
  inference(superposition,[],[f1838,f64])).
fof(f8650,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X49,X50),s(bool,X51)),bool,X49)) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl90),
  inference(resolution,[],[f4449,f74])).
fof(f8647,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X28,X29),s(bool,X30)),bool,X28)) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | $spl90),
  inference(resolution,[],[f4449,f73])).
fof(f8646,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X21,X22),s(bool,X23)),bool,X21)) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl90),
  inference(resolution,[],[f4449,f72])).
fof(f8645,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X16)),bool,X14)) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl90),
  inference(resolution,[],[f4449,f70])).
fof(f8643,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | $spl90),
  inference(resolution,[],[f4449,f68])).
fof(f4724,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) | $spl168),
  inference(cnf_transformation,[],[f4724_D])).
fof(f4724_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f8590,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) != s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X58,X59)),X58,bool)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl146),
  inference(resolution,[],[f3616,f86])).
fof(f8589,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) != s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X51,X52)),X51,bool)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl146),
  inference(resolution,[],[f3616,f74])).
fof(f8587,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X35),s(X37,X38)),X37,bool)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl146),
  inference(resolution,[],[f3616,f73])).
fof(f8585,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) != s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X23,X24)),X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl146),
  inference(resolution,[],[f3616,f72])).
fof(f8584,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X16,X17)),X16,bool)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl146),
  inference(resolution,[],[f3616,f70])).
fof(f8582,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl146),
  inference(resolution,[],[f3616,f68])).
fof(f4664,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) | $spl166),
  inference(cnf_transformation,[],[f4664_D])).
fof(f4664_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f8529,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) != s(bool,sK8('const.pair.,_2'(s(X58,X59),s(bool,X56)),bool,X58)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl104),
  inference(resolution,[],[f2224,f86])).
fof(f8528,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) != s(bool,sK8('const.pair.,_2'(s(X51,X52),s(bool,X49)),bool,X51)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl104),
  inference(resolution,[],[f2224,f74])).
fof(f8526,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X37,X38),s(bool,X35)),bool,X37)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl104),
  inference(resolution,[],[f2224,f73])).
fof(f8524,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,X21)),bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl104),
  inference(resolution,[],[f2224,f72])).
fof(f8523,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,sK8('const.pair.,_2'(s(X16,X17),s(bool,X14)),bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl104),
  inference(resolution,[],[f2224,f70])).
fof(f8521,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl104),
  inference(resolution,[],[f2224,f68])).
fof(f5229,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl188),
  inference(cnf_transformation,[],[f5229_D])).
fof(f5229_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f8463,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f86])).
fof(f8462,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X22) != s(bool,X24)) ) | $spl61),
  inference(resolution,[],[f409,f74])).
fof(f8460,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f73])).
fof(f8458,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f72])).
fof(f8457,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X7) != s(bool,X9)) ) | $spl61),
  inference(resolution,[],[f409,f70])).
fof(f8455,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f68])).
fof(f1478,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f1478_D])).
fof(f1478_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7982,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl70),
  inference(resolution,[],[f799,f86])).
fof(f7981,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl70),
  inference(resolution,[],[f799,f74])).
fof(f7978,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl70),
  inference(resolution,[],[f799,f73])).
fof(f7977,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl70),
  inference(resolution,[],[f799,f72])).
fof(f7976,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl70),
  inference(resolution,[],[f799,f70])).
fof(f7974,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl70),
  inference(resolution,[],[f799,f68])).
fof(f5066,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl182),
  inference(cnf_transformation,[],[f5066_D])).
fof(f5066_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7965,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f86])).
fof(f7964,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f249,f74])).
fof(f7961,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f73])).
fof(f7960,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f72])).
fof(f7959,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f249,f70])).
fof(f7957,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f68])).
fof(f4345,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) | $spl156),
  inference(cnf_transformation,[],[f4345_D])).
fof(f4345_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f7908,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X56)) | s(bool,X57) != s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X58,X59)),X58,bool)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl148),
  inference(resolution,[],[f3863,f86])).
fof(f7907,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X51,X52)),X51,bool)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl148),
  inference(resolution,[],[f3863,f74])).
fof(f7905,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X35),s(X37,X38)),X37,bool)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl148),
  inference(resolution,[],[f3863,f73])).
fof(f7903,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X22) != s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X23,X24)),X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl148),
  inference(resolution,[],[f3863,f72])).
fof(f7902,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X15) != s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X16,X17)),X16,bool)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl148),
  inference(resolution,[],[f3863,f70])).
fof(f7900,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl148),
  inference(resolution,[],[f3863,f68])).
fof(f2895,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) | $spl128),
  inference(cnf_transformation,[],[f2895_D])).
fof(f2895_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7852,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X56)) | s(bool,X57) != s(bool,sK8('const.pair.,_2'(s(X58,X59),s(bool,X56)),bool,X58)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl110),
  inference(resolution,[],[f2496,f86])).
fof(f7851,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,sK8('const.pair.,_2'(s(X51,X52),s(bool,X49)),bool,X51)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl110),
  inference(resolution,[],[f2496,f74])).
fof(f7849,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X37,X38),s(bool,X35)),bool,X37)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl110),
  inference(resolution,[],[f2496,f73])).
fof(f7847,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X22) != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,X21)),bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl110),
  inference(resolution,[],[f2496,f72])).
fof(f7846,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X15) != s(bool,sK8('const.pair.,_2'(s(X16,X17),s(bool,X14)),bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl110),
  inference(resolution,[],[f2496,f70])).
fof(f7844,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl110),
  inference(resolution,[],[f2496,f68])).
fof(f390,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl58),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5276,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) | $spl190),
  inference(cnf_transformation,[],[f5276_D])).
fof(f5276_D,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7795,plain,(
  ( ! [X121,X120] : (s(bool,'F_0') = s(bool,X121) | s(bool,X121) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,'T_0') = s(bool,X120)) ) | $spl6),
  inference(subsumption_resolution,[],[f7712,f85])).
fof(f7712,plain,(
  ( ! [X121,X120] : (s(bool,'F_0') != s(bool,X120) | s(bool,'F_0') = s(bool,X121) | s(bool,X121) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,'T_0') = s(bool,X120)) ) | $spl6),
  inference(superposition,[],[f780,f415])).
fof(f415,plain,(
  s(bool,'F_0') = s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))),
  inference(subsumption_resolution,[],[f404,f82])).
fof(f404,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))),
  inference(superposition,[],[f58,f85])).
fof(f58,plain,(
  ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))),
  inference(cnf_transformation,[],[f46])).
fof(f780,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) != s(bool,X48) | s(bool,'F_0') = s(bool,X46) | s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,X48)) ) | $spl6),
  inference(resolution,[],[f494,f488])).
fof(f7435,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl208),
  inference(cnf_transformation,[],[f7435_D])).
fof(f7435_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7616,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f86])).
fof(f410,plain,(
  p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | $spl62),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7615,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) ) | $spl62),
  inference(resolution,[],[f410,f74])).
fof(f7612,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f73])).
fof(f7611,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f72])).
fof(f7610,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) ) | $spl62),
  inference(resolution,[],[f410,f70])).
fof(f7608,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f68])).
fof(f3481,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6))) | ~p(s(bool,X8))) ) | $spl142),
  inference(cnf_transformation,[],[f3481_D])).
fof(f3481_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6))) | ~p(s(bool,X8))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7491,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) )),
  inference(resolution,[],[f87,f68])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X2)))) | p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X2))))) )),
  inference(resolution,[],[f55,f63])).
fof(f63,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3))))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4))))) | p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3)))) & p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4))))) => p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3)))) & p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4))))) => p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X9,X11,X23] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11)))) & p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X11),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X23))))) => p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X23)))))),
  file('thm.holSemanticsExtra.is_interpretation_reduce','thm.holSemanticsExtra.is_type_assignment_reduce')).
fof(f7493,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X18)))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f87,f70])).
fof(f7494,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X24)))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X24)))) )),
  inference(resolution,[],[f87,f72])).
fof(f7496,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X36)))) )),
  inference(resolution,[],[f87,f73])).
fof(f7498,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X48)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f87,f74])).
fof(f7499,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X54)))) | p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X54)))) )),
  inference(resolution,[],[f87,f86])).
fof(f7513,plain,(
  $false | $spl63),
  inference(subsumption_resolution,[],[f7490,f411])).
fof(f7490,plain,(
  p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f87,f248])).
fof(f248,plain,(
  p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f235,f65])).
fof(f235,plain,(
  p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f57,f59])).
fof(f59,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | ~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f7386,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl206),
  inference(cnf_transformation,[],[f7386_D])).
fof(f7386_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7478,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f86])).
fof(f7477,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X22) != s(bool,X24)) ) | $spl63),
  inference(resolution,[],[f411,f74])).
fof(f7475,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f73])).
fof(f7473,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f72])).
fof(f7472,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X7) != s(bool,X9)) ) | $spl63),
  inference(resolution,[],[f411,f70])).
fof(f7470,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f68])).
fof(f7429,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f86])).
fof(f7428,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) ) | $spl62),
  inference(resolution,[],[f410,f74])).
fof(f7425,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f73])).
fof(f7424,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f72])).
fof(f7423,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) ) | $spl62),
  inference(resolution,[],[f410,f70])).
fof(f7421,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl62),
  inference(resolution,[],[f410,f68])).
fof(f7376,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f86])).
fof(f7375,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X22) != s(bool,X24)) ) | $spl63),
  inference(resolution,[],[f411,f74])).
fof(f7373,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f73])).
fof(f7371,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f72])).
fof(f7370,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X7) != s(bool,X9)) ) | $spl63),
  inference(resolution,[],[f411,f70])).
fof(f7368,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl63),
  inference(resolution,[],[f411,f68])).
fof(f5279,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X47) | s(bool,X51) != s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool))) ) | $spl192),
  inference(cnf_transformation,[],[f5279_D])).
fof(f5279_D,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X47) | s(bool,X51) != s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7358,plain,(
  ( ! [X296,X298,X295,X297,X299] : (s(bool,'F_0') = s(bool,X299) | s(bool,X299) != s(bool,sK7('const.pair.,_2'(s(bool,X295),s(X296,X297)),X296,bool)) | s(bool,'T_0') = s(bool,X298) | s(bool,'T_0') = s(bool,X295)) ) | $spl6),
  inference(subsumption_resolution,[],[f7296,f85])).
fof(f7296,plain,(
  ( ! [X296,X298,X295,X297,X299] : (s(bool,'F_0') != s(bool,X298) | s(bool,'F_0') = s(bool,X299) | s(bool,X299) != s(bool,sK7('const.pair.,_2'(s(bool,X295),s(X296,X297)),X296,bool)) | s(bool,'T_0') = s(bool,X298) | s(bool,'T_0') = s(bool,X295)) ) | $spl6),
  inference(superposition,[],[f780,f3129])).
fof(f3616,plain,(
  ( ! [X30,X28,X26,X27] : (~p(s(bool,X30)) | s(bool,'T_0') = s(bool,X26) | s(bool,X30) != s(bool,sK7('const.pair.,_2'(s(bool,X26),s(X27,X28)),X27,bool))) ) | $spl146),
  inference(cnf_transformation,[],[f3616_D])).
fof(f3616_D,plain,(
  ( ! [X30,X28,X26,X27] : (~p(s(bool,X30)) | s(bool,'T_0') = s(bool,X26) | s(bool,X30) != s(bool,sK7('const.pair.,_2'(s(bool,X26),s(X27,X28)),X27,bool))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7350,plain,(
  ( ! [X237,X235,X233,X236,X234] : (p(s(bool,X236)) | ~p(s(bool,X237)) | s(bool,X237) != s(bool,sK7('const.pair.,_2'(s(bool,X233),s(X234,X235)),X234,bool)) | s(bool,'T_0') = s(bool,X233)) ) | $spl6),
  inference(subsumption_resolution,[],[f7280,f81])).
fof(f7280,plain,(
  ( ! [X237,X235,X233,X236,X234] : (p(s(bool,'F_0')) | p(s(bool,X236)) | ~p(s(bool,X237)) | s(bool,X237) != s(bool,sK7('const.pair.,_2'(s(bool,X233),s(X234,X235)),X234,bool)) | s(bool,'T_0') = s(bool,X233)) ) | $spl6),
  inference(superposition,[],[f86,f3129])).
fof(f7348,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,sK7('const.pair.,_2'(s(bool,X199),s(X200,X201)),X200,bool)) | s(bool,X202) != s(bool,X204) | s(bool,'T_0') = s(bool,X199)) ) | $spl6),
  inference(subsumption_resolution,[],[f7274,f81])).
fof(f7274,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (p(s(bool,'F_0')) | ~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,sK7('const.pair.,_2'(s(bool,X199),s(X200,X201)),X200,bool)) | s(bool,X202) != s(bool,X204) | s(bool,'T_0') = s(bool,X199)) ) | $spl6),
  inference(superposition,[],[f74,f3129])).
fof(f7346,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (p(s(bool,X154)) | ~p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,X155) != s(bool,sK7('const.pair.,_2'(s(bool,X151),s(X152,X153)),X152,bool)) | s(bool,'T_0') = s(bool,X151)) ) | $spl6),
  inference(subsumption_resolution,[],[f7266,f81])).
fof(f7266,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (p(s(bool,'F_0')) | p(s(bool,X154)) | ~p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,X155) != s(bool,sK7('const.pair.,_2'(s(bool,X151),s(X152,X153)),X152,bool)) | s(bool,'T_0') = s(bool,X151)) ) | $spl6),
  inference(superposition,[],[f73,f3129])).
fof(f7344,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,sK7('const.pair.,_2'(s(bool,X121),s(X122,X123)),X122,bool)) | s(bool,'T_0') = s(bool,X121)) ) | $spl6),
  inference(subsumption_resolution,[],[f7261,f81])).
fof(f7261,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (p(s(bool,'F_0')) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,sK7('const.pair.,_2'(s(bool,X121),s(X122,X123)),X122,bool)) | s(bool,'T_0') = s(bool,X121)) ) | $spl6),
  inference(superposition,[],[f72,f3129])).
fof(f7341,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X83) != s(bool,sK7('const.pair.,_2'(s(bool,X79),s(X80,X81)),X80,bool)) | s(bool,X82) != s(bool,X84) | s(bool,'T_0') = s(bool,X79)) ) | $spl6),
  inference(subsumption_resolution,[],[f7254,f81])).
fof(f7254,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (p(s(bool,'F_0')) | ~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X83) != s(bool,sK7('const.pair.,_2'(s(bool,X79),s(X80,X81)),X80,bool)) | s(bool,X82) != s(bool,X84) | s(bool,'T_0') = s(bool,X79)) ) | $spl6),
  inference(superposition,[],[f70,f3129])).
fof(f7339,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,sK7('const.pair.,_2'(s(bool,X37),s(X38,X39)),X38,bool)) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(subsumption_resolution,[],[f7247,f81])).
fof(f7247,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,'F_0')) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,sK7('const.pair.,_2'(s(bool,X37),s(X38,X39)),X38,bool)) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f68,f3129])).
fof(f5282,plain,(
  ( ! [X61,X62,X65,X63] : (s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X63) | s(bool,X65) != s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61))) ) | $spl194),
  inference(cnf_transformation,[],[f5282_D])).
fof(f5282_D,plain,(
  ( ! [X61,X62,X65,X63] : (s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X63) | s(bool,X65) != s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7193,plain,(
  ( ! [X300,X296,X298,X297,X299] : (s(bool,'F_0') = s(bool,X300) | s(bool,X300) != s(bool,sK8('const.pair.,_2'(s(X296,X297),s(bool,X298)),bool,X296)) | s(bool,'T_0') = s(bool,X299) | s(bool,'T_0') = s(bool,X298)) ) | $spl6),
  inference(subsumption_resolution,[],[f7132,f85])).
fof(f7132,plain,(
  ( ! [X300,X296,X298,X297,X299] : (s(bool,'F_0') != s(bool,X299) | s(bool,'F_0') = s(bool,X300) | s(bool,X300) != s(bool,sK8('const.pair.,_2'(s(X296,X297),s(bool,X298)),bool,X296)) | s(bool,'T_0') = s(bool,X299) | s(bool,'T_0') = s(bool,X298)) ) | $spl6),
  inference(superposition,[],[f780,f1764])).
fof(f2224,plain,(
  ( ! [X30,X28,X26,X27] : (~p(s(bool,X30)) | s(bool,'T_0') = s(bool,X28) | s(bool,X30) != s(bool,sK8('const.pair.,_2'(s(X26,X27),s(bool,X28)),bool,X26))) ) | $spl104),
  inference(cnf_transformation,[],[f2224_D])).
fof(f2224_D,plain,(
  ( ! [X30,X28,X26,X27] : (~p(s(bool,X30)) | s(bool,'T_0') = s(bool,X28) | s(bool,X30) != s(bool,sK8('const.pair.,_2'(s(X26,X27),s(bool,X28)),bool,X26))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f7185,plain,(
  ( ! [X237,X235,X238,X236,X234] : (p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X238) != s(bool,sK8('const.pair.,_2'(s(X234,X235),s(bool,X236)),bool,X234)) | s(bool,'T_0') = s(bool,X236)) ) | $spl6),
  inference(subsumption_resolution,[],[f7116,f81])).
fof(f7116,plain,(
  ( ! [X237,X235,X238,X236,X234] : (p(s(bool,'F_0')) | p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,X238) != s(bool,sK8('const.pair.,_2'(s(X234,X235),s(bool,X236)),bool,X234)) | s(bool,'T_0') = s(bool,X236)) ) | $spl6),
  inference(superposition,[],[f86,f1764])).
fof(f7183,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (~p(s(bool,X203)) | ~p(s(bool,X204)) | s(bool,X204) != s(bool,sK8('const.pair.,_2'(s(X200,X201),s(bool,X202)),bool,X200)) | s(bool,X203) != s(bool,X205) | s(bool,'T_0') = s(bool,X202)) ) | $spl6),
  inference(subsumption_resolution,[],[f7110,f81])).
fof(f7110,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (p(s(bool,'F_0')) | ~p(s(bool,X203)) | ~p(s(bool,X204)) | s(bool,X204) != s(bool,sK8('const.pair.,_2'(s(X200,X201),s(bool,X202)),bool,X200)) | s(bool,X203) != s(bool,X205) | s(bool,'T_0') = s(bool,X202)) ) | $spl6),
  inference(superposition,[],[f74,f1764])).
fof(f7181,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (p(s(bool,X155)) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X156) != s(bool,sK8('const.pair.,_2'(s(X152,X153),s(bool,X154)),bool,X152)) | s(bool,'T_0') = s(bool,X154)) ) | $spl6),
  inference(subsumption_resolution,[],[f7102,f81])).
fof(f7102,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (p(s(bool,'F_0')) | p(s(bool,X155)) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X156) != s(bool,sK8('const.pair.,_2'(s(X152,X153),s(bool,X154)),bool,X152)) | s(bool,'T_0') = s(bool,X154)) ) | $spl6),
  inference(superposition,[],[f73,f1764])).
fof(f7179,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X125) != s(bool,X127) | s(bool,X126) != s(bool,sK8('const.pair.,_2'(s(X122,X123),s(bool,X124)),bool,X122)) | s(bool,'T_0') = s(bool,X124)) ) | $spl6),
  inference(subsumption_resolution,[],[f7097,f81])).
fof(f7097,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,'F_0')) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X125) != s(bool,X127) | s(bool,X126) != s(bool,sK8('const.pair.,_2'(s(X122,X123),s(bool,X124)),bool,X122)) | s(bool,'T_0') = s(bool,X124)) ) | $spl6),
  inference(superposition,[],[f72,f1764])).
fof(f7176,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,sK8('const.pair.,_2'(s(X80,X81),s(bool,X82)),bool,X80)) | s(bool,X83) != s(bool,X85) | s(bool,'T_0') = s(bool,X82)) ) | $spl6),
  inference(subsumption_resolution,[],[f7090,f81])).
fof(f7090,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,'F_0')) | ~p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,sK8('const.pair.,_2'(s(X80,X81),s(bool,X82)),bool,X80)) | s(bool,X83) != s(bool,X85) | s(bool,'T_0') = s(bool,X82)) ) | $spl6),
  inference(superposition,[],[f70,f1764])).
fof(f7174,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,sK8('const.pair.,_2'(s(X38,X39),s(bool,X40)),bool,X38)) | s(bool,'T_0') = s(bool,X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f7083,f81])).
fof(f7083,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,'F_0')) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,sK8('const.pair.,_2'(s(X38,X39),s(bool,X40)),bool,X38)) | s(bool,'T_0') = s(bool,X40)) ) | $spl6),
  inference(superposition,[],[f68,f1764])).
fof(f840,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl72),
  inference(cnf_transformation,[],[f840_D])).
fof(f840_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7032,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f86])).
fof(f7031,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f248,f74])).
fof(f7028,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f73])).
fof(f7027,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f72])).
fof(f7026,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f248,f70])).
fof(f7024,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f68])).
fof(f3221,plain,(
  ( ! [X33,X36,X34] : (s(bool,X36) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34))) | s(bool,'F_0') = s(bool,sK7(X34,X33,bool)) | p(s(bool,X36))) ) | $spl136),
  inference(cnf_transformation,[],[f3221_D])).
fof(f3221_D,plain,(
  ( ! [X33,X36,X34] : (s(bool,X36) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34))) | s(bool,'F_0') = s(bool,sK7(X34,X33,bool)) | p(s(bool,X36))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6999,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X192),X193))) | s(bool,'F_0') = s(bool,sK7(X193,X192,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6935,f82])).
fof(f6935,plain,(
  ( ! [X194,X192,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X192),X193))) | s(bool,'F_0') = s(bool,sK7(X193,X192,bool))) ) | $spl6),
  inference(superposition,[],[f86,f3000])).
fof(f6997,plain,(
  ( ! [X167,X171,X169,X170,X168] : (p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X169) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X167),X168))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,sK7(X168,X167,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6929,f82])).
fof(f6929,plain,(
  ( ! [X167,X171,X169,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X169) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X167),X168))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,sK7(X168,X167,bool))) ) | $spl6),
  inference(superposition,[],[f74,f3000])).
fof(f6995,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | p(s(bool,X131)) | s(bool,X131) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X127),X128))) | s(bool,'F_0') = s(bool,sK7(X128,X127,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6921,f82])).
fof(f6921,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X129)) | ~p(s(bool,X130)) | p(s(bool,X131)) | s(bool,X131) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X127),X128))) | s(bool,'F_0') = s(bool,sK7(X128,X127,bool))) ) | $spl6),
  inference(superposition,[],[f73,f3000])).
fof(f6993,plain,(
  ( ! [X103,X105,X102,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))) | s(bool,'F_0') = s(bool,sK7(X103,X102,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6916,f82])).
fof(f6916,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))) | s(bool,'F_0') = s(bool,sK7(X103,X102,bool))) ) | $spl6),
  inference(superposition,[],[f72,f3000])).
fof(f6990,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),X68))) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,sK7(X68,X67,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6909,f82])).
fof(f6909,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),X68))) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,sK7(X68,X67,bool))) ) | $spl6),
  inference(superposition,[],[f70,f3000])).
fof(f6988,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),X33))) | s(bool,'F_0') = s(bool,sK7(X33,X32,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f6902,f82])).
fof(f6902,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),X33))) | s(bool,'F_0') = s(bool,sK7(X33,X32,bool))) ) | $spl6),
  inference(superposition,[],[f68,f3000])).
fof(f1838,plain,(
  ( ! [X33,X36,X34] : (s(bool,X36) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),X34))) | s(bool,'F_0') = s(bool,sK8(X34,bool,X33)) | p(s(bool,X36))) ) | $spl90),
  inference(cnf_transformation,[],[f1838_D])).
fof(f1838_D,plain,(
  ( ! [X33,X36,X34] : (s(bool,X36) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),X34))) | s(bool,'F_0') = s(bool,sK8(X34,bool,X33)) | p(s(bool,X36))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f6884,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X192,bool),X193))) | s(bool,'F_0') = s(bool,sK8(X193,bool,X192))) ) | $spl6),
  inference(subsumption_resolution,[],[f6820,f82])).
fof(f6820,plain,(
  ( ! [X194,X192,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X192,bool),X193))) | s(bool,'F_0') = s(bool,sK8(X193,bool,X192))) ) | $spl6),
  inference(superposition,[],[f86,f1655])).
fof(f6882,plain,(
  ( ! [X167,X171,X169,X170,X168] : (p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X169) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X167,bool),X168))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,sK8(X168,bool,X167))) ) | $spl6),
  inference(subsumption_resolution,[],[f6814,f82])).
fof(f6814,plain,(
  ( ! [X167,X171,X169,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X169) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X167,bool),X168))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,sK8(X168,bool,X167))) ) | $spl6),
  inference(superposition,[],[f74,f1655])).
fof(f6880,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | p(s(bool,X131)) | s(bool,X131) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),X128))) | s(bool,'F_0') = s(bool,sK8(X128,bool,X127))) ) | $spl6),
  inference(subsumption_resolution,[],[f6806,f82])).
fof(f6806,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X129)) | ~p(s(bool,X130)) | p(s(bool,X131)) | s(bool,X131) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),X128))) | s(bool,'F_0') = s(bool,sK8(X128,bool,X127))) ) | $spl6),
  inference(superposition,[],[f73,f1655])).
fof(f6878,plain,(
  ( ! [X103,X105,X102,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X102,bool),X103))) | s(bool,'F_0') = s(bool,sK8(X103,bool,X102))) ) | $spl6),
  inference(subsumption_resolution,[],[f6801,f82])).
fof(f6801,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X102,bool),X103))) | s(bool,'F_0') = s(bool,sK8(X103,bool,X102))) ) | $spl6),
  inference(superposition,[],[f72,f1655])).
fof(f6875,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X67,bool),X68))) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,sK8(X68,bool,X67))) ) | $spl6),
  inference(subsumption_resolution,[],[f6794,f82])).
fof(f6794,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X67,bool),X68))) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,sK8(X68,bool,X67))) ) | $spl6),
  inference(superposition,[],[f70,f1655])).
fof(f6873,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),X33))) | s(bool,'F_0') = s(bool,sK8(X33,bool,X32))) ) | $spl6),
  inference(subsumption_resolution,[],[f6787,f82])).
fof(f6787,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),X33))) | s(bool,'F_0') = s(bool,sK8(X33,bool,X32))) ) | $spl6),
  inference(superposition,[],[f68,f1655])).
fof(f2739,plain,(
  ( ! [X14,X12,X10,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | s(bool,X14) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))) | ~p(s(bool,X14)) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))))) ) | $spl126),
  inference(cnf_transformation,[],[f2739_D])).
fof(f2739_D,plain,(
  ( ! [X14,X12,X10,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | s(bool,X14) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))) | ~p(s(bool,X14)) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6670,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) )),
  inference(resolution,[],[f61,f86])).
fof(f61,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f6669,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))))) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))))) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f61,f74])).
fof(f6667,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) )),
  inference(resolution,[],[f61,f73])).
fof(f6665,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) )),
  inference(resolution,[],[f61,f72])).
fof(f6664,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f61,f70])).
fof(f6662,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) )),
  inference(resolution,[],[f61,f68])).
fof(f2738,plain,(
  ( ! [X14,X12,X10,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))) | ~p(s(bool,X14)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))))) ) | $spl124),
  inference(cnf_transformation,[],[f2738_D])).
fof(f2738_D,plain,(
  ( ! [X14,X12,X10,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))) | ~p(s(bool,X14)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6652,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) )),
  inference(resolution,[],[f61,f86])).
fof(f6651,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))))) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))))) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f61,f74])).
fof(f6649,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) )),
  inference(resolution,[],[f61,f73])).
fof(f6647,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) )),
  inference(resolution,[],[f61,f72])).
fof(f6646,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f61,f70])).
fof(f6644,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) )),
  inference(resolution,[],[f61,f68])).
fof(f2737,plain,(
  ( ! [X10,X8,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | s(bool,X13) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))) | p(s(bool,X13)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))))))) ) | $spl122),
  inference(cnf_transformation,[],[f2737_D])).
fof(f2737_D,plain,(
  ( ! [X10,X8,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | s(bool,X13) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))) | p(s(bool,X13)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f6638,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))))) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))) )),
  inference(resolution,[],[f61,f86])).
fof(f6637,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X57,fun(X57,bool)),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X57,fun(X57,bool)),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60)))))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X57,fun(X57,bool)),X58),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f61,f74])).
fof(f6634,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X36,fun(X36,bool)),X37),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X36,fun(X36,bool)),X37),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X42) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X36,fun(X36,bool)),X37),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))) )),
  inference(resolution,[],[f61,f73])).
fof(f6633,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X29,fun(X29,bool)),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X29,fun(X29,bool)),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X29,fun(X29,bool)),X30),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))) )),
  inference(resolution,[],[f61,f72])).
fof(f6632,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f61,f70])).
fof(f6630,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))) )),
  inference(resolution,[],[f61,f68])).
fof(f4410,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X0) != s(bool,X4) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool))) ) | $spl160),
  inference(cnf_transformation,[],[f4410_D])).
fof(f4410_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X0) != s(bool,X4) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6593,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X48),s(X49,X50)),X49,bool)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X48) != s(bool,X52)) ) | $spl148),
  inference(resolution,[],[f4316,f86])).
fof(f4316,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,X101)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X101),s(X102,X103)),X102,bool))) ) | $spl148),
  inference(resolution,[],[f3863,f82])).
fof(f6592,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X42),s(X43,X44)),X43,bool)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl148),
  inference(resolution,[],[f4316,f74])).
fof(f6589,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X24),s(X25,X26)),X25,bool)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X24) != s(bool,X29)) ) | $spl148),
  inference(resolution,[],[f4316,f73])).
fof(f6588,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X19,X20)),X19,bool)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | $spl148),
  inference(resolution,[],[f4316,f72])).
fof(f6587,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X13,X14)),X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl148),
  inference(resolution,[],[f4316,f70])).
fof(f6585,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | $spl148),
  inference(resolution,[],[f4316,f68])).
fof(f2946,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X2) != s(bool,X4) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0))) ) | $spl132),
  inference(cnf_transformation,[],[f2946_D])).
fof(f2946_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X2) != s(bool,X4) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f6550,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X48,X49),s(bool,X50)),bool,X48)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X50) != s(bool,X52)) ) | $spl110),
  inference(resolution,[],[f2874,f86])).
fof(f2874,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X90,X91),s(bool,X89)),bool,X90))) ) | $spl110),
  inference(resolution,[],[f2496,f82])).
fof(f6549,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X42,X43),s(bool,X44)),bool,X42)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl110),
  inference(resolution,[],[f2874,f74])).
fof(f6546,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X24,X25),s(bool,X26)),bool,X24)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | $spl110),
  inference(resolution,[],[f2874,f73])).
fof(f6545,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X18,X19),s(bool,X20)),bool,X18)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl110),
  inference(resolution,[],[f2874,f72])).
fof(f6544,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X12,X13),s(bool,X14)),bool,X12)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl110),
  inference(resolution,[],[f2874,f70])).
fof(f6542,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl110),
  inference(resolution,[],[f2874,f68])).
fof(f2567,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))) | ~p(s(bool,X5))) ) | $spl114),
  inference(cnf_transformation,[],[f2567_D])).
fof(f2567_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))) | ~p(s(bool,X5))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6489,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))) )),
  inference(resolution,[],[f60,f86])).
fof(f6488,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X49,fun(X49,bool)),X50),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f60,f74])).
fof(f6486,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X35,fun(X35,bool)),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))) )),
  inference(resolution,[],[f60,f73])).
fof(f6484,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X21,fun(X21,bool)),X22),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))) )),
  inference(resolution,[],[f60,f72])).
fof(f6483,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X14,fun(X14,bool)),X15),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f60,f70])).
fof(f6481,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) )),
  inference(resolution,[],[f60,f68])).
fof(f2565,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | s(bool,X10) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))))) | p(s(bool,X10))) ) | $spl112),
  inference(cnf_transformation,[],[f2565_D])).
fof(f2565_D,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | s(bool,X10) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))))) | p(s(bool,X10))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f6477,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X66) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X61,fun(X61,bool)),X62),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))))) )),
  inference(resolution,[],[f60,f86])).
fof(f6476,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X54,fun(X54,bool)),X55),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f60,f74])).
fof(f6473,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X33,fun(X33,bool)),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))))) )),
  inference(resolution,[],[f60,f73])).
fof(f6472,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X26,fun(X26,bool)),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))))) )),
  inference(resolution,[],[f60,f72])).
fof(f6471,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X19,fun(X19,bool)),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f60,f70])).
fof(f6469,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))))) )),
  inference(resolution,[],[f60,f68])).
fof(f6421,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,X46) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,sK7('const.pair.,_2'(s(bool,X46),s(X49,X50)),X49,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f86])).
fof(f3609,plain,(
  ( ! [X249,X250,X248] : (~p(s(bool,sK7('const.pair.,_2'(s(bool,X248),s(X249,X250)),X249,bool))) | s(bool,'T_0') = s(bool,X248)) ) | ($spl6 | $spl28)),
  inference(trivial_inequality_removal,[],[f3566])).
fof(f3566,plain,(
  ( ! [X249,X250,X248] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X248),s(X249,X250)),X249,bool))) | s(bool,'T_0') = s(bool,X248)) ) | ($spl6 | $spl28)),
  inference(superposition,[],[f225,f3129])).
fof(f6420,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X40) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK7('const.pair.,_2'(s(bool,X40),s(X43,X44)),X43,bool)) | s(bool,X41) != s(bool,X45)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f74])).
fof(f6418,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X30) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK7('const.pair.,_2'(s(bool,X30),s(X34,X35)),X34,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f73])).
fof(f6416,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X22,X23)),X22,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f72])).
fof(f6415,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X15,X16)),X15,bool)) | s(bool,X13) != s(bool,X17)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f70])).
fof(f6413,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X4,X5)),X4,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f68])).
fof(f3299,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X0) != s(bool,X4) | p(s(bool,X4))) ) | $spl138),
  inference(cnf_transformation,[],[f3299_D])).
fof(f3299_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X0) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f6379,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X48),s(X49,X50)),X49,bool)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X48) != s(bool,X52)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f86])).
fof(f3258,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X10),s(X9,X11)),X9,bool))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3216,f65])).
fof(f3216,plain,(
  ( ! [X216,X215] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X215),X216)))) | s(bool,'F_0') = s(bool,sK7(X216,X215,bool))) ) | ($spl0 | $spl6)),
  inference(trivial_inequality_removal,[],[f3177])).
fof(f3177,plain,(
  ( ! [X216,X215] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X215),X216)))) | s(bool,'F_0') = s(bool,sK7(X216,X215,bool))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f737,f3000])).
fof(f6378,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X42),s(X43,X44)),X43,bool)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f74])).
fof(f6375,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X24),s(X25,X26)),X25,bool)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X24) != s(bool,X29)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f73])).
fof(f6374,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X19,X20)),X19,bool)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f72])).
fof(f6373,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X13,X14)),X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f70])).
fof(f6371,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f68])).
fof(f2314,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))) | ~p(s(bool,X5))) ) | $spl108),
  inference(cnf_transformation,[],[f2314_D])).
fof(f2314_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))) | ~p(s(bool,X5))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f6323,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))) )),
  inference(resolution,[],[f59,f86])).
fof(f6322,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f59,f74])).
fof(f6320,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))) )),
  inference(resolution,[],[f59,f73])).
fof(f6318,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))) )),
  inference(resolution,[],[f59,f72])).
fof(f6317,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f59,f70])).
fof(f6315,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) )),
  inference(resolution,[],[f59,f68])).
fof(f2312,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | s(bool,X10) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))))) | p(s(bool,X10))) ) | $spl106),
  inference(cnf_transformation,[],[f2312_D])).
fof(f2312_D,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | s(bool,X10) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))))) | p(s(bool,X10))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6311,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X66) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))))) )),
  inference(resolution,[],[f59,f86])).
fof(f6310,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f59,f74])).
fof(f6307,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))))) )),
  inference(resolution,[],[f59,f73])).
fof(f6306,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))))) )),
  inference(resolution,[],[f59,f72])).
fof(f6305,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f59,f70])).
fof(f6303,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))))) )),
  inference(resolution,[],[f59,f68])).
fof(f6256,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,X46) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,sK8('const.pair.,_2'(s(X49,X50),s(bool,X46)),bool,X49))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f86])).
fof(f2217,plain,(
  ( ! [X249,X250,X248] : (~p(s(bool,sK8('const.pair.,_2'(s(X248,X249),s(bool,X250)),bool,X248))) | s(bool,'T_0') = s(bool,X250)) ) | ($spl6 | $spl28)),
  inference(trivial_inequality_removal,[],[f2192])).
fof(f2192,plain,(
  ( ! [X249,X250,X248] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,sK8('const.pair.,_2'(s(X248,X249),s(bool,X250)),bool,X248))) | s(bool,'T_0') = s(bool,X250)) ) | ($spl6 | $spl28)),
  inference(superposition,[],[f225,f1764])).
fof(f6255,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X40) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK8('const.pair.,_2'(s(X43,X44),s(bool,X40)),bool,X43)) | s(bool,X41) != s(bool,X45)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f74])).
fof(f6253,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X30) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK8('const.pair.,_2'(s(X34,X35),s(bool,X30)),bool,X34))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f73])).
fof(f6251,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,sK8('const.pair.,_2'(s(X22,X23),s(bool,X18)),bool,X22))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f72])).
fof(f6250,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,sK8('const.pair.,_2'(s(X15,X16),s(bool,X12)),bool,X15)) | s(bool,X13) != s(bool,X17)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f70])).
fof(f6248,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK8('const.pair.,_2'(s(X4,X5),s(bool,X0)),bool,X4))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f68])).
fof(f1997,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X2) != s(bool,X4) | p(s(bool,X4))) ) | $spl98),
  inference(cnf_transformation,[],[f1997_D])).
fof(f1997_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X2) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6216,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X48,X49),s(bool,X50)),bool,X48)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X50) != s(bool,X52)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f86])).
fof(f1873,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X9,X10),s(bool,X11)),bool,X9))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1833,f64])).
fof(f1833,plain,(
  ( ! [X216,X215] : (p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X216)))) | s(bool,'F_0') = s(bool,sK8(X216,bool,X215))) ) | ($spl0 | $spl6)),
  inference(trivial_inequality_removal,[],[f1812])).
fof(f1812,plain,(
  ( ! [X216,X215] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X216)))) | s(bool,'F_0') = s(bool,sK8(X216,bool,X215))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f737,f1655])).
fof(f6215,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X42,X43),s(bool,X44)),bool,X42)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f74])).
fof(f6212,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X24,X25),s(bool,X26)),bool,X24)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f73])).
fof(f6211,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X18,X19),s(bool,X20)),bool,X18)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f72])).
fof(f6210,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X12,X13),s(bool,X14)),bool,X12)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f70])).
fof(f6208,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f68])).
fof(f6197,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,sK7(X40,X41,bool)) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),X40)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f86])).
fof(f6196,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,sK7(X35,X36,bool)) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X36),X35))) | s(bool,X38) != s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f74])).
fof(f6193,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,sK7(X20,X21,bool)) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X20)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f73])).
fof(f6192,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,sK7(X15,X16,bool)) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X16),X15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f72])).
fof(f6191,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,sK7(X10,X11,bool)) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),X10))) | s(bool,X13) != s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f70])).
fof(f6189,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK7(X0,X1,bool)) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),X0)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f68])).
fof(f6178,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,sK8(X40,bool,X41)) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),X40)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f86])).
fof(f6177,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,sK8(X35,bool,X36)) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),X35))) | s(bool,X38) != s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f74])).
fof(f6174,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,sK8(X20,bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X20)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f73])).
fof(f6173,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,sK8(X15,bool,X16)) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),X15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f72])).
fof(f6172,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,sK8(X10,bool,X11)) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X11,bool),X10))) | s(bool,X13) != s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f70])).
fof(f6170,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK8(X0,bool,X1)) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),X0)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f68])).
fof(f799,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl70),
  inference(cnf_transformation,[],[f799_D])).
fof(f799_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f6142,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl6),
  inference(resolution,[],[f494,f86])).
fof(f6141,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f494,f74])).
fof(f6138,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f494,f73])).
fof(f6137,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f494,f72])).
fof(f6136,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f494,f70])).
fof(f6134,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f494,f68])).
fof(f1944,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | s(bool,X7) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5))) | ~p(s(bool,X7)) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4))))) ) | $spl96),
  inference(cnf_transformation,[],[f1944_D])).
fof(f1944_D,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | s(bool,X7) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5))) | ~p(s(bool,X7)) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6080,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X72,fun(X72,bool)),X73),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X76)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X72,fun(X72,bool)),X73),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X77),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X76)))) | p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X74),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X77)))) )),
  inference(resolution,[],[f62,f86])).
fof(f6079,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X63,fun(X63,bool)),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X65),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X67)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X63,fun(X63,bool)),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X68),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X67)))) | ~p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X68))) | s(bool,X69) != s(bool,X71)) )),
  inference(resolution,[],[f62,f74])).
fof(f6077,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X45,fun(X45,bool)),X46),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X47),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X49)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X45,fun(X45,bool)),X46),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X50),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X49)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X50)))) )),
  inference(resolution,[],[f62,f73])).
fof(f6075,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X27,fun(X27,bool)),X28),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X29),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X31)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X27,fun(X27,bool)),X28),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X31)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X29),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X32)))) )),
  inference(resolution,[],[f62,f72])).
fof(f6074,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X18,fun(X18,bool)),X19),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X20),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X21),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X22)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X18,fun(X18,bool)),X19),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X23),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X21),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X22)))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X23))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f62,f70])).
fof(f6072,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5)))) )),
  inference(resolution,[],[f62,f68])).
fof(f1943,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11))) | ~p(s(bool,X14)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12))))) ) | $spl94),
  inference(cnf_transformation,[],[f1943_D])).
fof(f1943_D,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11))) | ~p(s(bool,X14)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6069,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84)))) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X86) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) )),
  inference(resolution,[],[f62,f86])).
fof(f6068,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X77,X75,X73] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75)))) | ~p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74))) | s(bool,X76) != s(bool,X78)) )),
  inference(resolution,[],[f62,f74])).
fof(f6066,plain,(
  ( ! [X59,X57,X54,X52,X60,X58,X56,X55,X53] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X54),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X55),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X56)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X54),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X57),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X55),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X56)))) )),
  inference(resolution,[],[f62,f73])).
fof(f6064,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) )),
  inference(resolution,[],[f62,f72])).
fof(f6063,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f62,f70])).
fof(f6061,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) )),
  inference(resolution,[],[f62,f68])).
fof(f1942,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9))))) ) | $spl92),
  inference(cnf_transformation,[],[f1942_D])).
fof(f1942_D,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6056,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81)))) | p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X86) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) )),
  inference(resolution,[],[f62,f86])).
fof(f6055,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X77,X75,X73] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72)))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X76) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74))) | s(bool,X77) != s(bool,X78)) )),
  inference(resolution,[],[f62,f74])).
fof(f6052,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X43,fun(X43,bool)),X44),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X45),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X46),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X47)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X48),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X45)))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X43,fun(X43,bool)),X44),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X46),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X47)))) )),
  inference(resolution,[],[f62,f73])).
fof(f6051,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) )),
  inference(resolution,[],[f62,f72])).
fof(f6050,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f62,f70])).
fof(f6048,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) )),
  inference(resolution,[],[f62,f68])).
fof(f6002,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X255,X256)),X255,bool))) ) | $spl204),
  inference(cnf_transformation,[],[f6002_D])).
fof(f6002_D,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X255,X256)),X255,bool))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f5873,plain,(
  s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')),bool,bool)) | $spl201),
  inference(cnf_transformation,[],[f5873_D])).
fof(f5873_D,plain,(
  s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')),bool,bool)) <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f5952,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')),bool,bool)) | s(bool,'T_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X255,X256)),X255,bool))) ) | ($spl0 | $spl6 | $spl116)),
  inference(superposition,[],[f2636,f3273])).
fof(f3273,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X71,X72)),X71,bool))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f81])).
fof(f2636,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(bool,X14),s(bool,X14)),bool,bool)) | s(bool,'T_0') = s(bool,X14)) ) | $spl116),
  inference(superposition,[],[f2625,f85])).
fof(f6001,plain,(
  ( ! [X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X229,X230)),X229,bool))) ) | $spl202),
  inference(cnf_transformation,[],[f6001_D])).
fof(f6001_D,plain,(
  ( ! [X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X229,X230)),X229,bool))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f6000,plain,(
  ( ! [X231,X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X229,X230)),X229,bool)) | s(bool,'T_0') = s(bool,X231)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f5945,f85])).
fof(f5945,plain,(
  ( ! [X231,X229,X230,X232] : (s(bool,'F_0') != s(bool,X231) | s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X229,X230)),X229,bool)) | s(bool,'T_0') = s(bool,X231)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f780,f3273])).
fof(f5871,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') = s(bool,sK8('const.pair.,_2'(s(X255,X256),s(bool,'F_0')),bool,X255))) ) | $spl198),
  inference(cnf_transformation,[],[f5871_D])).
fof(f5871_D,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') = s(bool,sK8('const.pair.,_2'(s(X255,X256),s(bool,'F_0')),bool,X255))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f3302,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X10,X11)),X10,bool))) ) | $spl140),
  inference(cnf_transformation,[],[f3302_D])).
fof(f3302_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X10,X11)),X10,bool))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f5824,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X263,X264)),X263,bool)) | s(bool,'T_0') = s(bool,sK8('const.pair.,_2'(s(X261,X262),s(bool,'F_0')),bool,X261))) ) | ($spl0 | $spl6 | $spl47)),
  inference(superposition,[],[f3611,f1981])).
fof(f1981,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X65,X66),s(bool,'F_0')),bool,X65))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f81])).
fof(f5822,plain,(
  ( ! [X255,X256] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')),bool,bool)) | s(bool,'T_0') = s(bool,sK8('const.pair.,_2'(s(X255,X256),s(bool,'F_0')),bool,X255))) ) | ($spl0 | $spl6 | $spl116)),
  inference(superposition,[],[f2636,f1981])).
fof(f5870,plain,(
  ( ! [X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X229,X230),s(bool,'F_0')),bool,X229))) ) | $spl196),
  inference(cnf_transformation,[],[f5870_D])).
fof(f5870_D,plain,(
  ( ! [X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X229,X230),s(bool,'F_0')),bool,X229))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f5869,plain,(
  ( ! [X231,X229,X230,X232] : (s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X229,X230),s(bool,'F_0')),bool,X229)) | s(bool,'T_0') = s(bool,X231)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f5815,f85])).
fof(f5815,plain,(
  ( ! [X231,X229,X230,X232] : (s(bool,'F_0') != s(bool,X231) | s(bool,'F_0') = s(bool,X232) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X229,X230),s(bool,'F_0')),bool,X229)) | s(bool,'T_0') = s(bool,X231)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f780,f1981])).
fof(f1615,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4))) | ~p(s(bool,X6)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3))))) ) | $spl86),
  inference(cnf_transformation,[],[f1615_D])).
fof(f1615_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4))) | ~p(s(bool,X6)) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5698,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),X67)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X68),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X68)))) )),
  inference(resolution,[],[f63,f86])).
fof(f5697,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X58),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),X59)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X60),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f63,f74])).
fof(f5695,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X40,fun(X40,bool)),X41),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s(fun('type.mlstring.mlstring',fun('type.list.list'(X40),X40)),X43)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X40,fun(X40,bool)),X41),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X44),s(fun('type.mlstring.mlstring',fun('type.list.list'(X40),X40)),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X44)))) )),
  inference(resolution,[],[f63,f73])).
fof(f5693,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X24,fun(X24,bool)),X25),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X24),X24)),X27)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X24,fun(X24,bool)),X25),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X24),X24)),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X28)))) )),
  inference(resolution,[],[f63,f72])).
fof(f5692,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X18),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X19)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X20),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X19)))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X18),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X20))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f63,f70])).
fof(f5690,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4)))) )),
  inference(resolution,[],[f63,f68])).
fof(f1614,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9))) | ~p(s(bool,X12)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10))))) ) | $spl84),
  inference(cnf_transformation,[],[f1614_D])).
fof(f1614_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9))) | ~p(s(bool,X12)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5687,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) )),
  inference(resolution,[],[f63,f86])).
fof(f5686,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66)))) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f63,f74])).
fof(f5684,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X49)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X48),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X50)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X50),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X49)))) )),
  inference(resolution,[],[f63,f73])).
fof(f5682,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) )),
  inference(resolution,[],[f63,f72])).
fof(f5681,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26)))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f63,f70])).
fof(f5679,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) )),
  inference(resolution,[],[f63,f68])).
fof(f1613,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9))) | p(s(bool,X12)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8))))) ) | $spl82),
  inference(cnf_transformation,[],[f1613_D])).
fof(f1613_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9))) | p(s(bool,X12)) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5674,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72)))) | p(s(bool,X75)) | p(s(bool,X76)) | s(bool,X76) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) )),
  inference(resolution,[],[f63,f86])).
fof(f5673,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65))) | s(bool,X68) != s(bool,X69)) )),
  inference(resolution,[],[f63,f74])).
fof(f5670,plain,(
  ( ! [X39,X45,X43,X41,X38,X44,X42,X40] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X38,fun(X38,bool)),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X40),s(fun('type.mlstring.mlstring',fun('type.list.list'(X38),X38)),X41)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X40)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X38,fun(X38,bool)),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s(fun('type.mlstring.mlstring',fun('type.list.list'(X38),X38)),X41)))) )),
  inference(resolution,[],[f63,f73])).
fof(f5669,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32)))) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) )),
  inference(resolution,[],[f63,f72])).
fof(f5668,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25))) | s(bool,X28) != s(bool,X29)) )),
  inference(resolution,[],[f63,f70])).
fof(f5666,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8)))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) )),
  inference(resolution,[],[f63,f68])).
fof(f278,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl42),
  inference(cnf_transformation,[],[f278_D])).
fof(f278_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5624,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl22),
  inference(resolution,[],[f222,f86])).
fof(f5623,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f222,f74])).
fof(f5620,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f222,f73])).
fof(f5619,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f222,f72])).
fof(f5618,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f222,f70])).
fof(f5616,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f222,f68])).
fof(f225,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X5))) ) | $spl28),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X5))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5611,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f81,f86])).
fof(f5610,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f81,f74])).
fof(f5608,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f81,f73])).
fof(f5606,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f81,f72])).
fof(f5605,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f81,f70])).
fof(f5603,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f81,f68])).
fof(f5546,plain,(
  $false | ($spl58 | $spl160)),
  inference(subsumption_resolution,[],[f5545,f390])).
fof(f5545,plain,(
  ( ! [X105] : (s(bool,'T_0') != s(bool,X105)) ) | ($spl58 | $spl160)),
  inference(subsumption_resolution,[],[f5323,f390])).
fof(f5323,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,X105) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X105),s(X106,X107)),X106,bool))) ) | ($spl58 | $spl160)),
  inference(backward_demodulation,[],[f390,f5121])).
fof(f5121,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X105),s(X106,X107)),X106,bool))) ) | $spl160),
  inference(resolution,[],[f4410,f81])).
fof(f5544,plain,(
  $false | ($spl58 | $spl132)),
  inference(subsumption_resolution,[],[f5543,f390])).
fof(f5543,plain,(
  ( ! [X105] : (s(bool,'T_0') != s(bool,X105)) ) | ($spl58 | $spl132)),
  inference(subsumption_resolution,[],[f5322,f390])).
fof(f5322,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,X105) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X106,X107),s(bool,X105)),bool,X106))) ) | ($spl58 | $spl132)),
  inference(backward_demodulation,[],[f390,f5002])).
fof(f5002,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X106,X107),s(bool,X105)),bool,X106))) ) | $spl132),
  inference(resolution,[],[f2946,f81])).
fof(f5538,plain,(
  $false | ($spl58 | $spl150)),
  inference(subsumption_resolution,[],[f5315,f390])).
fof(f5315,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'T_0'),s(X1,X2)),X1,bool))) ) | ($spl58 | $spl150)),
  inference(backward_demodulation,[],[f390,f4070])).
fof(f5536,plain,(
  $false | ($spl58 | $spl116)),
  inference(subsumption_resolution,[],[f5306,f390])).
fof(f5306,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,'T_0')),bool,X1))) ) | ($spl58 | $spl116)),
  inference(backward_demodulation,[],[f390,f2625])).
fof(f5533,plain,(
  $false | $spl58),
  inference(subsumption_resolution,[],[f5290,f82])).
fof(f5290,plain,(
  ~p(s(bool,'T_0')) | $spl58),
  inference(backward_demodulation,[],[f390,f81])).
fof(f5517,plain,(
  $false | ($spl47 | $spl58)),
  inference(trivial_inequality_removal,[],[f5491])).
fof(f5491,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl47 | $spl58)),
  inference(superposition,[],[f295,f390])).
fof(f5281,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X65) | s(bool,X65) != s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61)) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) ) | $spl6),
  inference(subsumption_resolution,[],[f5248,f85])).
fof(f5248,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'F_0') != s(bool,X64) | s(bool,'F_0') = s(bool,X65) | s(bool,X65) != s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61)) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) ) | $spl6),
  inference(superposition,[],[f780,f1764])).
fof(f5278,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X51) != s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool)) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X47)) ) | $spl6),
  inference(subsumption_resolution,[],[f5245,f85])).
fof(f5245,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) != s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool)) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X47)) ) | $spl6),
  inference(superposition,[],[f780,f3129])).
fof(f5275,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') = s(bool,X19) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(subsumption_resolution,[],[f5238,f85])).
fof(f5238,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'F_0') = s(bool,X19) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(superposition,[],[f780,f415])).
fof(f5212,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f86])).
fof(f5211,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X22) != s(bool,X24)) ) | $spl61),
  inference(resolution,[],[f409,f74])).
fof(f5209,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f73])).
fof(f5207,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f72])).
fof(f5206,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X7) != s(bool,X9)) ) | $spl61),
  inference(resolution,[],[f409,f70])).
fof(f5204,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) ) | $spl61),
  inference(resolution,[],[f409,f68])).
fof(f5156,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X13),s(X14,X15)),X14,bool))) ) | $spl186),
  inference(cnf_transformation,[],[f5156_D])).
fof(f5156_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X13),s(X14,X15)),X14,bool))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f5155,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X13),s(X14,X15)),X14,bool))) ) | ($spl0 | $spl160)),
  inference(subsumption_resolution,[],[f5136,f737])).
fof(f5136,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X13),s(X14,X15)),X14,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl160),
  inference(superposition,[],[f4410,f85])).
fof(f5153,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) | $spl184),
  inference(cnf_transformation,[],[f5153_D])).
fof(f5153_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f5113,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X58,X59)),X58,bool)) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl160),
  inference(resolution,[],[f4410,f86])).
fof(f5112,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X51,X52)),X51,bool)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl160),
  inference(resolution,[],[f4410,f74])).
fof(f5109,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X28),s(X30,X31)),X30,bool)) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl160),
  inference(resolution,[],[f4410,f73])).
fof(f5108,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X23,X24)),X23,bool)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl160),
  inference(resolution,[],[f4410,f72])).
fof(f5107,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X16,X17)),X16,bool)) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl160),
  inference(resolution,[],[f4410,f70])).
fof(f5105,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool)) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl160),
  inference(resolution,[],[f4410,f68])).
fof(f5056,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f86])).
fof(f5055,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f249,f74])).
fof(f5052,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f73])).
fof(f5051,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f72])).
fof(f5050,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f249,f70])).
fof(f5048,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f249,f68])).
fof(f5037,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X13)),bool,X14))) ) | $spl180),
  inference(cnf_transformation,[],[f5037_D])).
fof(f5037_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X13)),bool,X14))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f5036,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X13)),bool,X14))) ) | ($spl0 | $spl132)),
  inference(subsumption_resolution,[],[f5017,f737])).
fof(f5017,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X13)),bool,X14)) | s(bool,'T_0') = s(bool,X12)) ) | $spl132),
  inference(superposition,[],[f2946,f85])).
fof(f5034,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) | $spl178),
  inference(cnf_transformation,[],[f5034_D])).
fof(f5034_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f4994,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X58,X59),s(bool,X56)),bool,X58)) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl132),
  inference(resolution,[],[f2946,f86])).
fof(f4993,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X51,X52),s(bool,X49)),bool,X51)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl132),
  inference(resolution,[],[f2946,f74])).
fof(f4990,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X30,X31),s(bool,X28)),bool,X30)) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl132),
  inference(resolution,[],[f2946,f73])).
fof(f4989,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,X21)),bool,X23)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl132),
  inference(resolution,[],[f2946,f72])).
fof(f4988,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X16,X17),s(bool,X14)),bool,X16)) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl132),
  inference(resolution,[],[f2946,f70])).
fof(f4986,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2)) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl132),
  inference(resolution,[],[f2946,f68])).
fof(f4930,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,sK7('const.pair.,_2'(s(bool,X57),s(X58,X59)),X58,bool)) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f86])).
fof(f4929,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,sK7('const.pair.,_2'(s(bool,X50),s(X51,X52)),X51,bool)) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f74])).
fof(f4926,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,sK7('const.pair.,_2'(s(bool,X29),s(X30,X31)),X30,bool)) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f73])).
fof(f4925,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,sK7('const.pair.,_2'(s(bool,X22),s(X23,X24)),X23,bool)) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f72])).
fof(f4924,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,sK7('const.pair.,_2'(s(bool,X15),s(X16,X17)),X16,bool)) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f70])).
fof(f4922,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK7('const.pair.,_2'(s(bool,X1),s(X2,X3)),X2,bool)) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f3876,f68])).
fof(f4866,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,sK8('const.pair.,_2'(s(X57,X58),s(bool,X59)),bool,X57)) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f86])).
fof(f4865,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,sK8('const.pair.,_2'(s(X50,X51),s(bool,X52)),bool,X50)) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f74])).
fof(f4862,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,sK8('const.pair.,_2'(s(X29,X30),s(bool,X31)),bool,X29)) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f73])).
fof(f4861,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,sK8('const.pair.,_2'(s(X22,X23),s(bool,X24)),bool,X22)) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f72])).
fof(f4860,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,sK8('const.pair.,_2'(s(X15,X16),s(bool,X17)),bool,X15)) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f70])).
fof(f4858,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,X3)),bool,X1)) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | ($spl0 | $spl6 | $spl47)),
  inference(resolution,[],[f2509,f68])).
fof(f4806,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X57,X58)),X57,bool)) | s(bool,X56) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | $spl136),
  inference(resolution,[],[f4546,f86])).
fof(f4805,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X50,X51)),X50,bool)) | s(bool,X49) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl136),
  inference(resolution,[],[f4546,f74])).
fof(f4802,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X28),s(X29,X30)),X29,bool)) | s(bool,X28) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | $spl136),
  inference(resolution,[],[f4546,f73])).
fof(f4801,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X22,X23)),X22,bool)) | s(bool,X21) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl136),
  inference(resolution,[],[f4546,f72])).
fof(f4800,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X15,X16)),X15,bool)) | s(bool,X14) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl136),
  inference(resolution,[],[f4546,f70])).
fof(f4798,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | s(bool,X0) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | $spl136),
  inference(resolution,[],[f4546,f68])).
fof(f4746,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X56,X57),s(bool,X58)),bool,X56)) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | $spl90),
  inference(resolution,[],[f4449,f86])).
fof(f4745,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X49,X50),s(bool,X51)),bool,X49)) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl90),
  inference(resolution,[],[f4449,f74])).
fof(f4742,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X28,X29),s(bool,X30)),bool,X28)) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) ) | $spl90),
  inference(resolution,[],[f4449,f73])).
fof(f4741,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X21,X22),s(bool,X23)),bool,X21)) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl90),
  inference(resolution,[],[f4449,f72])).
fof(f4740,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X14,X15),s(bool,X16)),bool,X14)) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl90),
  inference(resolution,[],[f4449,f70])).
fof(f4738,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) ) | $spl90),
  inference(resolution,[],[f4449,f68])).
fof(f4686,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) != s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X58,X59)),X58,bool)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl146),
  inference(resolution,[],[f3616,f86])).
fof(f4685,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) != s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X51,X52)),X51,bool)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl146),
  inference(resolution,[],[f3616,f74])).
fof(f4683,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X35),s(X37,X38)),X37,bool)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl146),
  inference(resolution,[],[f3616,f73])).
fof(f4681,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) != s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X23,X24)),X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl146),
  inference(resolution,[],[f3616,f72])).
fof(f4680,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X16,X17)),X16,bool)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl146),
  inference(resolution,[],[f3616,f70])).
fof(f4678,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl146),
  inference(resolution,[],[f3616,f68])).
fof(f4626,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) != s(bool,sK8('const.pair.,_2'(s(X58,X59),s(bool,X56)),bool,X58)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl104),
  inference(resolution,[],[f2224,f86])).
fof(f4625,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) != s(bool,sK8('const.pair.,_2'(s(X51,X52),s(bool,X49)),bool,X51)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl104),
  inference(resolution,[],[f2224,f74])).
fof(f4623,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X37,X38),s(bool,X35)),bool,X37)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl104),
  inference(resolution,[],[f2224,f73])).
fof(f4621,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,X21)),bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl104),
  inference(resolution,[],[f2224,f72])).
fof(f4620,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,sK8('const.pair.,_2'(s(X16,X17),s(bool,X14)),bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl104),
  inference(resolution,[],[f2224,f70])).
fof(f4618,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl104),
  inference(resolution,[],[f2224,f68])).
fof(f4550,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X5))) | s(bool,'F_0') = s(bool,sK7(X5,X4,bool))) ) | $spl164),
  inference(cnf_transformation,[],[f4550_D])).
fof(f4550_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X5))) | s(bool,'F_0') = s(bool,sK7(X5,X4,bool))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4567,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,sK7(X65,X64,bool)) | s(bool,'T_0') = s(bool,X63)) ) | ($spl6 | $spl28 | $spl136)),
  inference(subsumption_resolution,[],[f4536,f2217])).
fof(f4536,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,sK7(X65,X64,bool)) | p(s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl6 | $spl136)),
  inference(superposition,[],[f3221,f1764])).
fof(f4564,plain,(
  ( ! [X47,X50,X51] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51))) | s(bool,'F_0') = s(bool,sK7(X51,X50,bool)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl28 | $spl136)),
  inference(subsumption_resolution,[],[f4533,f3609])).
fof(f4533,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51))) | s(bool,'F_0') = s(bool,sK7(X51,X50,bool)) | p(s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl136)),
  inference(superposition,[],[f3221,f3129])).
fof(f4549,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X5))) | s(bool,'F_0') = s(bool,sK7(X5,X4,bool)) | p(s(bool,X3))) ) | ($spl0 | $spl136)),
  inference(subsumption_resolution,[],[f4522,f737])).
fof(f4522,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X5))) | s(bool,'F_0') = s(bool,sK7(X5,X4,bool)) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) ) | $spl136),
  inference(superposition,[],[f3221,f85])).
fof(f4453,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5))) | s(bool,'F_0') = s(bool,sK8(X5,bool,X4))) ) | $spl162),
  inference(cnf_transformation,[],[f4453_D])).
fof(f4453_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5))) | s(bool,'F_0') = s(bool,sK8(X5,bool,X4))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f4470,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X64,bool),X65))) | s(bool,'F_0') = s(bool,sK8(X65,bool,X64)) | s(bool,'T_0') = s(bool,X63)) ) | ($spl6 | $spl28 | $spl90)),
  inference(subsumption_resolution,[],[f4439,f2217])).
fof(f4439,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X64,bool),X65))) | s(bool,'F_0') = s(bool,sK8(X65,bool,X64)) | p(s(bool,sK8('const.pair.,_2'(s(X61,X62),s(bool,X63)),bool,X61))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl6 | $spl90)),
  inference(superposition,[],[f1838,f1764])).
fof(f4467,plain,(
  ( ! [X47,X50,X51] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X50,bool),X51))) | s(bool,'F_0') = s(bool,sK8(X51,bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl28 | $spl90)),
  inference(subsumption_resolution,[],[f4436,f3609])).
fof(f4436,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X50,bool),X51))) | s(bool,'F_0') = s(bool,sK8(X51,bool,X50)) | p(s(bool,sK7('const.pair.,_2'(s(bool,X47),s(X48,X49)),X48,bool))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl6 | $spl90)),
  inference(superposition,[],[f1838,f3129])).
fof(f4452,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5))) | s(bool,'F_0') = s(bool,sK8(X5,bool,X4)) | p(s(bool,X3))) ) | ($spl0 | $spl90)),
  inference(subsumption_resolution,[],[f4425,f737])).
fof(f4425,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5))) | s(bool,'F_0') = s(bool,sK8(X5,bool,X4)) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) ) | $spl90),
  inference(superposition,[],[f1838,f85])).
fof(f4372,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X48),s(X49,X50)),X49,bool)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X48) != s(bool,X52)) ) | $spl148),
  inference(resolution,[],[f4316,f86])).
fof(f4371,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X42),s(X43,X44)),X43,bool)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl148),
  inference(resolution,[],[f4316,f74])).
fof(f4368,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X24),s(X25,X26)),X25,bool)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X24) != s(bool,X29)) ) | $spl148),
  inference(resolution,[],[f4316,f73])).
fof(f4367,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X19,X20)),X19,bool)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | $spl148),
  inference(resolution,[],[f4316,f72])).
fof(f4366,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X13,X14)),X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl148),
  inference(resolution,[],[f4316,f70])).
fof(f4364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | $spl148),
  inference(resolution,[],[f4316,f68])).
fof(f358,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) | $spl54),
  inference(cnf_transformation,[],[f358_D])).
fof(f358_D,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4348,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool))) ) | $spl158),
  inference(cnf_transformation,[],[f4348_D])).
fof(f4348_D,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4347,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool)) | s(bool,'F_0') = s(bool,X8)) ) | ($spl22 | $spl148)),
  inference(subsumption_resolution,[],[f4327,f222])).
fof(f4327,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,X9),s(X10,X11)),X10,bool)) | s(bool,'F_0') = s(bool,X8)) ) | $spl148),
  inference(superposition,[],[f3863,f85])).
fof(f4308,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X56)) | s(bool,X57) != s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X58,X59)),X58,bool)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl148),
  inference(resolution,[],[f3863,f86])).
fof(f4307,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,sK7('const.pair.,_2'(s(bool,X49),s(X51,X52)),X51,bool)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl148),
  inference(resolution,[],[f3863,f74])).
fof(f4305,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X35),s(X37,X38)),X37,bool)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl148),
  inference(resolution,[],[f3863,f73])).
fof(f4303,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X22) != s(bool,sK7('const.pair.,_2'(s(bool,X21),s(X23,X24)),X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl148),
  inference(resolution,[],[f3863,f72])).
fof(f4302,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X15) != s(bool,sK7('const.pair.,_2'(s(bool,X14),s(X16,X17)),X16,bool)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl148),
  inference(resolution,[],[f3863,f70])).
fof(f4300,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X2,X3)),X2,bool)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl148),
  inference(resolution,[],[f3863,f68])).
fof(f409,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | $spl61),
  inference(cnf_transformation,[],[f409_D])).
fof(f409_D,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f411,plain,(
  ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | $spl63),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f4258,plain,(
  ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f4257,f65])).
fof(f4257,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f4256,f64])).
fof(f4256,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(subsumption_resolution,[],[f4156,f81])).
fof(f4156,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(superposition,[],[f61,f415])).
fof(f2629,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X22))) | s(bool,'F_0') = s(bool,sK8(X22,bool,X21))) ) | $spl120),
  inference(cnf_transformation,[],[f2629_D])).
fof(f2629_D,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X22))) | s(bool,'F_0') = s(bool,sK8(X22,bool,X21))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4070,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool))) ) | $spl150),
  inference(cnf_transformation,[],[f4070_D])).
fof(f4070_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4031,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X34,X35)),X34,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),X33))) | s(bool,'F_0') = s(bool,sK8(X33,bool,X32))) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f3414,f1655])).
fof(f3414,plain,(
  ( ! [X243,X244,X245] : (s(bool,X245) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X243,X244)),X243,bool)) | s(bool,'F_0') = s(bool,X245)) ) | ($spl0 | $spl6 | $spl70)),
  inference(trivial_inequality_removal,[],[f3387])).
fof(f3387,plain,(
  ( ! [X243,X244,X245] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X245) != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X243,X244)),X243,bool)) | s(bool,'F_0') = s(bool,X245)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f1469,f3273])).
fof(f1469,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,X65) != s(bool,X66) | s(bool,'F_0') = s(bool,X66)) ) | $spl70),
  inference(resolution,[],[f799,f81])).
fof(f2628,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),'const.pair.,_2'(s(X16,X17),s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) ) | $spl118),
  inference(cnf_transformation,[],[f2628_D])).
fof(f2628_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),'const.pair.,_2'(s(X16,X17),s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4030,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X30,X31)),X30,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29))))) | s(bool,'F_0') = s(bool,X29)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f3414,f1526])).
fof(f4075,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),X24))) | s(bool,'F_0') = s(bool,sK7(X24,X23,bool))) ) | $spl154),
  inference(cnf_transformation,[],[f4075_D])).
fof(f4075_D,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),X24))) | s(bool,'F_0') = s(bool,sK7(X24,X23,bool))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4029,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X25,X26)),X25,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),X24))) | s(bool,'F_0') = s(bool,sK7(X24,X23,bool))) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f3414,f3000])).
fof(f4074,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X19),s(X18,X20))))) | s(bool,'F_0') = s(bool,X19)) ) | $spl152),
  inference(cnf_transformation,[],[f4074_D])).
fof(f4074_D,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X19),s(X18,X20))))) | s(bool,'F_0') = s(bool,X19)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f4028,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X21,X22)),X21,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X19),s(X18,X20))))) | s(bool,'F_0') = s(bool,X19)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f3414,f1545])).
fof(f4069,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6 | $spl70)),
  inference(duplicate_literal_removal,[],[f4022])).
fof(f4022,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X1,X2)),X1,bool)) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f3414,f85])).
fof(f3863,plain,(
  ( ! [X33,X36,X34,X32] : (~p(s(bool,X36)) | p(s(bool,X32)) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X32),s(X33,X34)),X33,bool))) ) | $spl148),
  inference(cnf_transformation,[],[f3863_D])).
fof(f3863_D,plain,(
  ( ! [X33,X36,X34,X32] : (~p(s(bool,X36)) | p(s(bool,X32)) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X32),s(X33,X34)),X33,bool))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3874,plain,(
  ( ! [X231,X229,X230,X228,X232] : (p(s(bool,X228)) | p(s(bool,X231)) | ~p(s(bool,X232)) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,X228),s(X229,X230)),X229,bool))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3802,f737])).
fof(f3802,plain,(
  ( ! [X231,X229,X230,X228,X232] : (p(s(bool,X228)) | p(s(bool,X231)) | ~p(s(bool,X232)) | s(bool,X232) != s(bool,sK7('const.pair.,_2'(s(bool,X228),s(X229,X230)),X229,bool)) | s(bool,'T_0') = s(bool,X228)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f86,f3611])).
fof(f3872,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X194)) | ~p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X198) != s(bool,sK7('const.pair.,_2'(s(bool,X194),s(X195,X196)),X195,bool)) | s(bool,X197) != s(bool,X199)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3795,f737])).
fof(f3795,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X194)) | ~p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X198) != s(bool,sK7('const.pair.,_2'(s(bool,X194),s(X195,X196)),X195,bool)) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') = s(bool,X194)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f74,f3611])).
fof(f3870,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (p(s(bool,X146)) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,sK7('const.pair.,_2'(s(bool,X146),s(X147,X148)),X147,bool))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3787,f737])).
fof(f3787,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (p(s(bool,X146)) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,sK7('const.pair.,_2'(s(bool,X146),s(X147,X148)),X147,bool)) | s(bool,'T_0') = s(bool,X146)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f73,f3611])).
fof(f3868,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (p(s(bool,X116)) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,sK7('const.pair.,_2'(s(bool,X116),s(X117,X118)),X117,bool))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3782,f737])).
fof(f3782,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (p(s(bool,X116)) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,sK7('const.pair.,_2'(s(bool,X116),s(X117,X118)),X117,bool)) | s(bool,'T_0') = s(bool,X116)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f72,f3611])).
fof(f3865,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,X74)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,sK7('const.pair.,_2'(s(bool,X74),s(X75,X76)),X75,bool)) | s(bool,X77) != s(bool,X79)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3775,f737])).
fof(f3775,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,X74)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,sK7('const.pair.,_2'(s(bool,X74),s(X75,X76)),X75,bool)) | s(bool,X77) != s(bool,X79) | s(bool,'T_0') = s(bool,X74)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f70,f3611])).
fof(f3862,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,X32)) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X32),s(X33,X34)),X33,bool))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f3768,f737])).
fof(f3768,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,X32)) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,sK7('const.pair.,_2'(s(bool,X32),s(X33,X34)),X33,bool)) | s(bool,'T_0') = s(bool,X32)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f68,f3611])).
fof(f3719,plain,(
  $false | $spl40),
  inference(equality_resolution,[],[f251])).
fof(f3646,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,X46) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,sK7('const.pair.,_2'(s(bool,X46),s(X49,X50)),X49,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f86])).
fof(f3645,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X40) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK7('const.pair.,_2'(s(bool,X40),s(X43,X44)),X43,bool)) | s(bool,X41) != s(bool,X45)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f74])).
fof(f3643,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X30) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK7('const.pair.,_2'(s(bool,X30),s(X34,X35)),X34,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f73])).
fof(f3641,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X22,X23)),X22,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f72])).
fof(f3640,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X15,X16)),X15,bool)) | s(bool,X13) != s(bool,X17)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f70])).
fof(f3638,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X4,X5)),X4,bool))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f3609,f68])).
fof(f3613,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,sK7('const.pair.,_2'(s(bool,X8),s(X9,X10)),X9,bool))) | s(bool,'T_0') = s(bool,X8)) ) | $spl144),
  inference(cnf_transformation,[],[f3613_D])).
fof(f3613_D,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,sK7('const.pair.,_2'(s(bool,X8),s(X9,X10)),X9,bool))) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f163,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) | $spl20),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3628,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,'F_0') != s(bool,X244) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X241),s(X242,X243)),X242,bool))) | s(bool,'T_0') = s(bool,X241)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3564,f225])).
fof(f3564,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,'F_0') != s(bool,X244) | p(s(bool,X244)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X241),s(X242,X243)),X242,bool))) | s(bool,'T_0') = s(bool,X241)) ) | $spl6),
  inference(superposition,[],[f102,f3129])).
fof(f3627,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,sK7('const.pair.,_2'(s(bool,X222),s(X223,X224)),X223,bool)) | s(bool,'T_0') = s(bool,X222)) ) | $spl6),
  inference(subsumption_resolution,[],[f3560,f81])).
fof(f3560,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,'F_0')) | p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,sK7('const.pair.,_2'(s(bool,X222),s(X223,X224)),X223,bool)) | s(bool,'T_0') = s(bool,X222)) ) | $spl6),
  inference(superposition,[],[f86,f3129])).
fof(f3626,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,'F_0') != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X212),s(X213,X214)),X213,bool))) | s(bool,'T_0') = s(bool,X212)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3558,f225])).
fof(f3558,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,'F_0') != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X212),s(X213,X214)),X213,bool))) | p(s(bool,X215)) | s(bool,'T_0') = s(bool,X212)) ) | $spl6),
  inference(superposition,[],[f86,f3129])).
fof(f3625,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,sK7('const.pair.,_2'(s(bool,X188),s(X189,X190)),X189,bool)) | s(bool,X191) != s(bool,X193) | s(bool,'T_0') = s(bool,X188)) ) | $spl6),
  inference(subsumption_resolution,[],[f3554,f81])).
fof(f3554,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (p(s(bool,'F_0')) | ~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,sK7('const.pair.,_2'(s(bool,X188),s(X189,X190)),X189,bool)) | s(bool,X191) != s(bool,X193) | s(bool,'T_0') = s(bool,X188)) ) | $spl6),
  inference(superposition,[],[f74,f3129])).
fof(f3624,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,X185) | ~p(s(bool,X186)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X182),s(X183,X184)),X183,bool))) | s(bool,X186) != s(bool,X187) | s(bool,'T_0') = s(bool,X182)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3553,f225])).
fof(f3553,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,X185) | p(s(bool,X185)) | ~p(s(bool,X186)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X182),s(X183,X184)),X183,bool))) | s(bool,X186) != s(bool,X187) | s(bool,'T_0') = s(bool,X182)) ) | $spl6),
  inference(superposition,[],[f74,f3129])).
fof(f3623,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,X143)) | ~p(s(bool,X144)) | ~p(s(bool,X145)) | s(bool,X144) != s(bool,sK7('const.pair.,_2'(s(bool,X140),s(X141,X142)),X141,bool)) | s(bool,'T_0') = s(bool,X140)) ) | $spl6),
  inference(subsumption_resolution,[],[f3546,f81])).
fof(f3546,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~p(s(bool,X144)) | ~p(s(bool,X145)) | s(bool,X144) != s(bool,sK7('const.pair.,_2'(s(bool,X140),s(X141,X142)),X141,bool)) | s(bool,'T_0') = s(bool,X140)) ) | $spl6),
  inference(superposition,[],[f73,f3129])).
fof(f3622,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,'F_0') != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X128),s(X129,X130)),X129,bool))) | ~p(s(bool,X133)) | s(bool,'T_0') = s(bool,X128)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3544,f225])).
fof(f3544,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,'F_0') != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X128),s(X129,X130)),X129,bool))) | ~p(s(bool,X133)) | p(s(bool,X131)) | s(bool,'T_0') = s(bool,X128)) ) | $spl6),
  inference(superposition,[],[f73,f3129])).
fof(f3621,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,sK7('const.pair.,_2'(s(bool,X110),s(X111,X112)),X111,bool)) | s(bool,'T_0') = s(bool,X110)) ) | $spl6),
  inference(subsumption_resolution,[],[f3541,f81])).
fof(f3541,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,'F_0')) | p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,sK7('const.pair.,_2'(s(bool,X110),s(X111,X112)),X111,bool)) | s(bool,'T_0') = s(bool,X110)) ) | $spl6),
  inference(superposition,[],[f72,f3129])).
fof(f3619,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X86),s(X87,X88)),X87,bool))) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X86)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3537,f225])).
fof(f3537,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X86),s(X87,X88)),X87,bool))) | p(s(bool,X89)) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X86)) ) | $spl6),
  inference(superposition,[],[f72,f3129])).
fof(f3618,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,sK7('const.pair.,_2'(s(bool,X68),s(X69,X70)),X69,bool)) | s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X68)) ) | $spl6),
  inference(subsumption_resolution,[],[f3534,f81])).
fof(f3534,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (p(s(bool,'F_0')) | ~p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,sK7('const.pair.,_2'(s(bool,X68),s(X69,X70)),X69,bool)) | s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X68)) ) | $spl6),
  inference(superposition,[],[f70,f3129])).
fof(f3617,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,X60)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X57,X58)),X57,bool))) | s(bool,X60) != s(bool,X61) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3532,f225])).
fof(f3532,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X56),s(X57,X58)),X57,bool))) | s(bool,X60) != s(bool,X61) | s(bool,'T_0') = s(bool,X56)) ) | $spl6),
  inference(superposition,[],[f70,f3129])).
fof(f3615,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,sK7('const.pair.,_2'(s(bool,X26),s(X27,X28)),X27,bool)) | s(bool,'T_0') = s(bool,X26)) ) | $spl6),
  inference(subsumption_resolution,[],[f3527,f81])).
fof(f3527,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,sK7('const.pair.,_2'(s(bool,X26),s(X27,X28)),X27,bool)) | s(bool,'T_0') = s(bool,X26)) ) | $spl6),
  inference(superposition,[],[f68,f3129])).
fof(f3612,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X8),s(X9,X10)),X9,bool))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X8)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f3524,f225])).
fof(f3524,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,sK7('const.pair.,_2'(s(bool,X8),s(X9,X10)),X9,bool))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f68,f3129])).
fof(f3467,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X54)))) | p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X54)))) )),
  inference(resolution,[],[f87,f86])).
fof(f3466,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X48)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f87,f74])).
fof(f3464,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X36)))) )),
  inference(resolution,[],[f87,f73])).
fof(f3462,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X24)))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X24)))) )),
  inference(resolution,[],[f87,f72])).
fof(f3461,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X18)))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f87,f70])).
fof(f3459,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X4,fun(X4,bool)),X5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X4),X4)),X6)))) )),
  inference(resolution,[],[f87,f68])).
fof(f3301,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X10,X11)),X10,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3287,f737])).
fof(f3287,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,'F_0'),s(X10,X11)),X10,bool)) | s(bool,'T_0') = s(bool,X9)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3258,f85])).
fof(f3268,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X48),s(X49,X50)),X49,bool)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X48) != s(bool,X52)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f86])).
fof(f3267,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X42),s(X43,X44)),X43,bool)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f74])).
fof(f3264,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X24),s(X25,X26)),X25,bool)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X24) != s(bool,X29)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f73])).
fof(f3263,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X18),s(X19,X20)),X19,bool)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f72])).
fof(f3262,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X12),s(X13,X14)),X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f70])).
fof(f3260,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.,_2'(s(bool,X0),s(X1,X2)),X1,bool)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3258,f68])).
fof(f3246,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,sK7(X40,X41,bool)) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),X40)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f86])).
fof(f3245,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,sK7(X35,X36,bool)) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X36),X35))) | s(bool,X38) != s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f74])).
fof(f3242,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,sK7(X20,X21,bool)) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X20)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f73])).
fof(f3241,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,sK7(X15,X16,bool)) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X16),X15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f72])).
fof(f3240,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,sK7(X10,X11,bool)) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),X10))) | s(bool,X13) != s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f70])).
fof(f3238,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK7(X0,X1,bool)) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),X0)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3216,f68])).
fof(f3218,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4)))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | $spl134),
  inference(cnf_transformation,[],[f3218_D])).
fof(f3218_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4)))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f161,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl18),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3233,plain,(
  ( ! [X198,X199,X197] : (s(bool,'T_0') != s(bool,X199) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X197),X198)))) | s(bool,'F_0') = s(bool,sK7(X198,X197,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3170,f737])).
fof(f3170,plain,(
  ( ! [X198,X199,X197] : (s(bool,'T_0') != s(bool,X199) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X197),X198)))) | ~p(s(bool,X199)) | s(bool,'F_0') = s(bool,sK7(X198,X197,bool))) ) | $spl6),
  inference(superposition,[],[f102,f3000])).
fof(f3232,plain,(
  ( ! [X196,X194,X195,X193] : (p(s(bool,X195)) | p(s(bool,X196)) | s(bool,X196) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X193),X194))) | s(bool,'F_0') = s(bool,sK7(X194,X193,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3169,f82])).
fof(f3169,plain,(
  ( ! [X196,X194,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X195)) | p(s(bool,X196)) | s(bool,X196) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X193),X194))) | s(bool,'F_0') = s(bool,sK7(X194,X193,bool))) ) | $spl6),
  inference(superposition,[],[f86,f3000])).
fof(f3231,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X183) | p(s(bool,X184)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X181),X182)))) | s(bool,'F_0') = s(bool,sK7(X182,X181,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3166,f737])).
fof(f3166,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X183) | p(s(bool,X184)) | ~p(s(bool,X183)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X181),X182)))) | s(bool,'F_0') = s(bool,sK7(X182,X181,bool))) ) | $spl6),
  inference(superposition,[],[f86,f3000])).
fof(f3230,plain,(
  ( ! [X171,X169,X172,X170,X168] : (p(s(bool,X170)) | ~p(s(bool,X171)) | s(bool,X170) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X168),X169))) | s(bool,X171) != s(bool,X172) | s(bool,'F_0') = s(bool,sK7(X169,X168,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3163,f82])).
fof(f3163,plain,(
  ( ! [X171,X169,X172,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X170)) | ~p(s(bool,X171)) | s(bool,X170) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X168),X169))) | s(bool,X171) != s(bool,X172) | s(bool,'F_0') = s(bool,sK7(X169,X168,bool))) ) | $spl6),
  inference(superposition,[],[f74,f3000])).
fof(f3229,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,X150) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),X149)))) | ~p(s(bool,X151)) | s(bool,X151) != s(bool,X152) | s(bool,'F_0') = s(bool,sK7(X149,X148,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3159,f737])).
fof(f3159,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,X150) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),X149)))) | ~p(s(bool,X151)) | ~p(s(bool,X150)) | s(bool,X151) != s(bool,X152) | s(bool,'F_0') = s(bool,sK7(X149,X148,bool))) ) | $spl6),
  inference(superposition,[],[f74,f3000])).
fof(f3228,plain,(
  ( ! [X132,X130,X128,X131,X129] : (p(s(bool,X130)) | ~p(s(bool,X131)) | p(s(bool,X132)) | s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X128),X129))) | s(bool,'F_0') = s(bool,sK7(X129,X128,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3155,f82])).
fof(f3155,plain,(
  ( ! [X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X130)) | ~p(s(bool,X131)) | p(s(bool,X132)) | s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X128),X129))) | s(bool,'F_0') = s(bool,sK7(X129,X128,bool))) ) | $spl6),
  inference(superposition,[],[f73,f3000])).
fof(f3227,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,X115) | p(s(bool,X116)) | ~p(s(bool,X117)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X113),X114)))) | s(bool,'F_0') = s(bool,sK7(X114,X113,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3152,f737])).
fof(f3152,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,X115) | p(s(bool,X116)) | ~p(s(bool,X115)) | ~p(s(bool,X117)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X113),X114)))) | s(bool,'F_0') = s(bool,sK7(X114,X113,bool))) ) | $spl6),
  inference(superposition,[],[f73,f3000])).
fof(f3226,plain,(
  ( ! [X103,X107,X105,X106,X104] : (p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X105) != s(bool,X107) | s(bool,X106) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),X104))) | s(bool,'F_0') = s(bool,sK7(X104,X103,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3150,f82])).
fof(f3150,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X105) != s(bool,X107) | s(bool,X106) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),X104))) | s(bool,'F_0') = s(bool,sK7(X104,X103,bool))) ) | $spl6),
  inference(superposition,[],[f72,f3000])).
fof(f3224,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') != s(bool,X80) | p(s(bool,X81)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),X79)))) | s(bool,X81) != s(bool,X82) | s(bool,'F_0') = s(bool,sK7(X79,X78,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3145,f737])).
fof(f3145,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') != s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,X80)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),X79)))) | s(bool,X81) != s(bool,X82) | s(bool,'F_0') = s(bool,sK7(X79,X78,bool))) ) | $spl6),
  inference(superposition,[],[f72,f3000])).
fof(f3223,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X71) != s(bool,X72) | s(bool,'F_0') = s(bool,sK7(X69,X68,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3143,f82])).
fof(f3143,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X71) != s(bool,X72) | s(bool,'F_0') = s(bool,sK7(X69,X68,bool))) ) | $spl6),
  inference(superposition,[],[f70,f3000])).
fof(f3222,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54)))) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,sK7(X54,X53,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3140,f737])).
fof(f3140,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54)))) | ~p(s(bool,X56)) | ~p(s(bool,X55)) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,sK7(X54,X53,bool))) ) | $spl6),
  inference(superposition,[],[f70,f3000])).
fof(f3220,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34))) | s(bool,'F_0') = s(bool,sK7(X34,X33,bool))) ) | $spl6),
  inference(subsumption_resolution,[],[f3136,f82])).
fof(f3136,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34))) | s(bool,'F_0') = s(bool,sK7(X34,X33,bool))) ) | $spl6),
  inference(superposition,[],[f68,f3000])).
fof(f3217,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4)))) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3130,f737])).
fof(f3130,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X5)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4)))) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | $spl6),
  inference(superposition,[],[f68,f3000])).
fof(f2963,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f86])).
fof(f2962,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f248,f74])).
fof(f2959,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f73])).
fof(f2958,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f72])).
fof(f2957,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f248,f70])).
fof(f2955,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f68])).
fof(f2917,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X48,X49),s(bool,X50)),bool,X48)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X50) != s(bool,X52)) ) | $spl110),
  inference(resolution,[],[f2874,f86])).
fof(f2916,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X42,X43),s(bool,X44)),bool,X42)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl110),
  inference(resolution,[],[f2874,f74])).
fof(f2913,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X24,X25),s(bool,X26)),bool,X24)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | $spl110),
  inference(resolution,[],[f2874,f73])).
fof(f2912,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X18,X19),s(bool,X20)),bool,X18)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl110),
  inference(resolution,[],[f2874,f72])).
fof(f2911,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X12,X13),s(bool,X14)),bool,X12)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl110),
  inference(resolution,[],[f2874,f70])).
fof(f2909,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl110),
  inference(resolution,[],[f2874,f68])).
fof(f2898,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10))) ) | $spl130),
  inference(cnf_transformation,[],[f2898_D])).
fof(f2898_D,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2897,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10)) | s(bool,'F_0') = s(bool,X8)) ) | ($spl22 | $spl110)),
  inference(subsumption_resolution,[],[f2884,f222])).
fof(f2884,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,X9)),bool,X10)) | s(bool,'F_0') = s(bool,X8)) ) | $spl110),
  inference(superposition,[],[f2496,f85])).
fof(f2868,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X56)) | s(bool,X57) != s(bool,sK8('const.pair.,_2'(s(X58,X59),s(bool,X56)),bool,X58)) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl110),
  inference(resolution,[],[f2496,f86])).
fof(f2867,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,sK8('const.pair.,_2'(s(X51,X52),s(bool,X49)),bool,X51)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl110),
  inference(resolution,[],[f2496,f74])).
fof(f2865,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X37,X38),s(bool,X35)),bool,X37)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl110),
  inference(resolution,[],[f2496,f73])).
fof(f2863,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X22) != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,X21)),bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl110),
  inference(resolution,[],[f2496,f72])).
fof(f2862,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X15) != s(bool,sK8('const.pair.,_2'(s(X16,X17),s(bool,X14)),bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl110),
  inference(resolution,[],[f2496,f70])).
fof(f2860,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,sK8('const.pair.,_2'(s(X2,X3),s(bool,X0)),bool,X2)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl110),
  inference(resolution,[],[f2496,f68])).
fof(f2694,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) )),
  inference(resolution,[],[f61,f86])).
fof(f2693,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))))) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))))) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f61,f74])).
fof(f2691,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) )),
  inference(resolution,[],[f61,f73])).
fof(f2689,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) )),
  inference(resolution,[],[f61,f72])).
fof(f2688,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f61,f70])).
fof(f2686,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) )),
  inference(resolution,[],[f61,f68])).
fof(f2678,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X65,fun(X65,bool)),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65)),fun('type.mlstring.mlstring',fun('type.list.list'(X65),X65))),X68)))))) )),
  inference(resolution,[],[f61,f86])).
fof(f2677,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61)))))) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X58,fun(X58,bool)),X59),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58)),fun('type.mlstring.mlstring',fun('type.list.list'(X58),X58))),X61))))) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f61,f74])).
fof(f2675,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X44,fun(X44,bool)),X45),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X46))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44)),fun('type.mlstring.mlstring',fun('type.list.list'(X44),X44))),X47)))))) )),
  inference(resolution,[],[f61,f73])).
fof(f2673,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30))),X33)))))) )),
  inference(resolution,[],[f61,f72])).
fof(f2672,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X23,fun(X23,bool)),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23)),fun('type.mlstring.mlstring',fun('type.list.list'(X23),X23))),X26))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f61,f70])).
fof(f2670,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X9,fun(X9,bool)),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9)),fun('type.mlstring.mlstring',fun('type.list.list'(X9),X9))),X12)))))) )),
  inference(resolution,[],[f61,f68])).
fof(f2664,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))))) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X66),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64))),X67)))) )),
  inference(resolution,[],[f61,f86])).
fof(f2663,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X57,fun(X57,bool)),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X57,fun(X57,bool)),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60)))))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X57,fun(X57,bool)),X58),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X59),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57)),fun('type.mlstring.mlstring',fun('type.list.list'(X57),X57))),X60))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f61,f74])).
fof(f2660,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X36,fun(X36,bool)),X37),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X36,fun(X36,bool)),X37),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X42) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X36,fun(X36,bool)),X37),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X38),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36)),fun('type.mlstring.mlstring',fun('type.list.list'(X36),X36))),X39)))) )),
  inference(resolution,[],[f61,f73])).
fof(f2659,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X29,fun(X29,bool)),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X29,fun(X29,bool)),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X29,fun(X29,bool)),X30),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X31),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29)),fun('type.mlstring.mlstring',fun('type.list.list'(X29),X29))),X32)))) )),
  inference(resolution,[],[f61,f72])).
fof(f2658,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X24),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22))),X25))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f61,f70])).
fof(f2656,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X8,fun(X8,bool)),X9),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X10),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8)),fun('type.mlstring.mlstring',fun('type.list.list'(X8),X8))),X11)))) )),
  inference(resolution,[],[f61,f68])).
fof(f2736,plain,(
  ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f2735,f65])).
fof(f2735,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f2653,f64])).
fof(f2653,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f61,f58])).
fof(f2625,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,'F_0')),bool,X1))) ) | $spl116),
  inference(cnf_transformation,[],[f2625_D])).
fof(f2625_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,'F_0')),bool,X1))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2595,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X23,X24),s(bool,'F_0')),bool,X23)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X22))) | s(bool,'F_0') = s(bool,sK8(X22,bool,X21))) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f2085,f1655])).
fof(f2085,plain,(
  ( ! [X243,X244,X245] : (s(bool,X245) != s(bool,sK8('const.pair.,_2'(s(X243,X244),s(bool,'F_0')),bool,X243)) | s(bool,'F_0') = s(bool,X245)) ) | ($spl0 | $spl6 | $spl70)),
  inference(trivial_inequality_removal,[],[f2076])).
fof(f2076,plain,(
  ( ! [X243,X244,X245] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X245) != s(bool,sK8('const.pair.,_2'(s(X243,X244),s(bool,'F_0')),bool,X243)) | s(bool,'F_0') = s(bool,X245)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f1469,f1981])).
fof(f2594,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X19,X20),s(bool,'F_0')),bool,X19)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),'const.pair.,_2'(s(X16,X17),s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f2085,f1526])).
fof(f2624,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,'F_0')),bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6 | $spl70)),
  inference(duplicate_literal_removal,[],[f2589])).
fof(f2589,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,sK8('const.pair.,_2'(s(X1,X2),s(bool,'F_0')),bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6 | $spl70)),
  inference(superposition,[],[f2085,f85])).
fof(f2536,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))) )),
  inference(resolution,[],[f60,f86])).
fof(f2535,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X49,fun(X49,bool)),X50),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f60,f74])).
fof(f2533,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X35,fun(X35,bool)),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))) )),
  inference(resolution,[],[f60,f73])).
fof(f2531,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X21,fun(X21,bool)),X22),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))) )),
  inference(resolution,[],[f60,f72])).
fof(f2530,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X14,fun(X14,bool)),X15),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f60,f70])).
fof(f2528,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) )),
  inference(resolution,[],[f60,f68])).
fof(f2523,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X66) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X61,fun(X61,bool)),X62),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))))) )),
  inference(resolution,[],[f60,f86])).
fof(f2522,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X54,fun(X54,bool)),X55),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f60,f74])).
fof(f2519,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X33,fun(X33,bool)),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))))) )),
  inference(resolution,[],[f60,f73])).
fof(f2518,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X26,fun(X26,bool)),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))))) )),
  inference(resolution,[],[f60,f72])).
fof(f2517,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X19,fun(X19,bool)),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f60,f70])).
fof(f2515,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))))) )),
  inference(resolution,[],[f60,f68])).
fof(f2496,plain,(
  ( ! [X33,X36,X34,X32] : (~p(s(bool,X36)) | p(s(bool,X34)) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X32,X33),s(bool,X34)),bool,X32))) ) | $spl110),
  inference(cnf_transformation,[],[f2496_D])).
fof(f2496_D,plain,(
  ( ! [X33,X36,X34,X32] : (~p(s(bool,X36)) | p(s(bool,X34)) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X32,X33),s(bool,X34)),bool,X32))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2507,plain,(
  ( ! [X231,X229,X230,X228,X232] : (p(s(bool,X230)) | p(s(bool,X231)) | ~p(s(bool,X232)) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X228,X229),s(bool,X230)),bool,X228))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2453,f737])).
fof(f2453,plain,(
  ( ! [X231,X229,X230,X228,X232] : (p(s(bool,X230)) | p(s(bool,X231)) | ~p(s(bool,X232)) | s(bool,X232) != s(bool,sK8('const.pair.,_2'(s(X228,X229),s(bool,X230)),bool,X228)) | s(bool,'T_0') = s(bool,X230)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f86,f2219])).
fof(f2505,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X196)) | ~p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X198) != s(bool,sK8('const.pair.,_2'(s(X194,X195),s(bool,X196)),bool,X194)) | s(bool,X197) != s(bool,X199)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2446,f737])).
fof(f2446,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X196)) | ~p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X198) != s(bool,sK8('const.pair.,_2'(s(X194,X195),s(bool,X196)),bool,X194)) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') = s(bool,X196)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f74,f2219])).
fof(f2503,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (p(s(bool,X148)) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,sK8('const.pair.,_2'(s(X146,X147),s(bool,X148)),bool,X146))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2438,f737])).
fof(f2438,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (p(s(bool,X148)) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,sK8('const.pair.,_2'(s(X146,X147),s(bool,X148)),bool,X146)) | s(bool,'T_0') = s(bool,X148)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f73,f2219])).
fof(f2501,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (p(s(bool,X118)) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,sK8('const.pair.,_2'(s(X116,X117),s(bool,X118)),bool,X116))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2433,f737])).
fof(f2433,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (p(s(bool,X118)) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,sK8('const.pair.,_2'(s(X116,X117),s(bool,X118)),bool,X116)) | s(bool,'T_0') = s(bool,X118)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f72,f2219])).
fof(f2498,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,sK8('const.pair.,_2'(s(X74,X75),s(bool,X76)),bool,X74)) | s(bool,X77) != s(bool,X79)) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2426,f737])).
fof(f2426,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,sK8('const.pair.,_2'(s(X74,X75),s(bool,X76)),bool,X74)) | s(bool,X77) != s(bool,X79) | s(bool,'T_0') = s(bool,X76)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f70,f2219])).
fof(f2495,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X32,X33),s(bool,X34)),bool,X32))) ) | ($spl0 | $spl6 | $spl47)),
  inference(subsumption_resolution,[],[f2419,f737])).
fof(f2419,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,X34)) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,sK8('const.pair.,_2'(s(X32,X33),s(bool,X34)),bool,X32)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl47)),
  inference(superposition,[],[f68,f2219])).
fof(f2341,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,X46) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,sK8('const.pair.,_2'(s(X49,X50),s(bool,X46)),bool,X49))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f86])).
fof(f2340,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X40) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK8('const.pair.,_2'(s(X43,X44),s(bool,X40)),bool,X43)) | s(bool,X41) != s(bool,X45)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f74])).
fof(f2338,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X30) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK8('const.pair.,_2'(s(X34,X35),s(bool,X30)),bool,X34))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f73])).
fof(f2336,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,sK8('const.pair.,_2'(s(X22,X23),s(bool,X18)),bool,X22))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f72])).
fof(f2335,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,sK8('const.pair.,_2'(s(X15,X16),s(bool,X12)),bool,X15)) | s(bool,X13) != s(bool,X17)) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f70])).
fof(f2333,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,sK8('const.pair.,_2'(s(X4,X5),s(bool,X0)),bool,X4))) ) | ($spl6 | $spl28)),
  inference(resolution,[],[f2217,f68])).
fof(f2284,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X58),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56))),X59)))) )),
  inference(resolution,[],[f59,f86])).
fof(f2283,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X49,fun(X49,bool)),X50),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X51),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49)),fun('type.mlstring.mlstring',fun('type.list.list'(X49),X49))),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f59,f74])).
fof(f2281,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X35,fun(X35,bool)),X36),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X37),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35)),fun('type.mlstring.mlstring',fun('type.list.list'(X35),X35))),X38)))) )),
  inference(resolution,[],[f59,f73])).
fof(f2279,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X21,fun(X21,bool)),X22),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X23),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21)),fun('type.mlstring.mlstring',fun('type.list.list'(X21),X21))),X24)))) )),
  inference(resolution,[],[f59,f72])).
fof(f2278,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X14,fun(X14,bool)),X15),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X16),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14)),fun('type.mlstring.mlstring',fun('type.list.list'(X14),X14))),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f59,f70])).
fof(f2276,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X2),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0))),X3)))) )),
  inference(resolution,[],[f59,f68])).
fof(f2271,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X66) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X61,fun(X61,bool)),X62),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X63))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61)),fun('type.mlstring.mlstring',fun('type.list.list'(X61),X61))),X64)))))) )),
  inference(resolution,[],[f59,f86])).
fof(f2270,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X54,fun(X54,bool)),X55),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X56))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54)),fun('type.mlstring.mlstring',fun('type.list.list'(X54),X54))),X57))))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f59,f74])).
fof(f2267,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X33,fun(X33,bool)),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X35))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33)),fun('type.mlstring.mlstring',fun('type.list.list'(X33),X33))),X36)))))) )),
  inference(resolution,[],[f59,f73])).
fof(f2266,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X26,fun(X26,bool)),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X28))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26)),fun('type.mlstring.mlstring',fun('type.list.list'(X26),X26))),X29)))))) )),
  inference(resolution,[],[f59,f72])).
fof(f2265,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X19,fun(X19,bool)),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X21))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19)),fun('type.mlstring.mlstring',fun('type.list.list'(X19),X19))),X22))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f59,f70])).
fof(f2263,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X5,fun(X5,bool)),X6),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),X7))),s(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5)),fun('type.mlstring.mlstring',fun('type.list.list'(X5),X5))),X8)))))) )),
  inference(resolution,[],[f59,f68])).
fof(f2259,plain,(
  $false | ($spl0 | $spl6 | $spl20)),
  inference(trivial_inequality_removal,[],[f2253])).
fof(f2253,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl6 | $spl20)),
  inference(superposition,[],[f163,f1981])).
fof(f2257,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f163])).
fof(f2221,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,sK8('const.pair.,_2'(s(X8,X9),s(bool,X10)),bool,X8))) | s(bool,'T_0') = s(bool,X10)) ) | $spl102),
  inference(cnf_transformation,[],[f2221_D])).
fof(f2221_D,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,sK8('const.pair.,_2'(s(X8,X9),s(bool,X10)),bool,X8))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2236,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,'F_0') != s(bool,X244) | ~p(s(bool,sK8('const.pair.,_2'(s(X241,X242),s(bool,X243)),bool,X241))) | s(bool,'T_0') = s(bool,X243)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2190,f225])).
fof(f2190,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,'F_0') != s(bool,X244) | p(s(bool,X244)) | ~p(s(bool,sK8('const.pair.,_2'(s(X241,X242),s(bool,X243)),bool,X241))) | s(bool,'T_0') = s(bool,X243)) ) | $spl6),
  inference(superposition,[],[f102,f1764])).
fof(f2235,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,sK8('const.pair.,_2'(s(X222,X223),s(bool,X224)),bool,X222)) | s(bool,'T_0') = s(bool,X224)) ) | $spl6),
  inference(subsumption_resolution,[],[f2186,f81])).
fof(f2186,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,'F_0')) | p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,sK8('const.pair.,_2'(s(X222,X223),s(bool,X224)),bool,X222)) | s(bool,'T_0') = s(bool,X224)) ) | $spl6),
  inference(superposition,[],[f86,f1764])).
fof(f2234,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,'F_0') != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,sK8('const.pair.,_2'(s(X212,X213),s(bool,X214)),bool,X212))) | s(bool,'T_0') = s(bool,X214)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2184,f225])).
fof(f2184,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,'F_0') != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,sK8('const.pair.,_2'(s(X212,X213),s(bool,X214)),bool,X212))) | p(s(bool,X215)) | s(bool,'T_0') = s(bool,X214)) ) | $spl6),
  inference(superposition,[],[f86,f1764])).
fof(f2233,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,sK8('const.pair.,_2'(s(X188,X189),s(bool,X190)),bool,X188)) | s(bool,X191) != s(bool,X193) | s(bool,'T_0') = s(bool,X190)) ) | $spl6),
  inference(subsumption_resolution,[],[f2180,f81])).
fof(f2180,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (p(s(bool,'F_0')) | ~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,sK8('const.pair.,_2'(s(X188,X189),s(bool,X190)),bool,X188)) | s(bool,X191) != s(bool,X193) | s(bool,'T_0') = s(bool,X190)) ) | $spl6),
  inference(superposition,[],[f74,f1764])).
fof(f2232,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,X185) | ~p(s(bool,X186)) | ~p(s(bool,sK8('const.pair.,_2'(s(X182,X183),s(bool,X184)),bool,X182))) | s(bool,X186) != s(bool,X187) | s(bool,'T_0') = s(bool,X184)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2179,f225])).
fof(f2179,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,'F_0') != s(bool,X185) | p(s(bool,X185)) | ~p(s(bool,X186)) | ~p(s(bool,sK8('const.pair.,_2'(s(X182,X183),s(bool,X184)),bool,X182))) | s(bool,X186) != s(bool,X187) | s(bool,'T_0') = s(bool,X184)) ) | $spl6),
  inference(superposition,[],[f74,f1764])).
fof(f2231,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,X143)) | ~p(s(bool,X144)) | ~p(s(bool,X145)) | s(bool,X144) != s(bool,sK8('const.pair.,_2'(s(X140,X141),s(bool,X142)),bool,X140)) | s(bool,'T_0') = s(bool,X142)) ) | $spl6),
  inference(subsumption_resolution,[],[f2172,f81])).
fof(f2172,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~p(s(bool,X144)) | ~p(s(bool,X145)) | s(bool,X144) != s(bool,sK8('const.pair.,_2'(s(X140,X141),s(bool,X142)),bool,X140)) | s(bool,'T_0') = s(bool,X142)) ) | $spl6),
  inference(superposition,[],[f73,f1764])).
fof(f2230,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,'F_0') != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,sK8('const.pair.,_2'(s(X128,X129),s(bool,X130)),bool,X128))) | ~p(s(bool,X133)) | s(bool,'T_0') = s(bool,X130)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2170,f225])).
fof(f2170,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,'F_0') != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,sK8('const.pair.,_2'(s(X128,X129),s(bool,X130)),bool,X128))) | ~p(s(bool,X133)) | p(s(bool,X131)) | s(bool,'T_0') = s(bool,X130)) ) | $spl6),
  inference(superposition,[],[f73,f1764])).
fof(f2229,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,sK8('const.pair.,_2'(s(X110,X111),s(bool,X112)),bool,X110)) | s(bool,'T_0') = s(bool,X112)) ) | $spl6),
  inference(subsumption_resolution,[],[f2167,f81])).
fof(f2167,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,'F_0')) | p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,sK8('const.pair.,_2'(s(X110,X111),s(bool,X112)),bool,X110)) | s(bool,'T_0') = s(bool,X112)) ) | $spl6),
  inference(superposition,[],[f72,f1764])).
fof(f2227,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,sK8('const.pair.,_2'(s(X86,X87),s(bool,X88)),bool,X86))) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X88)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2163,f225])).
fof(f2163,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,sK8('const.pair.,_2'(s(X86,X87),s(bool,X88)),bool,X86))) | p(s(bool,X89)) | s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,X88)) ) | $spl6),
  inference(superposition,[],[f72,f1764])).
fof(f2226,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,sK8('const.pair.,_2'(s(X68,X69),s(bool,X70)),bool,X68)) | s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X70)) ) | $spl6),
  inference(subsumption_resolution,[],[f2160,f81])).
fof(f2160,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (p(s(bool,'F_0')) | ~p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,sK8('const.pair.,_2'(s(X68,X69),s(bool,X70)),bool,X68)) | s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X70)) ) | $spl6),
  inference(superposition,[],[f70,f1764])).
fof(f2225,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,X60)) | ~p(s(bool,sK8('const.pair.,_2'(s(X56,X57),s(bool,X58)),bool,X56))) | s(bool,X60) != s(bool,X61) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2158,f225])).
fof(f2158,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,sK8('const.pair.,_2'(s(X56,X57),s(bool,X58)),bool,X56))) | s(bool,X60) != s(bool,X61) | s(bool,'T_0') = s(bool,X58)) ) | $spl6),
  inference(superposition,[],[f70,f1764])).
fof(f2223,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,sK8('const.pair.,_2'(s(X26,X27),s(bool,X28)),bool,X26)) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(subsumption_resolution,[],[f2153,f81])).
fof(f2153,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,sK8('const.pair.,_2'(s(X26,X27),s(bool,X28)),bool,X26)) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f68,f1764])).
fof(f2220,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,sK8('const.pair.,_2'(s(X8,X9),s(bool,X10)),bool,X8))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X10)) ) | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2150,f225])).
fof(f2150,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'F_0') != s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,sK8('const.pair.,_2'(s(X8,X9),s(bool,X10)),bool,X8))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f68,f1764])).
fof(f2000,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,'F_0')),bool,X10))) ) | $spl100),
  inference(cnf_transformation,[],[f2000_D])).
fof(f2000_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,'F_0')),bool,X10))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1999,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,'F_0')),bool,X10))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1990,f737])).
fof(f1990,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X10,X11),s(bool,'F_0')),bool,X10)) | s(bool,'T_0') = s(bool,X9)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1873,f85])).
fof(f1977,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X48,X49),s(bool,X50)),bool,X48)) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X50) != s(bool,X52)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f86])).
fof(f1976,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X42,X43),s(bool,X44)),bool,X42)) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f74])).
fof(f1973,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X24,X25),s(bool,X26)),bool,X24)) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f73])).
fof(f1972,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X18,X19),s(bool,X20)),bool,X18)) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f72])).
fof(f1971,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X12,X13),s(bool,X14)),bool,X12)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f70])).
fof(f1969,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK8('const.pair.,_2'(s(X0,X1),s(bool,X2)),bool,X0)) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1873,f68])).
fof(f1910,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X72,fun(X72,bool)),X73),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X76)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X72,fun(X72,bool)),X73),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X77),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X72),X72)),X76)))) | p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X74),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X77)))) )),
  inference(resolution,[],[f62,f86])).
fof(f1909,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X63,fun(X63,bool)),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X65),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X67)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X63,fun(X63,bool)),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X68),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X63),X63)),X67)))) | ~p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X68))) | s(bool,X69) != s(bool,X71)) )),
  inference(resolution,[],[f62,f74])).
fof(f1907,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X45,fun(X45,bool)),X46),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X47),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X49)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X45,fun(X45,bool)),X46),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X50),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X45),X45)),X49)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X50)))) )),
  inference(resolution,[],[f62,f73])).
fof(f1905,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X27,fun(X27,bool)),X28),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X29),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X31)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X27,fun(X27,bool)),X28),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X27),X27)),X31)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X29),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X32)))) )),
  inference(resolution,[],[f62,f72])).
fof(f1904,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X18,fun(X18,bool)),X19),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X20),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X21),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X22)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X18,fun(X18,bool)),X19),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X23),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X21),s(fun('type.mlstring.mlstring',fun('type.list.list'(X18),X18)),X22)))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X20),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X23))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f62,f70])).
fof(f1902,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X4)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X5)))) )),
  inference(resolution,[],[f62,f68])).
fof(f1897,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84)))) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X86) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) )),
  inference(resolution,[],[f62,f86])).
fof(f1896,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X77,X75,X73] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75)))) | ~p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74))) | s(bool,X76) != s(bool,X78)) )),
  inference(resolution,[],[f62,f74])).
fof(f1894,plain,(
  ( ! [X59,X57,X54,X52,X60,X58,X56,X55,X53] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X54),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X55),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X56)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X54),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X57)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X52,fun(X52,bool)),X53),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X57),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X55),s(fun('type.mlstring.mlstring',fun('type.list.list'(X52),X52)),X56)))) )),
  inference(resolution,[],[f62,f73])).
fof(f1892,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) )),
  inference(resolution,[],[f62,f72])).
fof(f1891,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f62,f70])).
fof(f1889,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12)))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) )),
  inference(resolution,[],[f62,f68])).
fof(f1884,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X81)))) | p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X86) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X79,fun(X79,bool)),X80),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X84),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X82),s(fun('type.mlstring.mlstring',fun('type.list.list'(X79),X79)),X83)))) )),
  inference(resolution,[],[f62,f86])).
fof(f1883,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X77,X75,X73] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X72)))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X76) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X75),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X74))) | s(bool,X77) != s(bool,X78)) )),
  inference(resolution,[],[f62,f74])).
fof(f1880,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X43,fun(X43,bool)),X44),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X45),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X46),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X47)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X48),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X45)))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X43,fun(X43,bool)),X44),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X46),s(fun('type.mlstring.mlstring',fun('type.list.list'(X43),X43)),X47)))) )),
  inference(resolution,[],[f62,f73])).
fof(f1879,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X36)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X34,fun(X34,bool)),X35),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X39),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X37),s(fun('type.mlstring.mlstring',fun('type.list.list'(X34),X34)),X38)))) )),
  inference(resolution,[],[f62,f72])).
fof(f1878,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X27)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X25,fun(X25,bool)),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X30),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X25),X25)),X29))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f62,f70])).
fof(f1876,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X9)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(X7,fun(X7,bool)),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),X12),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X7),X7)),X11)))) )),
  inference(resolution,[],[f62,f68])).
fof(f1862,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,sK8(X40,bool,X41)) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),X40)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f86])).
fof(f1861,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,sK8(X35,bool,X36)) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),X35))) | s(bool,X38) != s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f74])).
fof(f1858,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,sK8(X20,bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X20)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f73])).
fof(f1857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,sK8(X15,bool,X16)) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),X15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f72])).
fof(f1856,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,sK8(X10,bool,X11)) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X11,bool),X10))) | s(bool,X13) != s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f70])).
fof(f1854,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK8(X0,bool,X1)) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),X0)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f1833,f68])).
fof(f1835,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),X4)))) | s(bool,'F_0') = s(bool,sK8(X4,bool,X3))) ) | $spl88),
  inference(cnf_transformation,[],[f1835_D])).
fof(f1835_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),X4)))) | s(bool,'F_0') = s(bool,sK8(X4,bool,X3))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1850,plain,(
  ( ! [X198,X199,X197] : (s(bool,'T_0') != s(bool,X199) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X197,bool),X198)))) | s(bool,'F_0') = s(bool,sK8(X198,bool,X197))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1805,f737])).
fof(f1805,plain,(
  ( ! [X198,X199,X197] : (s(bool,'T_0') != s(bool,X199) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X197,bool),X198)))) | ~p(s(bool,X199)) | s(bool,'F_0') = s(bool,sK8(X198,bool,X197))) ) | $spl6),
  inference(superposition,[],[f102,f1655])).
fof(f1849,plain,(
  ( ! [X196,X194,X195,X193] : (p(s(bool,X195)) | p(s(bool,X196)) | s(bool,X196) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X193,bool),X194))) | s(bool,'F_0') = s(bool,sK8(X194,bool,X193))) ) | $spl6),
  inference(subsumption_resolution,[],[f1804,f82])).
fof(f1804,plain,(
  ( ! [X196,X194,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X195)) | p(s(bool,X196)) | s(bool,X196) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X193,bool),X194))) | s(bool,'F_0') = s(bool,sK8(X194,bool,X193))) ) | $spl6),
  inference(superposition,[],[f86,f1655])).
fof(f1848,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X183) | p(s(bool,X184)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X181,bool),X182)))) | s(bool,'F_0') = s(bool,sK8(X182,bool,X181))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1801,f737])).
fof(f1801,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X183) | p(s(bool,X184)) | ~p(s(bool,X183)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X181,bool),X182)))) | s(bool,'F_0') = s(bool,sK8(X182,bool,X181))) ) | $spl6),
  inference(superposition,[],[f86,f1655])).
fof(f1847,plain,(
  ( ! [X171,X169,X172,X170,X168] : (p(s(bool,X170)) | ~p(s(bool,X171)) | s(bool,X170) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X168,bool),X169))) | s(bool,X171) != s(bool,X172) | s(bool,'F_0') = s(bool,sK8(X169,bool,X168))) ) | $spl6),
  inference(subsumption_resolution,[],[f1798,f82])).
fof(f1798,plain,(
  ( ! [X171,X169,X172,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X170)) | ~p(s(bool,X171)) | s(bool,X170) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X168,bool),X169))) | s(bool,X171) != s(bool,X172) | s(bool,'F_0') = s(bool,sK8(X169,bool,X168))) ) | $spl6),
  inference(superposition,[],[f74,f1655])).
fof(f1846,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,X150) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X148,bool),X149)))) | ~p(s(bool,X151)) | s(bool,X151) != s(bool,X152) | s(bool,'F_0') = s(bool,sK8(X149,bool,X148))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1794,f737])).
fof(f1794,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,X150) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X148,bool),X149)))) | ~p(s(bool,X151)) | ~p(s(bool,X150)) | s(bool,X151) != s(bool,X152) | s(bool,'F_0') = s(bool,sK8(X149,bool,X148))) ) | $spl6),
  inference(superposition,[],[f74,f1655])).
fof(f1845,plain,(
  ( ! [X132,X130,X128,X131,X129] : (p(s(bool,X130)) | ~p(s(bool,X131)) | p(s(bool,X132)) | s(bool,X132) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X128,bool),X129))) | s(bool,'F_0') = s(bool,sK8(X129,bool,X128))) ) | $spl6),
  inference(subsumption_resolution,[],[f1790,f82])).
fof(f1790,plain,(
  ( ! [X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X130)) | ~p(s(bool,X131)) | p(s(bool,X132)) | s(bool,X132) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X128,bool),X129))) | s(bool,'F_0') = s(bool,sK8(X129,bool,X128))) ) | $spl6),
  inference(superposition,[],[f73,f1655])).
fof(f1844,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,X115) | p(s(bool,X116)) | ~p(s(bool,X117)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),X114)))) | s(bool,'F_0') = s(bool,sK8(X114,bool,X113))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1787,f737])).
fof(f1787,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,X115) | p(s(bool,X116)) | ~p(s(bool,X115)) | ~p(s(bool,X117)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),X114)))) | s(bool,'F_0') = s(bool,sK8(X114,bool,X113))) ) | $spl6),
  inference(superposition,[],[f73,f1655])).
fof(f1843,plain,(
  ( ! [X103,X107,X105,X106,X104] : (p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X105) != s(bool,X107) | s(bool,X106) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X103,bool),X104))) | s(bool,'F_0') = s(bool,sK8(X104,bool,X103))) ) | $spl6),
  inference(subsumption_resolution,[],[f1785,f82])).
fof(f1785,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X105) != s(bool,X107) | s(bool,X106) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X103,bool),X104))) | s(bool,'F_0') = s(bool,sK8(X104,bool,X103))) ) | $spl6),
  inference(superposition,[],[f72,f1655])).
fof(f1841,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') != s(bool,X80) | p(s(bool,X81)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),X79)))) | s(bool,X81) != s(bool,X82) | s(bool,'F_0') = s(bool,sK8(X79,bool,X78))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1780,f737])).
fof(f1780,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') != s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,X80)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),X79)))) | s(bool,X81) != s(bool,X82) | s(bool,'F_0') = s(bool,sK8(X79,bool,X78))) ) | $spl6),
  inference(superposition,[],[f72,f1655])).
fof(f1840,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),X69))) | s(bool,X71) != s(bool,X72) | s(bool,'F_0') = s(bool,sK8(X69,bool,X68))) ) | $spl6),
  inference(subsumption_resolution,[],[f1778,f82])).
fof(f1778,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),X69))) | s(bool,X71) != s(bool,X72) | s(bool,'F_0') = s(bool,sK8(X69,bool,X68))) ) | $spl6),
  inference(superposition,[],[f70,f1655])).
fof(f1839,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X53,bool),X54)))) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,sK8(X54,bool,X53))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1775,f737])).
fof(f1775,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X55) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X53,bool),X54)))) | ~p(s(bool,X56)) | ~p(s(bool,X55)) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,sK8(X54,bool,X53))) ) | $spl6),
  inference(superposition,[],[f70,f1655])).
fof(f1837,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),X34))) | s(bool,'F_0') = s(bool,sK8(X34,bool,X33))) ) | $spl6),
  inference(subsumption_resolution,[],[f1771,f82])).
fof(f1771,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),X34))) | s(bool,'F_0') = s(bool,sK8(X34,bool,X33))) ) | $spl6),
  inference(superposition,[],[f68,f1655])).
fof(f1834,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),X4)))) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,sK8(X4,bool,X3))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1765,f737])).
fof(f1765,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X5)) | p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),X4)))) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,sK8(X4,bool,X3))) ) | $spl6),
  inference(superposition,[],[f68,f1655])).
fof(f1585,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),X67)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X64,fun(X64,bool)),X65),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X68),s(fun('type.mlstring.mlstring',fun('type.list.list'(X64),X64)),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X68)))) )),
  inference(resolution,[],[f63,f86])).
fof(f1584,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X58),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),X59)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X56,fun(X56,bool)),X57),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X60),s(fun('type.mlstring.mlstring',fun('type.list.list'(X56),X56)),X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X58),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f63,f74])).
fof(f1582,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X40,fun(X40,bool)),X41),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s(fun('type.mlstring.mlstring',fun('type.list.list'(X40),X40)),X43)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X40,fun(X40,bool)),X41),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X44),s(fun('type.mlstring.mlstring',fun('type.list.list'(X40),X40)),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X44)))) )),
  inference(resolution,[],[f63,f73])).
fof(f1580,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X24,fun(X24,bool)),X25),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X24),X24)),X27)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X24,fun(X24,bool)),X25),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X28),s(fun('type.mlstring.mlstring',fun('type.list.list'(X24),X24)),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X28)))) )),
  inference(resolution,[],[f63,f72])).
fof(f1579,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X18),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X19)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X16,fun(X16,bool)),X17),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X20),s(fun('type.mlstring.mlstring',fun('type.list.list'(X16),X16)),X19)))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X18),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X20))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f63,f70])).
fof(f1577,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X0,fun(X0,bool)),X1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4),s(fun('type.mlstring.mlstring',fun('type.list.list'(X0),X0)),X3)))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X2),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X4)))) )),
  inference(resolution,[],[f63,f68])).
fof(f1572,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) )),
  inference(resolution,[],[f63,f86])).
fof(f1571,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66)))) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f63,f74])).
fof(f1569,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X48),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X49)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X48),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X50)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X46,fun(X46,bool)),X47),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X50),s(fun('type.mlstring.mlstring',fun('type.list.list'(X46),X46)),X49)))) )),
  inference(resolution,[],[f63,f73])).
fof(f1567,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) )),
  inference(resolution,[],[f63,f72])).
fof(f1566,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26)))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f63,f70])).
fof(f1564,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) )),
  inference(resolution,[],[f63,f68])).
fof(f1559,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X72)))) | p(s(bool,X75)) | p(s(bool,X76)) | s(bool,X76) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X70,fun(X70,bool)),X71),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X74),s(fun('type.mlstring.mlstring',fun('type.list.list'(X70),X70)),X73)))) )),
  inference(resolution,[],[f63,f86])).
fof(f1558,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X64)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X62,fun(X62,bool)),X63),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X66),s(fun('type.mlstring.mlstring',fun('type.list.list'(X62),X62)),X65))) | s(bool,X68) != s(bool,X69)) )),
  inference(resolution,[],[f63,f74])).
fof(f1555,plain,(
  ( ! [X39,X45,X43,X41,X38,X44,X42,X40] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X38,fun(X38,bool)),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X40),s(fun('type.mlstring.mlstring',fun('type.list.list'(X38),X38)),X41)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X40)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X38,fun(X38,bool)),X39),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X42),s(fun('type.mlstring.mlstring',fun('type.list.list'(X38),X38)),X41)))) )),
  inference(resolution,[],[f63,f73])).
fof(f1554,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X32)))) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X30,fun(X30,bool)),X31),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X34),s(fun('type.mlstring.mlstring',fun('type.list.list'(X30),X30)),X33)))) )),
  inference(resolution,[],[f63,f72])).
fof(f1553,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X24)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X22,fun(X22,bool)),X23),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X26),s(fun('type.mlstring.mlstring',fun('type.list.list'(X22),X22)),X25))) | s(bool,X28) != s(bool,X29)) )),
  inference(resolution,[],[f63,f70])).
fof(f1551,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X8)))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(X6,fun(X6,bool)),X7),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),X10),s(fun('type.mlstring.mlstring',fun('type.list.list'(X6),X6)),X9)))) )),
  inference(resolution,[],[f63,f68])).
fof(f412,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) | $spl64),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1492,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f86])).
fof(f1491,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f58,f74])).
fof(f1489,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f73])).
fof(f1487,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f72])).
fof(f1486,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f58,f70])).
fof(f1484,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f68])).
fof(f1481,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl80),
  inference(cnf_transformation,[],[f1481_D])).
fof(f1481_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1480,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | ($spl0 | $spl70)),
  inference(subsumption_resolution,[],[f1473,f737])).
fof(f1473,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl70),
  inference(superposition,[],[f799,f85])).
fof(f1465,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl70),
  inference(resolution,[],[f799,f86])).
fof(f1464,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl70),
  inference(resolution,[],[f799,f74])).
fof(f1461,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl70),
  inference(resolution,[],[f799,f73])).
fof(f1460,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl70),
  inference(resolution,[],[f799,f72])).
fof(f1459,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl70),
  inference(resolution,[],[f799,f70])).
fof(f1457,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl70),
  inference(resolution,[],[f799,f68])).
fof(f1440,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | ($spl6 | $spl52)),
  inference(subsumption_resolution,[],[f1423,f958])).
fof(f958,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl52)),
  inference(subsumption_resolution,[],[f356,f494])).
fof(f1423,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl52)),
  inference(superposition,[],[f958,f85])).
fof(f1439,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl52)),
  inference(duplicate_literal_removal,[],[f1422])).
fof(f1422,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl52)),
  inference(superposition,[],[f958,f85])).
fof(f1420,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl48)),
  inference(duplicate_literal_removal,[],[f1404])).
fof(f1404,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl48)),
  inference(superposition,[],[f910,f85])).
fof(f910,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl48)),
  inference(subsumption_resolution,[],[f325,f488])).
fof(f1421,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl48)),
  inference(subsumption_resolution,[],[f1403,f910])).
fof(f1403,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl48)),
  inference(superposition,[],[f910,f85])).
fof(f250,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) | $spl38),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1396,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f86])).
fof(f1395,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f57,f74])).
fof(f1392,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f73])).
fof(f1391,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f72])).
fof(f1390,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f57,f70])).
fof(f1388,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f68])).
fof(f1303,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f1170,f81])).
fof(f1170,plain,(
  p(s(bool,'F_0')) | $spl54),
  inference(backward_demodulation,[],[f358,f57])).
fof(f1163,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f86])).
fof(f1162,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f57,f74])).
fof(f1159,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f73])).
fof(f1158,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f72])).
fof(f1157,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f57,f70])).
fof(f1155,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f68])).
fof(f1101,plain,(
  $false | ($spl22 | $spl30)),
  inference(subsumption_resolution,[],[f967,f81])).
fof(f967,plain,(
  p(s(bool,'F_0')) | ($spl22 | $spl30)),
  inference(backward_demodulation,[],[f959,f248])).
fof(f959,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) | ($spl22 | $spl30)),
  inference(subsumption_resolution,[],[f228,f222])).
fof(f1100,plain,(
  $false | ($spl22 | $spl30)),
  inference(subsumption_resolution,[],[f965,f81])).
fof(f965,plain,(
  p(s(bool,'F_0')) | ($spl22 | $spl30)),
  inference(backward_demodulation,[],[f959,f57])).
fof(f904,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) != s(bool,X8)) ) | $spl76),
  inference(cnf_transformation,[],[f904_D])).
fof(f904_D,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) != s(bool,X8)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f903,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl28 | $spl68)),
  inference(subsumption_resolution,[],[f895,f225])).
fof(f895,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X6)) ) | $spl68),
  inference(superposition,[],[f764,f85])).
fof(f295,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl47),
  inference(cnf_transformation,[],[f295_D])).
fof(f295_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f455,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X1))) ) | $spl28),
  inference(subsumption_resolution,[],[f448,f225])).
fof(f448,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl28),
  inference(superposition,[],[f225,f85])).
fof(f490,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f477,f229])).
fof(f477,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f854,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(forward_demodulation,[],[f853,f106])).
fof(f853,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f843])).
fof(f843,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f766,f85])).
fof(f869,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(forward_demodulation,[],[f868,f121])).
fof(f121,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))),
  inference(subsumption_resolution,[],[f118,f81])).
fof(f118,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))),
  inference(superposition,[],[f56,f85])).
fof(f868,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f859])).
fof(f859,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f767,f85])).
fof(f767,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f761,f82])).
fof(f761,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f488,f121])).
fof(f957,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4)) ) | ($spl6 | $spl34)),
  inference(subsumption_resolution,[],[f477,f230])).
fof(f229,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X3))) ) | $spl32),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X3))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f956,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X2)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f482,f222])).
fof(f482,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f619,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X22) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23)) )),
  inference(forward_demodulation,[],[f618,f106])).
fof(f618,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(subsumption_resolution,[],[f569,f82])).
fof(f569,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(superposition,[],[f68,f106])).
fof(f623,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') != s(bool,X42) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,X44)) )),
  inference(forward_demodulation,[],[f622,f106])).
fof(f622,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X43) != s(bool,X44)) )),
  inference(subsumption_resolution,[],[f576,f82])).
fof(f576,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X43) != s(bool,X44)) )),
  inference(superposition,[],[f70,f106])).
fof(f628,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X65)) )),
  inference(forward_demodulation,[],[f627,f106])).
fof(f627,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(subsumption_resolution,[],[f583,f82])).
fof(f583,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(superposition,[],[f72,f106])).
fof(f632,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,X80) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80))) )),
  inference(forward_demodulation,[],[f631,f106])).
fof(f631,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(subsumption_resolution,[],[f588,f82])).
fof(f588,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'T_0')) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(superposition,[],[f73,f106])).
fof(f636,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') != s(bool,X102) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X103) != s(bool,X104)) )),
  inference(forward_demodulation,[],[f635,f106])).
fof(f635,plain,(
  ( ! [X103,X102,X104] : (p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X103) != s(bool,X104)) )),
  inference(subsumption_resolution,[],[f596,f82])).
fof(f596,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X103) != s(bool,X104)) )),
  inference(superposition,[],[f74,f106])).
fof(f640,plain,(
  ( ! [X114,X113] : (s(bool,'T_0') != s(bool,X114) | p(s(bool,X113)) | p(s(bool,X114))) )),
  inference(forward_demodulation,[],[f639,f106])).
fof(f639,plain,(
  ( ! [X114,X113] : (p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X114) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(subsumption_resolution,[],[f602,f82])).
fof(f602,plain,(
  ( ! [X114,X113] : (~p(s(bool,'T_0')) | p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X114) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(superposition,[],[f86,f106])).
fof(f712,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24)) )),
  inference(forward_demodulation,[],[f711,f121])).
fof(f711,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(subsumption_resolution,[],[f662,f82])).
fof(f662,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(superposition,[],[f68,f121])).
fof(f716,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,X45)) )),
  inference(forward_demodulation,[],[f715,f121])).
fof(f715,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X44) != s(bool,X45)) )),
  inference(subsumption_resolution,[],[f669,f82])).
fof(f669,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X44) != s(bool,X45)) )),
  inference(superposition,[],[f70,f121])).
fof(f721,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X66)) )),
  inference(forward_demodulation,[],[f720,f121])).
fof(f720,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X66) | s(bool,X65) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(subsumption_resolution,[],[f676,f82])).
fof(f676,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X66) | s(bool,X65) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(superposition,[],[f72,f121])).
fof(f725,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81))) )),
  inference(forward_demodulation,[],[f724,f121])).
fof(f724,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | s(bool,X81) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(subsumption_resolution,[],[f681,f82])).
fof(f681,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | s(bool,X81) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(superposition,[],[f73,f121])).
fof(f729,plain,(
  ( ! [X103,X105,X104] : (s(bool,'T_0') != s(bool,X103) | p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X104) != s(bool,X105)) )),
  inference(forward_demodulation,[],[f728,f121])).
fof(f728,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X103) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X104) != s(bool,X105)) )),
  inference(subsumption_resolution,[],[f689,f82])).
fof(f689,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'T_0')) | p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X103) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X104) != s(bool,X105)) )),
  inference(superposition,[],[f74,f121])).
fof(f733,plain,(
  ( ! [X114,X115] : (s(bool,'T_0') != s(bool,X115) | p(s(bool,X114)) | p(s(bool,X115))) )),
  inference(forward_demodulation,[],[f732,f121])).
fof(f732,plain,(
  ( ! [X114,X115] : (p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X115) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(subsumption_resolution,[],[f695,f82])).
fof(f695,plain,(
  ( ! [X114,X115] : (~p(s(bool,'T_0')) | p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X115) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(superposition,[],[f86,f121])).
fof(f911,plain,(
  $false | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f309,f295])).
fof(f309,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl26),
  inference(resolution,[],[f224,f81])).
fof(f909,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X6)) ) | ($spl22 | $spl68)),
  inference(subsumption_resolution,[],[f895,f222])).
fof(f901,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl74),
  inference(cnf_transformation,[],[f901_D])).
fof(f901_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f881,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'T_0') = s(bool,X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl68),
  inference(resolution,[],[f764,f86])).
fof(f880,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl68),
  inference(resolution,[],[f764,f74])).
fof(f878,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl68),
  inference(resolution,[],[f764,f73])).
fof(f876,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl68),
  inference(resolution,[],[f764,f72])).
fof(f875,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl68),
  inference(resolution,[],[f764,f70])).
fof(f873,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl68),
  inference(resolution,[],[f764,f68])).
fof(f835,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f86])).
fof(f834,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f248,f74])).
fof(f831,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f73])).
fof(f830,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f72])).
fof(f829,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f248,f70])).
fof(f827,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) )),
  inference(resolution,[],[f248,f68])).
fof(f778,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl6),
  inference(resolution,[],[f494,f86])).
fof(f777,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f494,f74])).
fof(f774,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f494,f73])).
fof(f773,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f494,f72])).
fof(f772,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f494,f70])).
fof(f770,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f494,f68])).
fof(f764,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl68),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f746,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'T_0') = s(bool,X40) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl6),
  inference(resolution,[],[f488,f86])).
fof(f745,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f488,f74])).
fof(f743,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl6),
  inference(resolution,[],[f488,f73])).
fof(f741,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f488,f72])).
fof(f740,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f488,f70])).
fof(f738,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f488,f68])).
fof(f650,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f120])).
fof(f557,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f101])).
fof(f388,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl56),
  inference(cnf_transformation,[],[f388_D])).
fof(f388_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f530,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f230,f86])).
fof(f529,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f230,f74])).
fof(f527,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl34),
  inference(resolution,[],[f230,f73])).
fof(f525,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f230,f72])).
fof(f524,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f230,f70])).
fof(f522,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f230,f68])).
fof(f506,plain,(
  $false | ($spl22 | $spl36)),
  inference(resolution,[],[f474,f81])).
fof(f474,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) | ($spl22 | $spl36)),
  inference(subsumption_resolution,[],[f231,f222])).
fof(f511,plain,(
  $false | ($spl22 | $spl36)),
  inference(resolution,[],[f474,f58])).
fof(f230,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl34),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f493,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f483,f222])).
fof(f483,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f424,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl22),
  inference(resolution,[],[f222,f86])).
fof(f423,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f222,f74])).
fof(f102,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl6),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X6)) | ~p(s(bool,X7))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f279,plain,(
  ( ! [X4,X7] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) | $spl44),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  ( ! [X4,X7] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f422,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl22),
  inference(resolution,[],[f222,f74])).
fof(f420,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f222,f73])).
fof(f419,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f222,f72])).
fof(f418,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f222,f70])).
fof(f417,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl22),
  inference(resolution,[],[f222,f70])).
fof(f416,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f222,f68])).
fof(f402,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f86])).
fof(f401,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f58,f74])).
fof(f399,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f73])).
fof(f413,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) | $spl66),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f398,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f58,f72])).
fof(f397,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f72])).
fof(f396,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f58,f70])).
fof(f395,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f58,f68])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f58,f68])).
fof(f407,plain,(
  ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f406,f65])).
fof(f406,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(forward_demodulation,[],[f393,f64])).
fof(f393,plain,(
  ~p(s(bool,'const.holSemantics.is_term_assignment|48|_4'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),'const.pair.SND_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.SND_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))))) | ~p(s(bool,'const.holSemantics.is_type_assignment|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'const.pair.FST_1'(s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4))))),s(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),'const.pair.FST_1'(s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))))))),
  inference(resolution,[],[f58,f61])).
fof(f371,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl34),
  inference(resolution,[],[f230,f86])).
fof(f370,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f230,f86])).
fof(f369,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f230,f74])).
fof(f389,plain,(
  ( ! [X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl34),
  inference(subsumption_resolution,[],[f368,f86])).
fof(f368,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl34),
  inference(resolution,[],[f230,f73])).
fof(f367,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl34),
  inference(resolution,[],[f230,f73])).
fof(f326,plain,(
  ( ! [X4,X7] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) | $spl50),
  inference(cnf_transformation,[],[f326_D])).
fof(f326_D,plain,(
  ( ! [X4,X7] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f366,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl34),
  inference(resolution,[],[f230,f72])).
fof(f365,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f230,f72])).
fof(f364,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f230,f70])).
fof(f363,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl34),
  inference(resolution,[],[f230,f68])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f230,f68])).
fof(f339,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') = s(bool,X35) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl30),
  inference(resolution,[],[f228,f86])).
fof(f356,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl52),
  inference(cnf_transformation,[],[f356_D])).
fof(f356_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f338,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl30),
  inference(resolution,[],[f228,f86])).
fof(f337,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl30),
  inference(resolution,[],[f228,f74])).
fof(f357,plain,(
  ( ! [X24,X27,X25] : (s(bool,'F_0') = s(bool,X24) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl30),
  inference(subsumption_resolution,[],[f336,f86])).
fof(f336,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl30),
  inference(resolution,[],[f228,f73])).
fof(f335,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl30),
  inference(resolution,[],[f228,f73])).
fof(f334,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl30),
  inference(resolution,[],[f228,f72])).
fof(f333,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl30),
  inference(resolution,[],[f228,f72])).
fof(f332,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl30),
  inference(resolution,[],[f228,f70])).
fof(f331,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl30),
  inference(resolution,[],[f228,f68])).
fof(f330,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl30),
  inference(resolution,[],[f228,f68])).
fof(f325,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl48),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f308,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl26),
  inference(resolution,[],[f224,f86])).
fof(f307,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl26),
  inference(resolution,[],[f224,f74])).
fof(f306,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl26),
  inference(resolution,[],[f224,f74])).
fof(f304,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl26),
  inference(resolution,[],[f224,f73])).
fof(f303,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl26),
  inference(resolution,[],[f224,f72])).
fof(f302,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl26),
  inference(resolution,[],[f224,f70])).
fof(f301,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl26),
  inference(resolution,[],[f224,f70])).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl26),
  inference(resolution,[],[f224,f68])).
fof(f299,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f290,f82])).
fof(f290,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(equality_resolution,[],[f223])).
fof(f293,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X2))) ) | $spl24),
  inference(subsumption_resolution,[],[f286,f223])).
fof(f286,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(superposition,[],[f223,f85])).
fof(f261,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl22),
  inference(resolution,[],[f222,f86])).
fof(f260,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f222,f74])).
fof(f259,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl22),
  inference(resolution,[],[f222,f74])).
fof(f257,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f222,f73])).
fof(f256,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f222,f72])).
fof(f255,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f222,f70])).
fof(f254,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl22),
  inference(resolution,[],[f222,f70])).
fof(f253,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f222,f68])).
fof(f245,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f86])).
fof(f244,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f57,f74])).
fof(f251,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) | $spl40),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f243,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f74])).
fof(f241,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f73])).
fof(f240,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f72])).
fof(f239,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f57,f70])).
fof(f238,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f70])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.holSemantics.is_interpretation|48|_3'(s(fun(sK0,fun(sK0,bool)),sK1),s('type.pair.prod'('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),'type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type')),'const.pair.,_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))),s('type.pair.prod'(fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0)),fun('type.mlstring.mlstring',fun('type.list.list'(sK0),sK0))),sK2)))) )),
  inference(resolution,[],[f57,f68])).
fof(f231,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5))) ) | $spl36),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f216,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f228,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f215,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f224,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f210,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f222,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f223,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3))) ) | $spl24),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f209,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,'F_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f102,f85])).
fof(f194,plain,(
  $false | $spl10),
  inference(resolution,[],[f104,f82])).
fof(f195,plain,(
  $false | $spl10),
  inference(resolution,[],[f104,f55])).
fof(f196,plain,(
  $false | $spl10),
  inference(resolution,[],[f104,f56])).
fof(f180,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) ) | $spl8),
  inference(subsumption_resolution,[],[f168,f103])).
fof(f168,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f103,f85])).
fof(f179,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl8),
  inference(subsumption_resolution,[],[f167,f103])).
fof(f167,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f103,f85])).
fof(f162,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X4))) ) | $spl2),
  inference(subsumption_resolution,[],[f148,f100])).
fof(f148,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(superposition,[],[f100,f85])).
fof(f160,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2))) ) | $spl2),
  inference(subsumption_resolution,[],[f147,f100])).
fof(f147,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) ) | $spl2),
  inference(superposition,[],[f100,f85])).
fof(f131,plain,(
  $false | $spl12),
  inference(resolution,[],[f105,f81])).
fof(f119,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) ) | $spl14),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f116,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f86])).
fof(f115,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f56,f74])).
fof(f120,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) ) | $spl16),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f114,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f74])).
fof(f112,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f73])).
fof(f111,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f72])).
fof(f110,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f56,f70])).
fof(f109,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f70])).
fof(f108,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK4),s('type.finite_map.fmap'('type.mlstring.mlstring','type.holSyntax.type'),sK6)))) )),
  inference(resolution,[],[f56,f68])).
fof(f99,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) ) | $spl0),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f96,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f86])).
fof(f95,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f55,f74])).
fof(f101,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) ) | $spl4),
  inference(cnf_transformation,[],[f101_D])).
fof(f101_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f94,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X22) | s(bool,X23) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f74])).
fof(f92,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f73])).
fof(f91,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f72])).
fof(f90,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f55,f70])).
fof(f89,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X7) | s(bool,X8) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f70])).
fof(f88,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK3),s('type.finite_map.fmap'('type.mlstring.mlstring','type.num.num'),sK5)))) )),
  inference(resolution,[],[f55,f68])).
% SZS output end Proof for thm.holSemanticsExtra
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 178
% Passive clauses: 1490
% Generated clauses: 8914
% Final active clauses: 86
% Final passive clauses: 1129
% Input formulas: 15
% Initial clauses: 31
% 
% Selected by SInE selection: 13
% SInE iterations: 1
% 
% Duplicate literals: 83
% Trivial inequalities: 172
% Fw subsumption resolutions: 1047
% Fw demodulations: 840
% Bw demodulations: 65
% 
% Simple tautologies: 137
% Equational tautologies: 7
% Forward subsumptions: 4292
% 
% Binary resolution: 1742
% Forward superposition: 2391
% Backward superposition: 2246
% Self superposition: 30
% Equality resolution: 40
% 
% Split clauses: 750
% Split components: 1634
% SAT solver clauses: 775
% SAT solver unit clauses: 13
% SAT solver binary clauses: 236
% 
% Sat splits: 750
% Sat splitting refutations: 24
% 
% Memory used [KB]: 6012
% Time elapsed: 0.479 s
% ------------------------------
% ------------------------------
% Success in time 12.213 s
