Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.patSem
% (120940)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 163
% Passive clauses: 3400
% Generated clauses: 22930
% Final active clauses: 133
% Final passive clauses: 3053
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 792
% Trivial inequalities: 113
% Fw subsumption resolutions: 1626
% Fw demodulations: 3403
% Bw demodulations: 26
% 
% Simple tautologies: 139
% Equational tautologies: 47
% Forward subsumptions: 12979
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 470
% Factoring: 2
% Forward superposition: 8634
% Backward superposition: 7460
% Self superposition: 214
% Equality factoring: 94
% Equality resolution: 10
% 
% Split clauses: 211
% Split components: 422
% SAT solver clauses: 610
% SAT solver unit clauses: 40
% SAT solver binary clauses: 129
% 
% Sat splits: 211
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 610
% TWLsolver calls for satisfiability: 149
% 
% Memory used [KB]: 21108
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.patSem
% (120963)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 1263
% Generated clauses: 18220
% Final active clauses: 97
% Final passive clauses: 806
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 634
% Trivial inequalities: 105
% Fw subsumption resolutions: 1050
% Fw demodulations: 2165
% Bw demodulations: 348
% Condensations: 1624
% 
% Simple tautologies: 118
% Equational tautologies: 44
% Forward subsumptions: 9086
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 454
% Factoring: 1
% Forward superposition: 5485
% Backward superposition: 4676
% Self superposition: 184
% Equality factoring: 98
% Equality resolution: 8
% 
% Split clauses: 110
% Split components: 220
% SAT solver clauses: 1896
% SAT solver unit clauses: 131
% SAT solver binary clauses: 42
% 
% Sat splits: 110
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1896
% TWLsolver calls for satisfiability: 275
% 
% Memory used [KB]: 16886
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.patSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.patSem
% SZS output start Proof for thm.patSem
fof(f37533,plain,(
  $false),
  inference(subsumption_resolution,[],[f37532,f216])).
fof(f216,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.patSem.state_fupdfupds_comp|split|3',aHL_TRUTH)).
fof(f37532,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f37530,f3579])).
fof(f3579,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f3578])).
fof(f3578,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f3514,f300])).
fof(f300,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f219])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.patSem.state_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f218,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.patSem.state_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f3514,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f3003,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f219,f219])).
fof(f3003,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2914,f262])).
fof(f262,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X14] : ! [X17,X22] : s(X14,X22) = s(X14,'const.pair.SND_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X17),s(X14,X22)))))),
  file('thm.patSem.state_fupdfupds_comp|split|3','thm.pair.SND')).
fof(f2914,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f262,f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.patSem.state_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f37530,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.patSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.patSem.state'(sK1)),fun(fun(sK0,'type.patSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))))),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f37463,f13349])).
fof(f13349,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.patSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.patSem.state'(sK1)),fun(fun(sK0,'type.patSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))))),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4))),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f215,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f215,plain,(
  s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))) != s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4))),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))),s(fun(sK0,'type.patSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))))),s(fun(sK0,'type.patSem.state'(sK1)),sK2))) != s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK4))),s(fun(sK0,'type.patSem.state'(sK1)),'const.combin.o_2'(s(fun('type.patSem.state'(sK1),'type.patSem.state'(sK1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),sK3))),s(fun(sK0,'type.patSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))))),s(fun(X0,'type.patSem.state'(X1)),X2))) != s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4))),s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))),s(fun(X0,'type.patSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))))),s(fun(X0,'type.patSem.state'(X1)),X2))) = s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4))),s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))),s(fun(X0,'type.patSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))))),s(fun(X0,'type.patSem.state'(X1)),X2))) = s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X4))),s(fun(X0,'type.patSem.state'(X1)),'const.combin.o_2'(s(fun('type.patSem.state'(X1),'type.patSem.state'(X1)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X3))),s(fun(X0,'type.patSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9))))),s(fun(X6,'type.patSem.state'(X7)),X8))) = s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10))),s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9))),s(fun(X6,'type.patSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9))))),s(fun(X6,'type.patSem.state'(X7)),X8))) = s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10))),s(fun(X6,'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9))),s(fun(X6,'type.patSem.state'(X7)),X8)))))),
  file('thm.patSem.state_fupdfupds_comp|split|3','thm.patSem.state_fupdfupds_comp|split|3')).
fof(f37463,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2))))),s(fun(X3,'type.patSem.state'(X0)),X4))) = s(fun(X3,'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1))),s(fun(X3,'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2))),s(fun(X3,'type.patSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f277,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1))))) = s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2))),s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1))))) = s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2))),s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1))))) = s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.combin.o_2'(s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X2))),s(fun('type.patSem.state'(X0),'type.patSem.state'(X0)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9))))) = s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.combin.o_2'(s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X10))),s(fun('type.patSem.state'(X7),'type.patSem.state'(X7)),'const.patSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v'))),X9)))))),
  file('thm.patSem.state_fupdfupds_comp|split|3','thm.patSem.state_fupdfupds_comp|split|2')).
fof(f277,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14,X6,X21,X15] : ! [X10,X9,X8] : s(fun(X21,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X21,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X21,X15),X8))))) = s(fun(X21,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X21,X15),X8)))),
  file('thm.patSem.state_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.patSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 242
% Passive clauses: 3768
% Generated clauses: 40837
% Final active clauses: 112
% Final passive clauses: 1386
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 1056
% Trivial inequalities: 63
% Fw subsumption resolutions: 1128
% Fw demodulations: 6597
% Bw demodulations: 2083
% Condensations: 1532
% Global subsumptions: 353
% 
% Simple tautologies: 121
% Equational tautologies: 180
% Forward subsumptions: 25494
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 69
% 
% Unit resulting resolution: 1144
% Factoring: 1
% Forward superposition: 12437
% Backward superposition: 10243
% Self superposition: 280
% Equality factoring: 94
% Equality resolution: 2
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 4642
% SAT solver unit clauses: 1069
% SAT solver binary clauses: 2066
% 
% Sat splits: 39
% Sat splitting refutations: 441
% 
% Memory used [KB]: 44519
% Time elapsed: 1.950 s
% ------------------------------
% ------------------------------
% Success in time 3.28 s
