Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exh_to_patProof
% (114605)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 925
% Passive clauses: 4918
% Generated clauses: 17951
% Final active clauses: 911
% Final passive clauses: 3666
% Input formulas: 133
% Initial clauses: 441
% 
% Duplicate literals: 224
% Trivial inequalities: 48
% Fw subsumption resolutions: 1318
% Fw demodulations: 1406
% Bw demodulations: 6
% 
% Simple tautologies: 252
% Equational tautologies: 51
% Forward subsumptions: 9554
% 
% Binary resolution: 5905
% Factoring: 1156
% Forward superposition: 2887
% Backward superposition: 4434
% Self superposition: 21
% Equality factoring: 2
% Equality resolution: 34
% 
% Split clauses: 261
% Split components: 524
% SAT solver clauses: 499
% SAT solver unit clauses: 8
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 261
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 498
% TWLsolver calls for satisfiability: 438
% 
% Memory used [KB]: 10106
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exh_to_patProof
% (114623)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 492
% Passive clauses: 3345
% Generated clauses: 14154
% Final active clauses: 372
% Final passive clauses: 2247
% Input formulas: 133
% Initial clauses: 441
% 
% Duplicate literals: 84
% Trivial inequalities: 14
% Fw subsumption resolutions: 495
% Fw demodulations: 865
% Bw demodulations: 620
% Condensations: 68
% Global subsumptions: 22
% 
% Simple tautologies: 260
% Equational tautologies: 12
% Forward subsumptions: 5038
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 1523
% Factoring: 8
% Forward superposition: 3370
% Backward superposition: 2892
% Self superposition: 58
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 93
% Split components: 186
% SAT solver clauses: 4118
% SAT solver unit clauses: 469
% SAT solver binary clauses: 211
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 11
% 
% Sat splits: 93
% Sat splitting refutations: 59
% 
% TWLsolver clauses: 4029
% TWLsolver calls for satisfiability: 5628
% 
% Memory used [KB]: 18677
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exh_to_patProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exh_to_patProof
% SZS output start Proof for thm.exh_to_patProof
fof(f12090,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f633,f1101,f1099,f632,f1100,f1098,f2308,f5323,f5333,f5332,f6017,f6020,f6022,f6024,f6026,f6030,f6032,f6036,f6038,f6050,f6052,f6064,f6066,f6077,f6088,f6107,f6133,f6135,f6144,f6149,f6159,f6168,f6179,f6182,f6187,f6189,f6204,f6209,f6262,f6264,f6273,f6278,f6288,f6297,f6308,f6311,f6316,f6318,f6333,f6338,f6378,f6380,f6382,f6392,f6394,f6396,f6412,f6414,f6422,f6424,f6440,f6442,f6450,f6452,f6468,f6470,f6478,f6480,f6493,f6495,f7267,f7268,f7269,f7270,f6520,f6522,f6524,f6529,f6531,f6535,f6537,f6539,f6544,f6546,f6556,f6564,f6574,f6576,f6578,f6588,f6590,f6600,f6602,f6612,f6614,f6624,f6626,f6636,f6654,f6672,f6730,f6732,f6736,f6746,f6748,f6750,f6752,f6756,f6758,f6782,f6784,f6786,f6788,f6799,f6801,f6803,f6805,f6807,f6809,f6822,f6832,f6843,f6853,f6855,f6874,f6876,f6886,f6888,f6907,f6909,f6919,f6921,f6941,f6943,f6953,f6955,f6970,f6988,f7019,f7021,f7033,f7035,f7047,f7049,f7061,f7081,f7083,f7093,f7095,f7097,f7109,f7111,f7118,f7124,f7130,f7142,f7145,f7147,f7149,f7151,f7153,f7155,f7157,f7159,f7161,f7180,f7182,f7184,f7186,f7189,f7191,f7202,f7204,f7214,f7216,f7221,f7223,f7227,f7229,f7233,f7235,f7238,f7240,f7242,f7244,f7246,f7248,f7250,f7283,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7300,f7301,f7302,f7303,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7311,f7312,f7313,f7314,f7315,f7316,f7317,f7318,f7319,f7320,f7321,f7322,f7323,f7324,f7325,f7326,f7327,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7336,f7337,f7338,f7339,f7340,f7341,f7342,f7343,f7344,f7345,f7346,f7347,f7348,f7349,f7350,f7351,f7352,f7353,f7354,f7355,f7356,f7357,f7358,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7379,f7380,f7381,f7382,f7383,f7384,f7385,f7386,f7387,f7388,f7389,f7390,f7391,f7392,f7393,f7394,f7395,f7396,f7397,f7398,f7399,f7400,f7401,f7402,f7403,f7404,f7405,f7406,f7407,f7408,f7409,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7449,f7450,f7451,f7452,f7453,f7454,f7455,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7464,f7465,f7466,f7467,f7468,f7469,f7470,f7471,f7474,f7475,f7476,f7477,f7478,f7479,f7480,f7481,f7482,f7483,f7484,f7485,f7486,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f7504,f7505,f7506,f7507,f7508,f7509,f7510,f7511,f7512,f7513,f7514,f7515,f7516,f7517,f7518,f7519,f7520,f7521,f7522,f7523,f7524,f7525,f7526,f7527,f7528,f7529,f7530,f7531,f7532,f7533,f7534,f7535,f7536,f7537,f7538,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7547,f7548,f7549,f7550,f7551,f7552,f7553,f7554,f7555,f7556,f7557,f7558,f7559,f7560,f7561,f7562,f7563,f7564,f7565,f7566,f7567,f7568,f7569,f7570,f7571,f7572,f7573,f7930,f7932,f8714,f8716,f9326,f9327,f9378,f9379,f9430,f9431,f9484,f9485,f6018,f12089])).
fof(f12089,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12088,f1101])).
fof(f12088,plain,(
  s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) | $spl0),
  inference(forward_demodulation,[],[f12037,f11876])).
fof(f11876,plain,(
  s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(sK24,'const.patSem.Loc_1'(s('type.num.num',sK25)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2416,f2309,f3837,f3803,f8340,f891,f835])).
fof(f835,plain,(
  ( ! [X0,X1] : (~sP11(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(X1,X0)))) | sP9(X0,X1) | sP8(X0,X1) | sP7(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK128(X1,X0)))) | sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK128(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK128(X1,X0))))) | sP7(X0,X1) | sP8(X0,X1) | sP9(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(X1,X0))))) | sP10(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP7(X0,X1) & ~sP8(X0,X1) & ~sP9(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP10(X0,X1)) | sP11(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK128,sK129])],[f534])).
fof(f534,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ? [X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5)))) | sP7(X0,X1) | sP8(X0,X1) | sP9(X0,X1) | ? [X4] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP10(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP7(X0,X1) & ~sP8(X0,X1) & ~sP9(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP10(X0,X1)) | sP11(X0,X1)))),
  inference(rectify,[],[f533])).
fof(f533,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP7(X1,X0) | sP8(X1,X0) | sP9(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP10(X1,X0)) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP7(X1,X0) & ~sP8(X1,X0) & ~sP9(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP10(X1,X0)) | sP11(X1,X0)))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP7(X1,X0) | sP8(X1,X0) | sP9(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP10(X1,X0))) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP7(X1,X0) & ~sP8(X1,X0) & ~sP9(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP10(X1,X0)) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f408])).
fof(f408,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP7(X1,X0) | sP8(X1,X0) | sP9(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP10(X1,X0)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f891,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X56,X57] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X56))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X57)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.v_distinct|split|3')).
fof(f8340,plain,(
  ( ! [X0,X1] : (~sP9('const.patSem.Loc_1'(s('type.num.num',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1055,f849])).
fof(f849,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK133(X1,X0)),s('type.list.list'('type.patLang.exp'),sK135(X1,X0)),s('type.num.num',sK136(X1,X0))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK132(X1,X0)),s('type.list.list'('type.patLang.exp'),sK134(X1,X0)),s('type.num.num',sK136(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK133(X1,X0)),s('type.list.list'('type.patLang.exp'),sK135(X1,X0)),s('type.num.num',sK136(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK134(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK132(X1,X0)))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK135(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK133(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK134(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK132(X1,X0)),s('type.list.list'('type.patSem.v'),sK133(X1,X0)))))))),s('type.list.list'('type.patLang.exp'),sK134(X1,X0)),s('type.list.list'('type.patLang.exp'),sK135(X1,X0))))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP9(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK132,sK133,sK134,sK135,sK136])],[f540])).
fof(f540,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ? [X7,X8,X9,X10,X11] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patLang.exp'),X9),s('type.num.num',X11))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X11))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patSem.v'),X8))))))),s('type.list.list'('type.patLang.exp'),X9),s('type.list.list'('type.patLang.exp'),X10)))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP9(X0,X1)))),
  inference(rectify,[],[f539])).
fof(f539,plain,(
  ! [X1,X0] : ((~sP9(X1,X0) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8)))))) & (! [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | sP9(X1,X0)))),
  inference(nnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X1,X0] : (sP9(X1,X0) <=> ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f1055,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X85,X29,X28,X57] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X57))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X28),s('type.list.list'('type.patLang.exp'),X29),s('type.num.num',X85)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.v_distinct|split|12')).
fof(f3803,plain,(
  ( ! [X0,X1] : (~sP8('const.patSem.Loc_1'(s('type.num.num',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1016,f853])).
fof(f853,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK138(X1,X0)),s('type.patLang.exp',sK140(X1,X0))))) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK137(X1,X0)),s('type.patLang.exp',sK139(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK138(X1,X0)),s('type.patLang.exp',sK140(X1,X0)))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK137(X1,X0)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK138(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK137(X1,X0)),s('type.list.list'('type.patSem.v'),sK138(X1,X0)))))))),s('type.patLang.exp',sK139(X1,X0)))),s('type.patLang.exp',sK140(X1,X0))))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP8(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK137,sK138,sK139,sK140])],[f543])).
fof(f543,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | ? [X6,X7,X8,X9] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X6),s('type.patLang.exp',X8))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X7),s('type.patLang.exp',X9))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7))))))),s('type.patLang.exp',X8))),s('type.patLang.exp',X9)))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP8(X0,X1)))),
  inference(rectify,[],[f542])).
fof(f542,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13)))))) & (! [X10,X11,X12,X13] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | sP8(X1,X0)))),
  inference(nnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X1,X0] : (sP8(X1,X0) <=> ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f1016,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X29,X28,X57] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X57))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X28),s('type.patLang.exp',X29)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.v_distinct|split|10')).
fof(f3837,plain,(
  ( ! [X0,X1] : (~sP7('const.patSem.Loc_1'(s('type.num.num',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1015,f857])).
fof(f857,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK141(X1,X0)),s('type.list.list'('type.patSem.v'),sK143(X1,X0))))) )),
  inference(cnf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0,X1] : ((~sP7(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK141(X1,X0)),s('type.list.list'('type.patSem.v'),sK142(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK141(X1,X0)),s('type.list.list'('type.patSem.v'),sK143(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK142(X1,X0)),s('type.list.list'('type.patSem.v'),sK143(X1,X0))))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP7(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK141,sK142,sK143])],[f546])).
fof(f546,plain,(
  ! [X0,X1] : ((~sP7(X0,X1) | ? [X5,X6,X7] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X6))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7)))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP7(X0,X1)))),
  inference(rectify,[],[f545])).
fof(f545,plain,(
  ! [X1,X0] : ((~sP7(X1,X0) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16)))))) & (! [X14,X15,X16] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | sP7(X1,X0)))),
  inference(nnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X1,X0] : (sP7(X1,X0) <=> ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1015,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X29,X28,X57] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X57))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X28),s('type.list.list'('type.patSem.v'),X29)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.v_distinct|split|7')).
fof(f2309,plain,(
  sP11('const.patSem.Loc_1'(s('type.num.num',sK25)),sK24) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f1116])).
fof(f1116,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP11(X1,X0)) )),
  inference(forward_demodulation,[],[f860,f783])).
fof(f783,plain,(
  ( ! [X0,X1] : (s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X8,X9] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X8),s('type.patSem.v',X9))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X8))),s('type.patSem.v',X9)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','ARITY_0_2_const.exh_to_patProof.v_rel')).
fof(f860,plain,(
  ( ! [X0,X1] : (sP11(X1,X0) | ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f548])).
fof(f548,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | sP11(X1,X0)) & (~sP11(X1,X0) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> sP11(X1,X0))),
  inference(definition_folding,[],[f228,f408,f407,f406,f405,f404])).
fof(f407,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f228,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X28,X29] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X28))),s('type.patSem.v',X29)))) <=> (? [X7] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X18,X19,X20] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20))))) | ? [X21,X22,X23,X24] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X22),s('type.patLang.exp',X24))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.patLang.exp',X23))),s('type.patLang.exp',X24))))) | ? [X21,X22,X25,X26,X13] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X13))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X22),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X26))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.list.list'('type.patLang.exp'),X25),s('type.list.list'('type.patLang.exp'),X26))))) | ? [X13] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13)))) | ? [X19,X20] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20)))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_cases')).
fof(f2416,plain,(
  ( ! [X0,X1] : (~sP10('const.patSem.Loc_1'(s('type.num.num',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f892,f845])).
fof(f845,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK131(X1,X0))))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK130(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK131(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK130(X1,X0)),s('type.list.list'('type.patSem.v'),sK131(X1,X0))))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP10(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK130,sK131])],[f537])).
fof(f537,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ? [X4,X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X4),s('type.list.list'('type.patSem.v'),X5)))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP10(X0,X1)))),
  inference(rectify,[],[f536])).
fof(f536,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))) & (! [X2,X3] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f407])).
fof(f892,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X56,X57] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X57))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X56)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.v_distinct|split|14')).
fof(f12037,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(sK24,'const.patSem.Loc_1'(s('type.num.num',sK25)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2416,f2309,f3837,f3803,f8340,f891,f837])).
fof(f837,plain,(
  ( ! [X0,X1] : (~sP11(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK129(X1,X0)))) | sP9(X0,X1) | sP8(X0,X1) | sP7(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK128(X1,X0)))) | sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f535])).
fof(f6018,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f6018_D])).
fof(f6018_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9485,plain,(
  ( ! [X5] : (p(s(bool,sK102(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f9485_D])).
fof(f9485_D,plain,(
  ( ! [X5] : (p(s(bool,sK102(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f9484,plain,(
  ( ! [X4,X5] : (p(s(bool,sK102(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9472,f741])).
fof(f741,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X104] : ! [X106] : s(X104,X106) = s(X104,happ(s(fun(X104,X104),'_i_0_0'),s(X104,X106)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','_i0_ax')).
fof(f9472,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK102(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f766,f741])).
fof(f766,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK102(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK101(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK102(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK101,sK102])],[f495])).
fof(f495,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f494])).
fof(f494,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f493])).
fof(f493,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X35] : ! [X41,X40] : ((? [X38] : p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38)))) | p(s(bool,X40))) <=> ? [X38] : (p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38)))) | p(s(bool,X40))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f9431,plain,(
  ( ! [X5] : (p(s(bool,sK101(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl20),
  inference(cnf_transformation,[],[f9431_D])).
fof(f9431_D,plain,(
  ( ! [X5] : (p(s(bool,sK101(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9430,plain,(
  ( ! [X4,X5] : (p(s(bool,sK101(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9418,f741])).
fof(f9418,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK101(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f764,f741])).
fof(f764,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK101(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f496])).
fof(f9379,plain,(
  ( ! [X5] : (p(s(bool,sK100('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f9379_D])).
fof(f9379_D,plain,(
  ( ! [X5] : (p(s(bool,sK100('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9378,plain,(
  ( ! [X4,X5] : (p(s(bool,sK100('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9366,f741])).
fof(f9366,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK100('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f763,f741])).
fof(f763,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK100(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK99(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK100(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK99,sK100])],[f491])).
fof(f491,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f490])).
fof(f490,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X35] : ! [X41,X40] : ((p(s(bool,X41)) | ? [X38] : p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X38))))) <=> ? [X38] : (p(s(bool,X41)) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X38))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f9327,plain,(
  ( ! [X5] : (p(s(bool,sK99('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl16),
  inference(cnf_transformation,[],[f9327_D])).
fof(f9327_D,plain,(
  ( ! [X5] : (p(s(bool,sK99('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f9326,plain,(
  ( ! [X4,X5] : (p(s(bool,sK99('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9314,f741])).
fof(f9314,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK99('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f761,f741])).
fof(f761,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK99(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f492])).
fof(f8716,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl14),
  inference(cnf_transformation,[],[f8716_D])).
fof(f8716_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8714,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f8703,f741])).
fof(f8703,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f753,f741])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X35] : ! [X41,X38] : (p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38)))) => p(s(bool,happ(s(fun(X35,bool),X41),s(X35,'const.min.@_1'(s(fun(X35,bool),X41)))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.SELECT_AX')).
fof(f7932,plain,(
  p(s(bool,sK95('_i_0_0',bool))) | $spl12),
  inference(cnf_transformation,[],[f7932_D])).
fof(f7932_D,plain,(
  p(s(bool,sK95('_i_0_0',bool))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f7930,plain,(
  ( ! [X3] : (p(s(bool,sK95('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f7919,f741])).
fof(f7919,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK95('_i_0_0',bool)))))) )),
  inference(superposition,[],[f747,f741])).
fof(f747,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK95(X1,X0)))))) )),
  inference(cnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK94(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK95(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK94,sK95])],[f479])).
fof(f479,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f478])).
fof(f478,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X35] : ! [X41] : (~? [X38] : p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38)))) <=> ! [X38] : ~p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.NOT_EXISTS_THM')).
fof(f7573,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6031,f1102])).
fof(f1102,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f637,f634])).
fof(f634,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_loc|split|1',aHL_TRUTH)).
fof(f637,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f430])).
fof(f430,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f84])).
fof(f84,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f6031,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1161,f6018,f784])).
fof(f784,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X30,X31] : ((p(s(bool,X30)) => p(s(bool,X31))) => ((p(s(bool,X31)) => p(s(bool,X30))) => s(bool,X30) = s(bool,X31)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1161,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1102,f686])).
fof(f686,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X32] : (s(bool,'T_0') = s(bool,X32) <=> p(s(bool,X32)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f7572,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6037,f1102])).
fof(f6037,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1161,f6018,f784])).
fof(f7571,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6039,f1102])).
fof(f6039,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1199,f6018,f877])).
fof(f877,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1)))),
  inference(rectify,[],[f556])).
fof(f556,plain,(
  ! [X1,X0] : ((~sP13(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP13(X1,X0)))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X1,X0] : ((~sP13(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP13(X1,X0)))),
  inference(nnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X1,X0] : (sP13(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f1199,plain,(
  sP13('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1102,f634,f879])).
fof(f879,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f557])).
fof(f7570,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6040,f6018])).
fof(f6040,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4661,f6018,f877])).
fof(f4661,plain,(
  sP13('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f879])).
fof(f4610,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1101,f826])).
fof(f826,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f528])).
fof(f528,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X27,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))),s('type.patSem.v',X27)))) <=> s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_loc|split|0')).
fof(f7569,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6042,f6018])).
fof(f6042,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2324,f6018,f877])).
fof(f2324,plain,(
  sP13('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f882])).
fof(f882,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f557])).
fof(f7568,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6043,f6018])).
fof(f6043,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1230,f6018,f877])).
fof(f1230,plain,(
  ( ! [X0] : (sP13('F_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1102,f1211])).
fof(f1211,plain,(
  ( ! [X0,X1] : (sP13(X1,'const.bool.BOUNDED_1'(s(bool,X0))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1209,f634])).
fof(f1209,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP13(X1,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(superposition,[],[f882,f641])).
fof(f641,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X62] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X62)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.BOUNDED_THM')).
fof(f7567,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6044,f6018])).
fof(f6044,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f6018,f877])).
fof(f1423,plain,(
  ( ! [X0] : (sP13('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f1102,f642,f882])).
fof(f642,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X35] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X35),'_i_0_0'),s(fun(X35,X35),'_i_0_0'))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.ind_type.ISO_REFL')).
fof(f7566,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6045,f1102])).
fof(f6045,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2323,f6018,f877])).
fof(f2323,plain,(
  sP13('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f879])).
fof(f7565,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6047,f1102])).
fof(f6047,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1214,f6018,f877])).
fof(f1214,plain,(
  ( ! [X0] : (sP13('const.bool.BOUNDED_1'(s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1102,f1202])).
fof(f1202,plain,(
  ( ! [X0,X1] : (sP13('const.bool.BOUNDED_1'(s(bool,X0)),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1200,f634])).
fof(f1200,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP13('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f879,f641])).
fof(f7564,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6048,f1102])).
fof(f6048,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1422,f6018,f877])).
fof(f1422,plain,(
  ( ! [X0] : (sP13('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1102,f642,f879])).
fof(f7563,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6055,f1102])).
fof(f6055,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1208,f6018,f877])).
fof(f1208,plain,(
  sP13('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1102,f634,f882])).
fof(f7562,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6056,f1102])).
fof(f6056,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2324,f6018,f877])).
fof(f7561,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6057,f1102])).
fof(f6057,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1230,f6018,f877])).
fof(f7560,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6058,f1102])).
fof(f6058,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1423,f6018,f877])).
fof(f7559,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6059,f6018])).
fof(f6059,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2323,f6018,f877])).
fof(f7558,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6060,f6018])).
fof(f6060,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4662,f6018,f877])).
fof(f4662,plain,(
  sP13('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f882])).
fof(f7557,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6061,f6018])).
fof(f6061,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1214,f6018,f877])).
fof(f7556,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6062,f6018])).
fof(f6062,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1422,f6018,f877])).
fof(f7555,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6106,f6258])).
fof(f6258,plain,(
  ( ! [X0] : (sP17('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1129,f1134,f6018,f946])).
fof(f946,plain,(
  ( ! [X2,X0,X1] : (~sP15(X1,X2,X0) | p(s(bool,X0)) | ~sP16(X0,X2,X1) | sP17(X0,X1,X2) | ~sP14(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f588])).
fof(f588,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | (sP14(X2,X1,X0) & sP15(X1,X2,X0) & sP16(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP14(X2,X1,X0) | ~sP15(X1,X2,X0) | ~sP16(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP17(X0,X1,X2)))),
  inference(rectify,[],[f587])).
fof(f587,plain,(
  ! [X2,X1,X0] : ((~sP17(X2,X1,X0) | (sP14(X0,X1,X2) & sP15(X1,X0,X2) & sP16(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP14(X0,X1,X2) | ~sP15(X1,X0,X2) | ~sP16(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP17(X2,X1,X0)))),
  inference(flattening,[],[f586])).
fof(f586,plain,(
  ! [X2,X1,X0] : ((~sP17(X2,X1,X0) | (sP14(X0,X1,X2) & sP15(X1,X0,X2) & sP16(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP14(X0,X1,X2) | ~sP15(X1,X0,X2) | ~sP16(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP17(X2,X1,X0)))),
  inference(nnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X2,X1,X0] : (sP17(X2,X1,X0) <=> (sP14(X0,X1,X2) & sP15(X1,X0,X2) & sP16(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f1134,plain,(
  ( ! [X0,X1] : (sP15(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f952])).
fof(f952,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f593])).
fof(f593,plain,(
  ! [X1,X0,X2] : ((~sP15(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP15(X1,X0,X2)))),
  inference(flattening,[],[f592])).
fof(f592,plain,(
  ! [X1,X0,X2] : ((~sP15(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP15(X1,X0,X2)))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X1,X0,X2] : (sP15(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f1129,plain,(
  ( ! [X0,X1] : (sP16(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f948])).
fof(f948,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f590])).
fof(f590,plain,(
  ! [X2,X0,X1] : ((~sP16(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP16(X2,X0,X1)))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ! [X2,X0,X1] : ((~sP16(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP16(X2,X0,X1)))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X2,X0,X1] : (sP16(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f1139,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f956])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f596])).
fof(f596,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP14(X0,X1,X2)))),
  inference(flattening,[],[f595])).
fof(f595,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP14(X0,X1,X2)))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1,X2] : (sP14(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f6106,plain,(
  ( ! [X0] : (~sP17('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f634,f6018,f943])).
fof(f943,plain,(
  ( ! [X2,X0,X1] : (~sP17(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f588])).
fof(f7554,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6108,f6395])).
fof(f6395,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f954])).
fof(f954,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f594])).
fof(f6108,plain,(
  ~sP15('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f6018,f945])).
fof(f945,plain,(
  ( ! [X2,X0,X1] : (~sP15(X1,X2,X0) | p(s(bool,X1)) | ~sP16(X0,X2,X1) | sP17(X0,X1,X2) | ~sP14(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f588])).
fof(f1978,plain,(
  ~sP17('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f634,f1161,f961])).
fof(f961,plain,(
  ( ! [X2,X0,X1] : (~sP17(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP17(X2,X1,X0)) & (~sP17(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP17(X2,X1,X0))),
  inference(definition_folding,[],[f286,f417,f416,f415,f414])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X101,X99,X100] : ((p(s(bool,X100)) <=> s(bool,X99) = s(bool,X101)) <=> ((p(s(bool,X100)) | p(s(bool,X99)) | p(s(bool,X101))) & (p(s(bool,X100)) | ~p(s(bool,X101)) | ~p(s(bool,X99))) & (p(s(bool,X99)) | ~p(s(bool,X101)) | ~p(s(bool,X100))) & (p(s(bool,X101)) | ~p(s(bool,X99)) | ~p(s(bool,X100)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f1140,plain,(
  ( ! [X0,X1] : (sP14(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f634,f957])).
fof(f957,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f596])).
fof(f7553,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6109,f6395])).
fof(f6109,plain,(
  ~sP15('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f6018,f945])).
fof(f4663,plain,(
  ~sP17('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f634,f4610,f943])).
fof(f7552,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6113,f6395])).
fof(f6113,plain,(
  ~sP15('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f6018,f945])).
fof(f2338,plain,(
  ~sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1161,f1098,f961])).
fof(f7551,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6115,f6395])).
fof(f6115,plain,(
  ~sP15('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f6018,f945])).
fof(f1130,plain,(
  ( ! [X0,X1] : (sP16(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f949])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f591])).
fof(f7550,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6116,f6395])).
fof(f6116,plain,(
  ~sP15('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f6018,f945])).
fof(f7549,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6117,f6395])).
fof(f6117,plain,(
  ~sP15('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f6018,f945])).
fof(f4675,plain,(
  ( ! [X0,X1] : (sP16(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4610,f949])).
fof(f7548,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6124,f6381])).
fof(f6381,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f950])).
fof(f950,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f591])).
fof(f6124,plain,(
  ~sP16('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1134,f6018,f945])).
fof(f7547,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6125,f6381])).
fof(f6125,plain,(
  ~sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1134,f6018,f945])).
fof(f7546,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6150,f6381])).
fof(f6150,plain,(
  ~sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f2326,f6018,f945])).
fof(f2326,plain,(
  ( ! [X0,X1] : (sP15(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f952])).
fof(f7545,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6172,f6381])).
fof(f6172,plain,(
  ~sP16('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1135,f6018,f945])).
fof(f1135,plain,(
  ( ! [X0,X1] : (sP15(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f953])).
fof(f953,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f594])).
fof(f7544,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6173,f6381])).
fof(f6173,plain,(
  ~sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1135,f6018,f945])).
fof(f7543,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6198,f6381])).
fof(f6198,plain,(
  ~sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4680,f6018,f945])).
fof(f4680,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4610,f953])).
fof(f7542,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6237,f6395])).
fof(f6237,plain,(
  ~sP15('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f6018,f946])).
fof(f7541,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6238,f6395])).
fof(f6238,plain,(
  ~sP15('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f6018,f946])).
fof(f7540,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6242,f6395])).
fof(f6242,plain,(
  ~sP15('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f6018,f946])).
fof(f7539,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6244,f6395])).
fof(f6244,plain,(
  ~sP15('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f6018,f946])).
fof(f7538,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6245,f6395])).
fof(f6245,plain,(
  ~sP15('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f6018,f946])).
fof(f7537,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6246,f6395])).
fof(f6246,plain,(
  ~sP15('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f6018,f946])).
fof(f7536,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6253,f6381])).
fof(f6253,plain,(
  ~sP16('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1134,f6018,f946])).
fof(f7535,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6254,f6381])).
fof(f6254,plain,(
  ~sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1134,f6018,f946])).
fof(f7534,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6279,f6381])).
fof(f6279,plain,(
  ~sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f2326,f6018,f946])).
fof(f7533,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6286,f6536])).
fof(f6536,plain,(
  ( ! [X0] : (~sP17(X0,'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2315,f6018,f962])).
fof(f962,plain,(
  ( ! [X2,X0,X1] : (~sP17(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f597])).
fof(f2315,plain,(
  s(bool,'T_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f787])).
fof(f787,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f364])).
fof(f6286,plain,(
  sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f1129,f2326,f6018,f946])).
fof(f2329,plain,(
  ( ! [X0,X1] : (sP14('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f958])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f596])).
fof(f7532,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6301,f6381])).
fof(f6301,plain,(
  ~sP16('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1135,f6018,f946])).
fof(f7531,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6302,f6381])).
fof(f6302,plain,(
  ~sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1135,f6018,f946])).
fof(f7530,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6327,f6381])).
fof(f6327,plain,(
  ~sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4680,f6018,f946])).
fof(f7529,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6335,f4663])).
fof(f6335,plain,(
  sP17('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1139,f1129,f4680,f6018,f946])).
fof(f7528,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6377,f6381])).
fof(f6377,plain,(
  ( ! [X0] : (~sP16('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f634,f6018,f947])).
fof(f947,plain,(
  ( ! [X2,X0,X1] : (~sP16(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f591])).
fof(f7527,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6391,f6395])).
fof(f6391,plain,(
  ( ! [X0] : (~sP15('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f634,f6018,f951])).
fof(f951,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f594])).
fof(f7526,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6398,f6018])).
fof(f6398,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f1102,f6018,f955])).
fof(f955,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f596])).
fof(f2328,plain,(
  ( ! [X0,X1] : (sP14(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f957])).
fof(f7525,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6399,f6018])).
fof(f6399,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f1102,f6018,f955])).
fof(f1141,plain,(
  ( ! [X0,X1] : (sP14('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f634,f958])).
fof(f7524,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6400,f6018])).
fof(f6400,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f1102,f6018,f955])).
fof(f7523,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6401,f6018])).
fof(f6401,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f955])).
fof(f7522,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6402,f6018])).
fof(f6402,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f2327,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f956])).
fof(f7521,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6404,f6018])).
fof(f6404,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7520,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6405,f6018])).
fof(f6405,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f6018,f6018,f955])).
fof(f7519,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6406,f6018])).
fof(f6406,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7518,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6408,f6018])).
fof(f6408,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f1102,f6018,f955])).
fof(f7517,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6409,f6018])).
fof(f6409,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f1102,f6018,f955])).
fof(f7516,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6410,f6018])).
fof(f6410,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f1102,f6018,f955])).
fof(f7515,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6416,f6018])).
fof(f6416,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f7514,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6417,f6018])).
fof(f6417,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f6018,f6018,f955])).
fof(f7513,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6418,f6018])).
fof(f6418,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7512,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6419,f6018])).
fof(f6419,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f6018,f6018,f955])).
fof(f7511,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6420,f6018])).
fof(f6420,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7510,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6426,f6018])).
fof(f6426,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f1102,f6018,f955])).
fof(f7509,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6427,f6018])).
fof(f6427,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1102,f6018,f955])).
fof(f7508,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6428,f6018])).
fof(f6428,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f1102,f6018,f955])).
fof(f7507,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6430,f6018])).
fof(f6430,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f7506,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6431,f6018])).
fof(f6431,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f6018,f6018,f955])).
fof(f7505,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6432,f6018])).
fof(f6432,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7504,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6433,f6018])).
fof(f6433,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f6018,f6018,f955])).
fof(f7503,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6434,f6018])).
fof(f6434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7502,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6435,f6018])).
fof(f6435,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1102,f6018,f955])).
fof(f7501,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6436,f6018])).
fof(f6436,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f1102,f6018,f955])).
fof(f7500,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6438,f6018])).
fof(f6438,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f1102,f6018,f955])).
fof(f7499,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6443,f6018])).
fof(f6443,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f955])).
fof(f7498,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6444,f6018])).
fof(f6444,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f7497,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6445,f6018])).
fof(f6445,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f6018,f6018,f955])).
fof(f7496,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6446,f6018])).
fof(f6446,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7495,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6448,f6018])).
fof(f6448,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6453,f6018])).
fof(f6453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1102,f6018,f955])).
fof(f7493,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6454,f6018])).
fof(f6454,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f1102,f6018,f955])).
fof(f7492,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6456,f6018])).
fof(f6456,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f1102,f6018,f955])).
fof(f7491,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6457,f6018])).
fof(f6457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f955])).
fof(f7490,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6458,f6018])).
fof(f6458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f7489,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6460,f6018])).
fof(f6460,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7488,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6461,f6018])).
fof(f6461,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f6018,f6018,f955])).
fof(f7487,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6462,f6018])).
fof(f6462,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7486,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6463,f6018])).
fof(f6463,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1102,f6018,f955])).
fof(f7485,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6464,f6018])).
fof(f6464,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f1102,f6018,f955])).
fof(f7484,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6466,f6018])).
fof(f6466,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f1102,f6018,f955])).
fof(f7483,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6471,f6018])).
fof(f6471,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f955])).
fof(f7482,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6472,f6018])).
fof(f6472,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2327,f6018,f6018,f955])).
fof(f7481,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6473,f6018])).
fof(f6473,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f6018,f6018,f955])).
fof(f7480,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6474,f6018])).
fof(f6474,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2328,f6018,f6018,f955])).
fof(f7479,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6476,f6018])).
fof(f6476,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f6018,f6018,f955])).
fof(f7478,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6481,f1161])).
fof(f6481,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1974,f6018,f959])).
fof(f959,plain,(
  ( ! [X2,X0,X1] : (sP17(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f597])).
fof(f1974,plain,(
  ~sP17('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f634,f1161,f961])).
fof(f7477,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6482,f1161])).
fof(f6482,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1978,f6018,f959])).
fof(f7476,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6483,f4839])).
fof(f4839,plain,(
  s(bool,'T_0') != s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f1084])).
fof(f1084,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1083])).
fof(f1083,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1030])).
fof(f1030,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X38,X110,X37,X111] : ((s(bool,X38) = s(bool,X110) & (p(s(bool,X110)) => s(bool,X37) = s(bool,X111))) => ((p(s(bool,X38)) => p(s(bool,X37))) <=> (p(s(bool,X110)) => p(s(bool,X111)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.IMP_CONG')).
fof(f6483,plain,(
  s(bool,'T_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4663,f6018,f959])).
fof(f7475,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6490,f1161])).
fof(f6490,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2338,f6018,f959])).
fof(f7474,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6491,f1161])).
fof(f6491,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2339,f6018,f959])).
fof(f2339,plain,(
  ~sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1161,f1098,f961])).
fof(f7471,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6505,f4655])).
fof(f4655,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | $spl3),
  inference(unit_resulting_resolution,[],[f1102,f4610,f784])).
fof(f6505,plain,(
  s(bool,'F_0') != s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4666,f6018,f962])).
fof(f4666,plain,(
  sP17('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f1139,f1130,f1135,f4610,f945])).
fof(f7470,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6507,f741])).
fof(f6507,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2335,f6018,f962])).
fof(f2335,plain,(
  ( ! [X0] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f741,f1098,f960])).
fof(f960,plain,(
  ( ! [X2,X0,X1] : (sP17(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f597])).
fof(f7469,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6508,f894])).
fof(f894,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X34,X35] : ! [X38] : s(X35,X38) = s(X35,'const.sum.OUTL_1'(s('type.sum.sum'(X35,X34),'const.sum.INL_1'(s(X35,X38)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sum.OUTL')).
fof(f6508,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2336,f6018,f962])).
fof(f2336,plain,(
  ( ! [X0,X1] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f894,f1098,f960])).
fof(f7468,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6509,f895])).
fof(f895,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X35,X34] : ! [X38] : s(X34,X38) = s(X34,'const.sum.OUTR_1'(s('type.sum.sum'(X35,X34),'const.sum.INR_1'(s(X34,X38)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sum.OUTR')).
fof(f6509,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2337,f6018,f962])).
fof(f2337,plain,(
  ( ! [X0,X1] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f895,f1098,f960])).
fof(f7467,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6513,f741])).
fof(f6513,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2331,f6018,f962])).
fof(f2331,plain,(
  ( ! [X0] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f741,f1098,f960])).
fof(f7466,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6515,f895])).
fof(f6515,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2333,f6018,f962])).
fof(f2333,plain,(
  ( ! [X0,X1] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f895,f1098,f960])).
fof(f7465,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6516,f894])).
fof(f6516,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2332,f6018,f962])).
fof(f2332,plain,(
  ( ! [X0,X1] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f894,f1098,f960])).
fof(f7464,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6555,f7173])).
fof(f7173,plain,(
  ( ! [X0,X1] : (sP19('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1142,f6018,f1087])).
fof(f1087,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X0)) | ~sP18(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f969])).
fof(f969,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP18(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1,X2] : ((~sP19(X0,X1,X2) | (sP18(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP18(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP19(X0,X1,X2)))),
  inference(rectify,[],[f599])).
fof(f599,plain,(
  ! [X2,X0,X1] : ((~sP19(X2,X0,X1) | (sP18(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP18(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP19(X2,X0,X1)))),
  inference(flattening,[],[f598])).
fof(f598,plain,(
  ! [X2,X0,X1] : ((~sP19(X2,X0,X1) | (sP18(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP18(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP19(X2,X0,X1)))),
  inference(nnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X2,X0,X1] : (sP19(X2,X0,X1) <=> (sP18(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP19])])).
fof(f1142,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f971])).
fof(f971,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f602])).
fof(f602,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(flattening,[],[f601])).
fof(f601,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : (sP18(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f6555,plain,(
  ( ! [X0,X1] : (~sP19('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f964])).
fof(f964,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f600])).
fof(f7463,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6563,f7173])).
fof(f6563,plain,(
  ( ! [X0,X1] : (~sP19('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f965])).
fof(f965,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f600])).
fof(f7462,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6573,f6577])).
fof(f6577,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f973])).
fof(f973,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f602])).
fof(f6573,plain,(
  ( ! [X0] : (~sP18('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f634,f6018,f970])).
fof(f970,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f602])).
fof(f7461,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6583,f1102])).
fof(f6583,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2477,f6018,f974])).
fof(f974,plain,(
  ( ! [X2,X0,X1] : (sP19(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP19(X2,X0,X1)) & (~sP19(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP19(X2,X0,X1)) & (~sP19(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP19(X2,X0,X1))),
  inference(definition_folding,[],[f288,f420,f419])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X101,X99,X100] : ((p(s(bool,X100)) <=> (p(s(bool,X99)) & p(s(bool,X101)))) <=> ((p(s(bool,X100)) | ~p(s(bool,X99)) | ~p(s(bool,X101))) & (p(s(bool,X99)) | ~p(s(bool,X100))) & (p(s(bool,X101)) | ~p(s(bool,X100)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sat.sat_23_prove_ltb')).
fof(f2477,plain,(
  ~sP19('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2472,f963])).
fof(f963,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | sP18(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f600])).
fof(f2472,plain,(
  ~sP18('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f634,f1102,f634,f970])).
fof(f7460,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6584,f6018])).
fof(f6584,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2340,f6018,f974])).
fof(f2340,plain,(
  ( ! [X0] : (~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f964])).
fof(f7459,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6585,f6018])).
fof(f6585,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2341,f6018,f974])).
fof(f2341,plain,(
  ( ! [X0] : (~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f965])).
fof(f7458,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6586,f6018])).
fof(f6586,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4747,f6018,f974])).
fof(f4747,plain,(
  ~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f634,f634,f4610,f979])).
fof(f979,plain,(
  ( ! [X2,X0,X1] : (~sP19(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f604])).
fof(f7457,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6591,f1102])).
fof(f6591,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f974])).
fof(f1238,plain,(
  ( ! [X0] : (~sP19('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f634,f1102,f964])).
fof(f7456,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6592,f6018])).
fof(f6592,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4726,f6018,f974])).
fof(f4726,plain,(
  ( ! [X0] : (~sP19('T_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f964])).
fof(f7455,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6593,f6018])).
fof(f6593,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f6018,f974])).
fof(f1239,plain,(
  ( ! [X0] : (~sP19('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f1102,f965])).
fof(f7454,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6594,f6018])).
fof(f6594,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4728,f6018,f974])).
fof(f4728,plain,(
  ( ! [X0] : (~sP19('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f965])).
fof(f7453,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6596,f1102])).
fof(f6596,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2340,f6018,f974])).
fof(f7452,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6597,f6018])).
fof(f6597,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2341,f6018,f974])).
fof(f7451,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6607,f1102])).
fof(f6607,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2477,f6018,f975])).
fof(f975,plain,(
  ( ! [X2,X0,X1] : (sP19(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f604])).
fof(f7450,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6608,f6018])).
fof(f6608,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2340,f6018,f975])).
fof(f7449,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6609,f6018])).
fof(f6609,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2341,f6018,f975])).
fof(f7448,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6610,f6018])).
fof(f6610,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4747,f6018,f975])).
fof(f7447,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6615,f6018])).
fof(f6615,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f975])).
fof(f7446,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6616,f6018])).
fof(f6616,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4726,f6018,f975])).
fof(f7445,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6617,f1102])).
fof(f6617,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f6018,f975])).
fof(f7444,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6618,f6018])).
fof(f6618,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4728,f6018,f975])).
fof(f7443,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6620,f6018])).
fof(f6620,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2340,f6018,f975])).
fof(f7442,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6621,f1102])).
fof(f6621,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2341,f6018,f975])).
fof(f7441,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6653,f6954])).
fof(f6954,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f6018,f6018,f991])).
fof(f991,plain,(
  ( ! [X2,X0,X1] : (sP21(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP21(X2,X0,X1)) & (~sP21(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f611])).
fof(f611,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP21(X2,X0,X1)) & (~sP21(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP21(X2,X0,X1))),
  inference(definition_folding,[],[f290,f423,f422])).
fof(f422,plain,(
  ! [X2,X0,X1] : (sP20(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP20])])).
fof(f423,plain,(
  ! [X2,X0,X1] : (sP21(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP20(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP21])])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X101,X99,X100] : ((p(s(bool,X100)) <=> (p(s(bool,X99)) | p(s(bool,X101)))) <=> ((p(s(bool,X100)) | ~p(s(bool,X99))) & (p(s(bool,X100)) | ~p(s(bool,X101))) & (p(s(bool,X99)) | p(s(bool,X101)) | ~p(s(bool,X100)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sat.sat_25_prove_ltb')).
fof(f6653,plain,(
  ( ! [X0,X1] : (~sP21(X0,X1,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f980])).
fof(f980,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP20(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP20(X0,X1,X2) | sP21(X0,X1,X2)))),
  inference(rectify,[],[f606])).
fof(f606,plain,(
  ! [X2,X0,X1] : ((~sP21(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP20(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP20(X2,X0,X1) | sP21(X2,X0,X1)))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X2,X0,X1] : ((~sP21(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP20(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP20(X2,X0,X1)) | sP21(X2,X0,X1)))),
  inference(nnf_transformation,[],[f423])).
fof(f7440,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6671,f6954])).
fof(f6671,plain,(
  ( ! [X0,X1] : (~sP21(X0,'T_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f981])).
fof(f981,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f607])).
fof(f7439,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6710,f6018])).
fof(f6710,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2343,f6018,f986])).
fof(f986,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2) | ~sP20(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f607])).
fof(f2343,plain,(
  ( ! [X0] : (~sP21('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f980])).
fof(f1149,plain,(
  ( ! [X0,X1] : (sP20('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f990])).
fof(f990,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0,X1,X2] : ((~sP20(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP20(X0,X1,X2)))),
  inference(rectify,[],[f609])).
fof(f609,plain,(
  ! [X2,X0,X1] : ((~sP20(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP20(X2,X0,X1)))),
  inference(flattening,[],[f608])).
fof(f608,plain,(
  ! [X2,X0,X1] : ((~sP20(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP20(X2,X0,X1)))),
  inference(nnf_transformation,[],[f422])).
fof(f7438,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6711,f6018])).
fof(f6711,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1253,f6018,f986])).
fof(f1253,plain,(
  ( ! [X0] : (~sP21('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f1102,f981])).
fof(f7437,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6712,f6018])).
fof(f6712,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1253,f6018,f986])).
fof(f1148,plain,(
  ( ! [X0,X1] : (sP20(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f634,f989])).
fof(f989,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f610])).
fof(f7436,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6714,f6018])).
fof(f6714,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2351,f1253,f6018,f986])).
fof(f2351,plain,(
  ( ! [X0,X1] : (sP20(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f988])).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f610])).
fof(f7435,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6715,f6018])).
fof(f6715,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2347,f6018,f986])).
fof(f2347,plain,(
  ( ! [X0] : (~sP21('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f981])).
fof(f7434,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6716,f6018])).
fof(f6716,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2352,f2347,f6018,f986])).
fof(f2352,plain,(
  ( ! [X0,X1] : (sP20(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f989])).
fof(f7433,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6718,f6018])).
fof(f6718,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2351,f2347,f6018,f986])).
fof(f7432,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6720,f6018])).
fof(f6720,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4782,f4749,f6018,f986])).
fof(f4749,plain,(
  ( ! [X0] : (~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f981])).
fof(f4782,plain,(
  ( ! [X0,X1] : (sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4610,f990])).
fof(f7431,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6721,f6018])).
fof(f6721,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4749,f6018,f986])).
fof(f7430,plain,(
  $false | ($spl0 | $spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6723,f6018])).
fof(f6723,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f2351,f4749,f6018,f986])).
fof(f7429,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6729,f6854])).
fof(f6854,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f990])).
fof(f6729,plain,(
  ~sP20('F_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1102,f6018,f986])).
fof(f7428,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6731,f6854])).
fof(f6731,plain,(
  ~sP20('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1102,f6018,f986])).
fof(f7427,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6735,f6854])).
fof(f6735,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0','F_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f1102,f6018,f986])).
fof(f7426,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6745,f6854])).
fof(f6745,plain,(
  ( ! [X0] : (~sP20('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f986])).
fof(f1249,plain,(
  ( ! [X0] : (~sP21('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f1102,f980])).
fof(f7425,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6747,f6854])).
fof(f6747,plain,(
  ( ! [X0] : (~sP20('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f986])).
fof(f7424,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6749,f6854])).
fof(f6749,plain,(
  ( ! [X0] : (~sP20('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f986])).
fof(f7423,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6751,f6854])).
fof(f6751,plain,(
  ( ! [X0] : (~sP20('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f986])).
fof(f7422,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6755,f6854])).
fof(f6755,plain,(
  ( ! [X0] : (~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f986])).
fof(f4748,plain,(
  ( ! [X0] : (~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f634,f4610,f980])).
fof(f7421,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6757,f6854])).
fof(f6757,plain,(
  ( ! [X0] : (~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0',X0)) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f986])).
fof(f7420,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6759,f6018])).
fof(f6759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1147,f6018,f986])).
fof(f1147,plain,(
  ( ! [X0,X1] : (sP20(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f988])).
fof(f7419,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6761,f6018])).
fof(f6761,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1147,f6018,f986])).
fof(f7418,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6762,f6018])).
fof(f6762,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f1147,f6018,f986])).
fof(f7417,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6764,f6018])).
fof(f6764,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f2351,f6018,f986])).
fof(f7416,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6766,f6018])).
fof(f6766,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f2351,f6018,f986])).
fof(f7415,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6769,f6018])).
fof(f6769,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f2352,f6018,f986])).
fof(f7414,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6770,f6018])).
fof(f6770,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1149,f6018,f986])).
fof(f7413,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6771,f6018])).
fof(f6771,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1149,f6018,f986])).
fof(f7412,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6773,f6018])).
fof(f6773,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1149,f6018,f986])).
fof(f7411,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6774,f6018])).
fof(f6774,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f4782,f6018,f986])).
fof(f7410,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6776,f6854])).
fof(f6776,plain,(
  ~sP20('F_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f6018,f986])).
fof(f7409,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6777,f6854])).
fof(f6777,plain,(
  ~sP20('F_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1102,f6018,f986])).
fof(f7408,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6780,f6854])).
fof(f6780,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'F_0','T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f1102,f6018,f986])).
fof(f7407,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6790,f6854])).
fof(f6790,plain,(
  ( ! [X0] : (~sP20('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f986])).
fof(f7406,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6791,f6854])).
fof(f6791,plain,(
  ( ! [X0] : (~sP20('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f986])).
fof(f7405,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6792,f6854])).
fof(f6792,plain,(
  ( ! [X0] : (~sP20('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f986])).
fof(f7404,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6793,f6854])).
fof(f6793,plain,(
  ( ! [X0] : (~sP20('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f986])).
fof(f7403,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6796,f6854])).
fof(f6796,plain,(
  ( ! [X0] : (~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f986])).
fof(f7402,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6797,f6854])).
fof(f6797,plain,(
  ( ! [X0] : (~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0',X0)) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f986])).
fof(f7401,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6813,f6018])).
fof(f6813,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2351,f634,f6018,f987])).
fof(f987,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f610])).
fof(f7400,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6814,f6018])).
fof(f6814,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f634,f6018,f987])).
fof(f7399,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6815,f6018])).
fof(f6815,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2352,f634,f6018,f987])).
fof(f7398,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6821,f6854])).
fof(f6821,plain,(
  ( ! [X0] : (~sP20('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f1102,f6018,f987])).
fof(f7397,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6831,f6854])).
fof(f6831,plain,(
  ( ! [X0,X1] : (~sP20('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f6018,f987])).
fof(f7396,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6833,f6018])).
fof(f6833,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f634,f6018,f987])).
fof(f7395,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6834,f6018])).
fof(f6834,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2351,f634,f6018,f987])).
fof(f7394,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6836,f6018])).
fof(f6836,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2352,f634,f6018,f987])).
fof(f7393,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6842,f6854])).
fof(f6842,plain,(
  ( ! [X0] : (~sP20('T_0','F_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f1102,f6018,f987])).
fof(f7392,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6852,f6854])).
fof(f6852,plain,(
  ( ! [X0,X1] : (~sP20('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018,f6018,f987])).
fof(f7391,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6857,f6018])).
fof(f6857,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1102,f6018,f991])).
fof(f7390,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6858,f6018])).
fof(f6858,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1102,f6018,f991])).
fof(f7389,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6859,f6018])).
fof(f6859,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1102,f6018,f991])).
fof(f7388,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6861,f6018])).
fof(f6861,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7387,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6862,f6018])).
fof(f6862,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f991])).
fof(f7386,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6863,f6018])).
fof(f6863,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7385,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6864,f1102])).
fof(f6864,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f6018,f6018,f991])).
fof(f2543,plain,(
  ~sP21('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2538,f982])).
fof(f982,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | sP20(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f607])).
fof(f2538,plain,(
  ~sP20('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f634,f1102,f1102,f987])).
fof(f7384,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6865,f1102])).
fof(f6865,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f4545,plain,(
  ~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1102,f2814])).
fof(f2814,plain,(
  ( ! [X12,X11] : (~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X12,X11) | p(s(bool,X12)) | p(s(bool,X11))) ) | $spl0),
  inference(subsumption_resolution,[],[f2761,f634])).
fof(f2761,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | p(s(bool,X12)) | ~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X12,X11)) ) | $spl0),
  inference(superposition,[],[f994,f2315])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP21(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f612])).
fof(f7383,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6867,f6018])).
fof(f6867,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f991])).
fof(f7382,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6868,f1102])).
fof(f6868,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1102,f6018,f991])).
fof(f7381,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6869,f1102])).
fof(f6869,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1102,f6018,f991])).
fof(f7380,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6871,f6018])).
fof(f6871,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f1102,f6018,f991])).
fof(f7379,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6872,f6018])).
fof(f6872,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f1102,f6018,f991])).
fof(f7378,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6877,f1102])).
fof(f6877,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f991])).
fof(f7377,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6878,f1102])).
fof(f6878,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7376,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6879,f1102])).
fof(f6879,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f991])).
fof(f7375,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6880,f1102])).
fof(f6880,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7374,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6882,f6018])).
fof(f6882,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f7373,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6883,f6018])).
fof(f6883,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f991])).
fof(f7372,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6884,f6018])).
fof(f6884,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f991])).
fof(f7371,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6889,f6018])).
fof(f6889,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f6018,f991])).
fof(f7370,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6890,f6018])).
fof(f6890,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1102,f6018,f991])).
fof(f7369,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6892,f6018])).
fof(f6892,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1102,f6018,f991])).
fof(f7368,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6893,f6018])).
fof(f6893,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f991])).
fof(f7367,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6894,f6018])).
fof(f6894,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7366,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6896,f6018])).
fof(f6896,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7365,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6897,f1102])).
fof(f6897,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f6018,f6018,f991])).
fof(f7364,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6898,f1102])).
fof(f6898,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f7363,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6899,f6018])).
fof(f6899,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f991])).
fof(f7362,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6901,f1102])).
fof(f6901,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f6018,f991])).
fof(f7361,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6902,f1102])).
fof(f6902,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1102,f6018,f991])).
fof(f7360,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6904,f6018])).
fof(f6904,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f1102,f6018,f991])).
fof(f7359,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6905,f6018])).
fof(f6905,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f1102,f6018,f991])).
fof(f7358,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6910,f1102])).
fof(f6910,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f991])).
fof(f7357,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6911,f1102])).
fof(f6911,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7356,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6912,f1102])).
fof(f6912,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f991])).
fof(f7355,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6913,f1102])).
fof(f6913,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7354,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6915,f6018])).
fof(f6915,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f7353,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6916,f6018])).
fof(f6916,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f991])).
fof(f7352,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6917,f6018])).
fof(f6917,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f991])).
fof(f7351,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6923,f6018])).
fof(f6923,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f1102,f6018,f991])).
fof(f7350,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6924,f1102])).
fof(f6924,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f1102,f6018,f991])).
fof(f7349,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6925,f1102])).
fof(f6925,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f1102,f6018,f991])).
fof(f7348,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6927,f6018])).
fof(f6927,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f991])).
fof(f7347,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6928,f6018])).
fof(f6928,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7346,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6930,f6018])).
fof(f6930,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7345,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6931,f1102])).
fof(f6931,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f6018,f6018,f991])).
fof(f7344,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6932,f1102])).
fof(f6932,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f7343,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6933,f6018])).
fof(f6933,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f6018,f6018,f991])).
fof(f7342,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6936,f6018])).
fof(f6936,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f1102,f6018,f991])).
fof(f7341,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6937,f1102])).
fof(f6937,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f1102,f6018,f991])).
fof(f7340,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6938,f1102])).
fof(f6938,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f1102,f6018,f991])).
fof(f7339,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6945,f6018])).
fof(f6945,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f6018,f6018,f991])).
fof(f7338,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6946,f6018])).
fof(f6946,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f6018,f991])).
fof(f7337,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6947,f6018])).
fof(f6947,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f6018,f6018,f991])).
fof(f7336,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6948,f1102])).
fof(f6948,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f6018,f6018,f991])).
fof(f7335,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6949,f1102])).
fof(f6949,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4545,f6018,f6018,f991])).
fof(f7334,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6951,f6018])).
fof(f6951,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4749,f6018,f6018,f991])).
fof(f7333,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6957,f6018])).
fof(f6957,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2353,f634,f6018,f994])).
fof(f2353,plain,(
  ( ! [X0] : (sP21('T_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f992])).
fof(f992,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP21(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f612])).
fof(f7332,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6958,f6018])).
fof(f6958,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1262,f634,f6018,f994])).
fof(f1262,plain,(
  ( ! [X0] : (sP21('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f634,f993])).
fof(f993,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP21(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f612])).
fof(f7331,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6959,f6018])).
fof(f6959,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2358,f634,f6018,f994])).
fof(f2358,plain,(
  ( ! [X0] : (sP21('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f993])).
fof(f7330,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7014,f6018])).
fof(f7014,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1256,f634,f6018,f994])).
fof(f1256,plain,(
  ( ! [X0] : (sP21('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f634,f992])).
fof(f7329,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7015,f6018])).
fof(f7015,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2353,f634,f6018,f994])).
fof(f7328,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7017,f6018])).
fof(f7017,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2358,f634,f6018,f994])).
fof(f7327,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7022,f1102])).
fof(f7022,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1287,f6018,f997])).
fof(f997,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP22(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP22(X0,X1,X2) | sP23(X0,X1,X2)))),
  inference(rectify,[],[f614])).
fof(f614,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP22(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP22(X2,X0,X1) | sP23(X2,X0,X1)))),
  inference(flattening,[],[f613])).
fof(f613,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP22(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP22(X2,X0,X1)) | sP23(X2,X0,X1)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X2,X0,X1] : (sP23(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP22(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f1287,plain,(
  ( ! [X0] : (sP23('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1102,f634,f1009])).
fof(f1009,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP23(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP23(X2,X0,X1)) & (~sP23(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP23(X2,X0,X1)) & (~sP23(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP23(X2,X0,X1))),
  inference(definition_folding,[],[f379,f426,f425])).
fof(f425,plain,(
  ! [X2,X0,X1] : (sP22(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f379,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X101,X99,X100] : ((p(s(bool,X100)) <=> (p(s(bool,X99)) => p(s(bool,X101)))) <=> ((p(s(bool,X100)) | p(s(bool,X99))) & (p(s(bool,X100)) | ~p(s(bool,X101))) & (~p(s(bool,X99)) | p(s(bool,X101)) | ~p(s(bool,X100)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.sat.sat_27_prove_ltb')).
fof(f7326,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7023,f6018])).
fof(f7023,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1303,f6018,f997])).
fof(f1303,plain,(
  ( ! [X0] : (sP23('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f634,f1010])).
fof(f1010,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP23(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f620])).
fof(f7325,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7024,f6018])).
fof(f7024,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2368,f6018,f997])).
fof(f2368,plain,(
  ( ! [X0] : (sP23('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f1010])).
fof(f7324,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7026,f1102])).
fof(f7026,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2377,f6018,f997])).
fof(f2377,plain,(
  ( ! [X0] : (sP23('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1151,f1098,f1085])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP22(X0,X1,X2) | sP23(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1000])).
fof(f1000,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | ~sP22(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f615])).
fof(f1151,plain,(
  ( ! [X0,X1] : (sP22(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1102,f1005])).
fof(f1005,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP22(X0,X1,X2)))),
  inference(rectify,[],[f617])).
fof(f617,plain,(
  ! [X2,X0,X1] : ((~sP22(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP22(X2,X0,X1)))),
  inference(flattening,[],[f616])).
fof(f616,plain,(
  ! [X2,X0,X1] : ((~sP22(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP22(X2,X0,X1)))),
  inference(nnf_transformation,[],[f425])).
fof(f7323,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7027,f6018])).
fof(f7027,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2378,f6018,f997])).
fof(f2378,plain,(
  ( ! [X0] : (sP23('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1153,f1098,f1085])).
fof(f1153,plain,(
  ( ! [X0,X1] : (sP22(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f634,f1006])).
fof(f1006,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f618])).
fof(f7322,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7028,f1102])).
fof(f7028,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1370,f6018,f997])).
fof(f1370,plain,(
  ( ! [X0,X1] : (sP23('const.bool.BOUNDED_1'(s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1102,f1290])).
fof(f1290,plain,(
  ( ! [X2,X0,X1] : (sP23('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1288,f634])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP23('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1)) )),
  inference(superposition,[],[f1009,f641])).
fof(f7321,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7029,f6018])).
fof(f7029,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1389,f6018,f997])).
fof(f1389,plain,(
  ( ! [X0,X1] : (sP23('const.bool.BOUNDED_1'(s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f634,f1306])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP23('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(subsumption_resolution,[],[f1304,f634])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP23('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(superposition,[],[f1010,f641])).
fof(f7320,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7030,f6018])).
fof(f7030,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2383,f6018,f997])).
fof(f2383,plain,(
  ( ! [X0,X1] : (sP23('const.bool.BOUNDED_1'(s(bool,X0)),'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f1306])).
fof(f7319,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7031,f6018])).
fof(f7031,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1392,f6018,f997])).
fof(f1392,plain,(
  ( ! [X2,X0,X1] : (sP23('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(subsumption_resolution,[],[f1390,f634])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP23('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(superposition,[],[f1306,f641])).
fof(f7318,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7039,f1102])).
fof(f7039,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2619,f6018,f997])).
fof(f2619,plain,(
  sP23('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1154,f634,f1102,f1003])).
fof(f1003,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | ~sP22(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f615])).
fof(f1154,plain,(
  ( ! [X0,X1] : (sP22('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f1007])).
fof(f1007,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f618])).
fof(f7317,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7040,f6018])).
fof(f7040,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2377,f6018,f997])).
fof(f7316,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7041,f6018])).
fof(f7041,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2378,f6018,f997])).
fof(f7315,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7042,f6018])).
fof(f7042,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1370,f6018,f997])).
fof(f7314,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7043,f6018])).
fof(f7043,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1389,f6018,f997])).
fof(f7313,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7044,f6018])).
fof(f7044,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2383,f6018,f997])).
fof(f7312,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7045,f6018])).
fof(f7045,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1392,f6018,f997])).
fof(f7311,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7080,f7096])).
fof(f7096,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f1007])).
fof(f7080,plain,(
  ~sP22('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1277,f634,f6018,f1003])).
fof(f1277,plain,(
  ( ! [X0] : (~sP23('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f1102,f998])).
fof(f998,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f615])).
fof(f7310,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7082,f7096])).
fof(f7082,plain,(
  ~sP22('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2363,f634,f6018,f1003])).
fof(f2363,plain,(
  ( ! [X0] : (~sP23('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1102,f1098,f998])).
fof(f7309,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7092,f7096])).
fof(f7092,plain,(
  ( ! [X0] : (~sP22('T_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f634,f634,f6018,f1004])).
fof(f1004,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f618])).
fof(f7308,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7105,f1102])).
fof(f7105,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1277,f634,f6018,f1008])).
fof(f1008,plain,(
  ( ! [X2,X0,X1] : (sP23(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f620])).
fof(f7307,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7106,f1102])).
fof(f7106,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2363,f634,f6018,f1008])).
fof(f7306,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7113,f6018])).
fof(f7113,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2363,f634,f6018,f1008])).
fof(f7305,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7114,f1102])).
fof(f7114,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2635,f634,f6018,f1008])).
fof(f2635,plain,(
  ~sP23('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2629,f999])).
fof(f999,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | sP22(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f615])).
fof(f2629,plain,(
  ~sP22('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f634,f634,f1102,f1004])).
fof(f7304,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7179,f1142])).
fof(f7179,plain,(
  ( ! [X0] : (~sP18(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f1087])).
fof(f7303,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f7181,f1142])).
fof(f7181,plain,(
  ( ! [X0] : (~sP18(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4726,f6018,f1087])).
fof(f7302,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7183,f1142])).
fof(f7183,plain,(
  ( ! [X0] : (~sP18('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f6018,f1087])).
fof(f7301,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f7185,f1142])).
fof(f7185,plain,(
  ( ! [X0] : (~sP18('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4728,f6018,f1087])).
fof(f7300,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7188,f2342])).
fof(f2342,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f971])).
fof(f7188,plain,(
  ( ! [X0] : (~sP18(X0,'F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2340,f6018,f1087])).
fof(f7299,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7190,f2342])).
fof(f7190,plain,(
  ( ! [X0] : (~sP18('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2341,f6018,f1087])).
fof(f7298,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7194,f6018])).
fof(f7194,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2317,f6018,f1117])).
fof(f1117,plain,(
  ( ! [X0,X1] : (~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f862,f1102])).
fof(f862,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f549])).
fof(f549,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f2317,plain,(
  ( ! [X0] : (sP12(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f550])).
fof(f7297,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7195,f6018])).
fof(f7195,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1160,f6018,f1117])).
fof(f1160,plain,(
  ( ! [X2,X3] : (sP12(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(subsumption_resolution,[],[f1158,f634])).
fof(f1158,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP12(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(superposition,[],[f863,f641])).
fof(f7296,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7196,f6018])).
fof(f7196,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1416,f6018,f1117])).
fof(f1416,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f642,f863])).
fof(f7295,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7197,f6018])).
fof(f7197,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1128,f6018,f1117])).
fof(f1128,plain,(
  ( ! [X0] : (sP12('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f634,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f550])).
fof(f7294,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7198,f6018])).
fof(f7198,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2318,f6018,f1117])).
fof(f2318,plain,(
  ( ! [X0] : (sP12('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f864])).
fof(f7293,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7199,f6018])).
fof(f7199,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1159,f6018,f1117])).
fof(f1159,plain,(
  ( ! [X0,X1] : (sP12('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(subsumption_resolution,[],[f1157,f634])).
fof(f1157,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP12('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f864,f641])).
fof(f7292,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7200,f6018])).
fof(f7200,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6018,f1117])).
fof(f1417,plain,(
  ( ! [X0,X1] : (sP12('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f642,f864])).
fof(f7291,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7205,f6018])).
fof(f7205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1127,f6018,f1117])).
fof(f1127,plain,(
  ( ! [X0] : (sP12(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f863])).
fof(f7290,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7206,f6018])).
fof(f7206,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2317,f6018,f1117])).
fof(f7289,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7207,f6018])).
fof(f7207,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1160,f6018,f1117])).
fof(f7288,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7208,f6018])).
fof(f7208,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1416,f6018,f1117])).
fof(f7287,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7210,f6018])).
fof(f7210,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2318,f6018,f1117])).
fof(f7286,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7211,f6018])).
fof(f7211,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1159,f6018,f1117])).
fof(f7285,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7212,f6018])).
fof(f7212,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1417,f6018,f1117])).
fof(f7284,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7228,f6018])).
fof(f7228,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2380,f6018,f2814])).
fof(f2380,plain,(
  ( ! [X0] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1148,f1098,f1086])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP20(X0,X1,X2) | sP21(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f983])).
fof(f983,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2) | ~sP20(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f607])).
fof(f7283,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7232,f6018])).
fof(f7232,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2379,f6018,f2814])).
fof(f2379,plain,(
  ( ! [X0] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1147,f1098,f1086])).
fof(f7250,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f646,f6018,f829])).
fof(f829,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X27,X33] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X33))),s('type.patSem.v',X27)))) <=> s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X33))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_lit|split|2')).
fof(f646,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X33] : s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X33))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X33),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.patSem.Boolv_def')).
fof(f7248,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f741,f6018,f831])).
fof(f831,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X27,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))),s('type.patSem.v',X27)))) <=> s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_lit|split|0')).
fof(f7246,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f643,f6018,f827])).
fof(f827,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f528])).
fof(f643,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.exh_to_patProof.compile_v_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.exh_to_patProof.compile_v_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X13] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))) = s('type.patSem.v','const.exh_to_patProof.compile_v_1'(s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X13)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.compile_v_def|split|4')).
fof(f7244,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f741,f6018,f825])).
fof(f825,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))))))),
  inference(nnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_lit|split|1')).
fof(f7242,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f646,f6018,f823])).
fof(f823,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))))))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X33] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X33)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X33))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_lit|split|3')).
fof(f7240,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2310,f6018,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP6(X0,X1)) & (~sP6(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> sP6(X0,X1))),
  inference(definition_folding,[],[f365,f402,f401,f400,f399,f398])).
fof(f398,plain,(
  ! [X2,X4,X3] : (sP2(X2,X4,X3) <=> ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f399,plain,(
  ! [X2,X4,X3] : (sP3(X2,X4,X3) <=> ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f400,plain,(
  ! [X2,X4,X3] : (sP4(X2,X4,X3) <=> ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f401,plain,(
  ! [X2,X4,X3] : (sP5(X2,X4,X3) <=> ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f402,plain,(
  ! [X0,X1] : (sP6(X0,X1) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | sP2(X2,X4,X3) | sP3(X2,X4,X3) | sP4(X2,X4,X3) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | sP5(X2,X4,X3)) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f365,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(ennf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (! [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X5,X14] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X5),s('type.patSem.v',X14)))) <=> ! [X15] : (! [X16,X17] : ((? [X7] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X18,X19,X20] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20))))) | ? [X21,X22,X23,X24] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X22),s('type.patLang.exp',X24))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.patLang.exp',X23))),s('type.patLang.exp',X24))))) | ? [X21,X22,X25,X26,X13] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X13))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X22),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X26))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.list.list'('type.patLang.exp'),X25),s('type.list.list'('type.patLang.exp'),X26))))) | ? [X13] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13)))) | ? [X19,X20] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.patSem.v',X16))),s('type.patSem.v',X17))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.patSem.v',X5))),s('type.patSem.v',X14))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_def')).
fof(f2310,plain,(
  sP6(sK24,'const.patSem.Loc_1'(s('type.num.num',sK25))) | $spl0),
  inference(unit_resulting_resolution,[],[f1098,f816])).
fof(f816,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP6(X0,X1)) )),
  inference(cnf_transformation,[],[f523])).
fof(f7238,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2309,f6018,f1115])).
fof(f1115,plain,(
  ( ! [X0,X1] : (~sP11(X1,X0) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(forward_demodulation,[],[f861,f783])).
fof(f861,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | ~sP11(X1,X0)) )),
  inference(cnf_transformation,[],[f548])).
fof(f7235,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6018,f2380,f6018,f2814])).
fof(f7233,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1102,f2379,f6018,f2814])).
fof(f7229,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1102,f2380,f6018,f2814])).
fof(f7227,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6018,f2379,f6018,f2814])).
fof(f7223,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2321,f6018,f2811])).
fof(f2811,plain,(
  ( ! [X6] : (sP13(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | p(s(bool,X6))) ) | $spl0),
  inference(subsumption_resolution,[],[f2758,f634])).
fof(f2758,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP13(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) ) | $spl0),
  inference(superposition,[],[f882,f2315])).
fof(f2321,plain,(
  ~sP13('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f878])).
fof(f878,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f557])).
fof(f7221,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2319,f6018,f2810])).
fof(f2810,plain,(
  ( ! [X5] : (sP13('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X5) | p(s(bool,X5))) ) | $spl0),
  inference(subsumption_resolution,[],[f2757,f634])).
fof(f2757,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP13('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),X5)) ) | $spl0),
  inference(superposition,[],[f879,f2315])).
fof(f2319,plain,(
  ~sP13('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f634,f1098,f878])).
fof(f7216,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1127,f6018,f6018,f1117])).
fof(f7214,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1127,f1102,f6018,f1117])).
fof(f7204,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1127,f6018,f6018,f1117])).
fof(f7202,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1128,f1102,f6018,f1117])).
fof(f7191,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1145,f2341,f6018,f1087])).
fof(f1145,plain,(
  ( ! [X0,X1] : (sP18('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f973])).
fof(f7189,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1143,f2340,f6018,f1087])).
fof(f1143,plain,(
  ( ! [X0,X1] : (sP18(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1102,f972])).
fof(f972,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f602])).
fof(f7186,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4734,f4728,f6018,f1087])).
fof(f4734,plain,(
  ( ! [X0,X1] : (sP18('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4610,f973])).
fof(f7184,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1145,f1239,f6018,f1087])).
fof(f7182,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4733,f4726,f6018,f1087])).
fof(f4733,plain,(
  ( ! [X0,X1] : (sP18(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))),s('type.patSem.v',sK24)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4610,f972])).
fof(f7180,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1143,f1238,f6018,f1087])).
fof(f7161,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f752,f6018,f1084])).
fof(f752,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X35] : ! [X57,X113] : s(X35,X57) = s(X35,'const.ind_type.FCONS_3'(s(X35,X57),s(fun('type.num.num',X35),X113),s('type.num.num','const.num.|48|_0')))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.ind_type.FCONS|split|0')).
fof(f7159,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f894,f6018,f1084])).
fof(f7157,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f895,f6018,f1084])).
fof(f7155,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f899,f6018,f1084])).
fof(f899,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X35,X34] : ! [X38,X37] : s(X34,X37) = s(X34,'const.pair.SND_1'(s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X38),s(X34,X37)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.pair.SND')).
fof(f7153,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f898,f6018,f1084])).
fof(f898,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X34,X35] : ! [X38,X37] : s(X35,X38) = s(X35,'const.pair.FST_1'(s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X38),s(X34,X37)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.pair.FST')).
fof(f7151,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f918,f6018,f1084])).
fof(f918,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK153(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f575])).
fof(f575,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK153(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK153(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK153(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK153])],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X34,X35] : ! [X115,X116] : ? [X117] : (s(X34,X115) = s(X34,happ(s(fun('type.list.list'(X35),X34),X117),s('type.list.list'(X35),'const.list.NIL_0'))) & ! [X28,X29] : s(X34,happ(s(fun('type.list.list'(X35),X34),X117),s('type.list.list'(X35),'const.list.CONS_2'(s(X35,X28),s('type.list.list'(X35),X29))))) = s(X34,happ(s(fun(X34,X34),happ(s(fun('type.list.list'(X35),fun(X34,X34)),happ(s(fun(X35,fun('type.list.list'(X35),fun(X34,X34))),X116),s(X35,X28))),s('type.list.list'(X35),X29))),s(X34,happ(s(fun('type.list.list'(X35),X34),X117),s('type.list.list'(X35),X29))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.list.list_Axiom')).
fof(f7149,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f741,f6018,f1084])).
fof(f7147,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f634,f2315,f6018,f1084])).
fof(f7145,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f641,f6018,f1084])).
fof(f7142,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1303,f634,f634,f6018,f1011])).
fof(f1011,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP23(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f620])).
fof(f7130,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2635,f634,f6018,f1009])).
fof(f7124,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1277,f634,f6018,f6018,f1008])).
fof(f7118,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2635,f634,f1102,f6018,f1008])).
fof(f7111,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1277,f6018,f634,f6018,f1008])).
fof(f7109,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1277,f1102,f634,f6018,f1008])).
fof(f7097,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2629,f6018,f1007])).
fof(f7095,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2629,f6018,f1005])).
fof(f7093,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1153,f634,f634,f6018,f1004])).
fof(f7083,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1154,f2363,f634,f6018,f1003])).
fof(f7081,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1154,f1277,f634,f6018,f1003])).
fof(f7061,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1303,f634,f6018,f998])).
fof(f7049,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1287,f6018,f6018,f997])).
fof(f7047,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2619,f1102,f6018,f997])).
fof(f7035,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1287,f6018,f6018,f997])).
fof(f7033,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1287,f1102,f6018,f997])).
fof(f7021,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1256,f6018,f634,f6018,f994])).
fof(f7019,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1256,f1102,f634,f6018,f994])).
fof(f6988,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1256,f634,f6018,f6018,f994])).
fof(f6970,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1262,f634,f1102,f6018,f994])).
fof(f6955,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f6018,f991])).
fof(f6953,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1102,f6018,f6018,f991])).
fof(f6943,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f1102,f6018,f991])).
fof(f6941,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2543,f1102,f1102,f6018,f991])).
fof(f6921,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f6018,f991])).
fof(f6919,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f6018,f6018,f991])).
fof(f6909,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f1102,f6018,f991])).
fof(f6907,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f1102,f6018,f991])).
fof(f6888,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6018,f6018,f6018,f991])).
fof(f6886,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1102,f6018,f6018,f991])).
fof(f6876,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f6018,f1102,f6018,f991])).
fof(f6874,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1102,f1102,f6018,f991])).
fof(f6855,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2538,f6018,f990])).
fof(f6853,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f634,f6018,f6018,f987])).
fof(f6843,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f634,f1102,f6018,f987])).
fof(f6832,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f634,f6018,f6018,f987])).
fof(f6822,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f634,f1102,f6018,f987])).
fof(f6809,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f4782,f6018,f6018,f986])).
fof(f6807,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1149,f6018,f6018,f986])).
fof(f6805,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2347,f2352,f6018,f6018,f986])).
fof(f6803,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1253,f1148,f6018,f6018,f986])).
fof(f6801,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f2351,f6018,f6018,f986])).
fof(f6799,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1147,f6018,f6018,f986])).
fof(f6788,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4748,f4782,f1102,f6018,f986])).
fof(f6786,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1149,f1102,f6018,f986])).
fof(f6784,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2343,f2351,f1102,f6018,f986])).
fof(f6782,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f1147,f1102,f6018,f986])).
fof(f6758,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4782,f4749,f6018,f6018,f986])).
fof(f6756,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4782,f4748,f6018,f6018,f986])).
fof(f6752,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2347,f6018,f6018,f986])).
fof(f6750,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1253,f6018,f6018,f986])).
fof(f6748,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2343,f6018,f6018,f986])).
fof(f6746,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1249,f6018,f6018,f986])).
fof(f6736,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4782,f4749,f1102,f6018,f986])).
fof(f6732,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2347,f1102,f6018,f986])).
fof(f6730,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1253,f1102,f6018,f986])).
fof(f6672,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1262,f634,f6018,f981])).
fof(f6654,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1256,f634,f6018,f980])).
fof(f6636,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2463,f634,f634,f6018,f979])).
fof(f2463,plain,(
  sP19('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1142,f634,f634,f966])).
fof(f966,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP18(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f600])).
fof(f6626,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f6018,f975])).
fof(f6624,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f1102,f6018,f975])).
fof(f6614,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f6018,f975])).
fof(f6612,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2477,f1102,f6018,f975])).
fof(f6602,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f6018,f974])).
fof(f6600,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f1102,f6018,f974])).
fof(f6590,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1238,f6018,f6018,f974])).
fof(f6588,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2477,f1102,f6018,f974])).
fof(f6578,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2472,f6018,f973])).
fof(f6576,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2472,f6018,f972])).
fof(f6574,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1142,f634,f634,f6018,f970])).
fof(f6564,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2463,f634,f6018,f965])).
fof(f6556,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2463,f634,f6018,f964])).
fof(f6546,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2336,f894,f6018,f962])).
fof(f6544,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2337,f895,f6018,f962])).
fof(f6539,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2335,f741,f6018,f962])).
fof(f6537,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2742,f2315,f6018,f962])).
fof(f2742,plain,(
  sP17('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))) | $spl0),
  inference(unit_resulting_resolution,[],[f634,f2315,f960])).
fof(f6535,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1960,f641,f6018,f962])).
fof(f1960,plain,(
  ( ! [X0] : (sP17('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f634,f641,f960])).
fof(f6531,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2332,f894,f6018,f962])).
fof(f6529,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2333,f895,f6018,f962])).
fof(f6524,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2331,f741,f6018,f962])).
fof(f6522,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2740,f2315,f6018,f962])).
fof(f2740,plain,(
  sP17('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f634,f2315,f960])).
fof(f6520,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1955,f641,f6018,f962])).
fof(f1955,plain,(
  ( ! [X0] : (sP17('T_0','const.bool.BOUNDED_1'(s(bool,X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f634,f641,f960])).
fof(f7270,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f6512])).
fof(f6512,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3732,f6018,f962])).
fof(f3732,plain,(
  sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2327,f1130,f1135,f1102,f945])).
fof(f7269,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f6511])).
fof(f6511,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3698,f6018,f962])).
fof(f3698,plain,(
  sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1141,f1129,f2326,f634,f944])).
fof(f944,plain,(
  ( ! [X2,X0,X1] : (~sP15(X1,X2,X0) | ~p(s(bool,X2)) | ~sP16(X0,X2,X1) | sP17(X0,X1,X2) | ~sP14(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f588])).
fof(f7268,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6503])).
fof(f6503,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3731,f6018,f962])).
fof(f3731,plain,(
  sP17('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1139,f1130,f1135,f1102,f945])).
fof(f7267,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6502])).
fof(f6502,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3693,f6018,f962])).
fof(f3693,plain,(
  sP17('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1141,f1129,f1134,f634,f944])).
fof(f6495,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1978,f1161,f6018,f959])).
fof(f6493,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1974,f1161,f6018,f959])).
fof(f6480,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f6018,f955])).
fof(f6478,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1102,f6018,f6018,f955])).
fof(f6470,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f1102,f6018,f955])).
fof(f6468,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1102,f1102,f6018,f955])).
fof(f6452,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f6018,f955])).
fof(f6450,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f1102,f6018,f6018,f955])).
fof(f6442,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f6018,f1102,f6018,f955])).
fof(f6440,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1102,f1102,f6018,f955])).
fof(f6424,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f6018,f6018,f955])).
fof(f6422,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1102,f6018,f6018,f955])).
fof(f6414,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1139,f6018,f1102,f6018,f955])).
fof(f6412,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1141,f1102,f1102,f6018,f955])).
fof(f6396,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2445,f6018,f954])).
fof(f2445,plain,(
  ~sP15('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f634,f1102,f634,f951])).
fof(f6394,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2445,f6018,f953])).
fof(f6392,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1134,f634,f634,f6018,f951])).
fof(f6382,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2434,f6018,f950])).
fof(f2434,plain,(
  ~sP16('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f634,f1102,f634,f947])).
fof(f6380,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2434,f6018,f949])).
fof(f6378,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2181,f634,f634,f6018,f947])).
fof(f2181,plain,(
  ( ! [X0] : (sP16('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1955,f942])).
fof(f942,plain,(
  ( ! [X2,X0,X1] : (~sP17(X0,X1,X2) | sP16(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f588])).
fof(f6338,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f4680,f6018,f946])).
fof(f6333,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f4680,f6018,f946])).
fof(f6318,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f1135,f6018,f946])).
fof(f6316,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f1135,f6018,f946])).
fof(f6311,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f1135,f6018,f946])).
fof(f6308,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f1135,f6018,f946])).
fof(f6297,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f2326,f6018,f946])).
fof(f6288,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f2326,f6018,f946])).
fof(f6278,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f1134,f6018,f946])).
fof(f6273,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f1134,f6018,f946])).
fof(f6264,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f1134,f6018,f946])).
fof(f6262,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f1134,f6018,f946])).
fof(f6209,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f4680,f6018,f945])).
fof(f6204,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f4680,f6018,f945])).
fof(f6189,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f1135,f6018,f945])).
fof(f6187,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f1135,f6018,f945])).
fof(f6182,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f1135,f6018,f945])).
fof(f6179,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f1135,f6018,f945])).
fof(f6168,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1130,f2326,f6018,f945])).
fof(f6159,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f2338,f1129,f2326,f6018,f945])).
fof(f6149,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f4675,f1134,f6018,f945])).
fof(f6144,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1130,f1134,f6018,f945])).
fof(f6135,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1140,f4663,f1129,f1134,f6018,f945])).
fof(f6133,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1140,f1978,f1129,f1134,f6018,f945])).
fof(f6107,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1960,f634,f634,f6018,f943])).
fof(f6088,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1197,f634,f6018,f882])).
fof(f1197,plain,(
  ~sP13('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f634,f634,f878])).
fof(f6077,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1197,f634,f6018,f879])).
fof(f6066,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1199,f6018,f6018,f877])).
fof(f6064,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1208,f1102,f6018,f877])).
fof(f6052,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1199,f6018,f6018,f877])).
fof(f6050,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1199,f1102,f6018,f877])).
fof(f6038,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1102,f1161,f6018,f784])).
fof(f6036,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f1161,f6018,f784])).
fof(f6032,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1102,f1161,f6018,f784])).
fof(f6030,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6018,f1161,f6018,f784])).
fof(f6026,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1161,f6018,f691])).
fof(f691,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X32] : (s(bool,'F_0') = s(bool,X32) <=> ~p(s(bool,X32)))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f6024,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f641,f6018,f686])).
fof(f6022,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f634,f6018])).
fof(f6020,plain,(
  p(s(bool,sK94('_i_0_0',bool))) | $spl10),
  inference(cnf_transformation,[],[f6020_D])).
fof(f6020_D,plain,(
  p(s(bool,sK94('_i_0_0',bool))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f6017,plain,(
  ( ! [X3] : (p(s(bool,sK94('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f6006,f741])).
fof(f6006,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK94('_i_0_0',bool)))))) )),
  inference(superposition,[],[f746,f741])).
fof(f746,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK94(X1,X0)))))) )),
  inference(cnf_transformation,[],[f480])).
fof(f5332,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f5332_D])).
fof(f5332_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5333,plain,(
  ( ! [X3] : (~p(s(bool,sK96(X3,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f5333_D])).
fof(f5333_D,plain,(
  ( ! [X3] : (~p(s(bool,sK96(X3,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5323,plain,(
  ( ! [X3] : (~p(s(bool,sK96(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f755,f741])).
fof(f755,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK96(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK96(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK96(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK96])],[f360])).
fof(f360,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X35] : ! [X41,X32] : (! [X38] : (s(X35,X32) = s(X35,X38) => p(s(bool,happ(s(fun(X35,bool),X41),s(X35,X38))))) => p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X41)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.bool.JRH_INDUCT_UTIL')).
fof(f2308,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2307,f1102])).
fof(f2307,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2288,f1180])).
fof(f1180,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v',sK24))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1126,f691])).
fof(f1126,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v',sK24)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1099,f1100])).
fof(f2288,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v',sK24)))) | $spl2),
  inference(superposition,[],[f645,f1100])).
fof(f645,plain,(
  ( ! [X0] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_rules|split|4')).
fof(f1098,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | $spl0),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1100,plain,(
  s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) | $spl2),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f632,plain,(
  s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))))),
  inference(cnf_transformation,[],[f429])).
fof(f429,plain,(
  (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | s('type.patSem.v',sK24) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | s('type.patSem.v',sK24) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f428])).
fof(f428,plain,(
  ? [X0,X1] : ((p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f338])).
fof(f338,plain,(
  ? [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))) <~> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|1','thm.exh_to_patProof.v_rel_loc|split|1')).
fof(f1099,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) | $spl1),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1101,plain,(
  s('type.patSem.v',sK24) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) | $spl3),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  s('type.patSem.v',sK24) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f633,plain,(
  s('type.patSem.v',sK24) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))) | ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK24),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK25))))))),
  inference(cnf_transformation,[],[f429])).
% SZS output end Proof for thm.exh_to_patProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 611
% Passive clauses: 2144
% Generated clauses: 11473
% Final active clauses: 563
% Final passive clauses: 1390
% Input formulas: 133
% Initial clauses: 441
% 
% Duplicate literals: 8
% Trivial inequalities: 8
% Fw subsumption resolutions: 337
% Fw demodulations: 1963
% Bw demodulations: 40
% Condensations: 19
% 
% Simple tautologies: 92
% Equational tautologies: 39
% Forward subsumptions: 5666
% Fw demodulations to eq. taut.: 720
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 4973
% Factoring: 1
% Forward superposition: 2206
% Backward superposition: 1397
% Self superposition: 20
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 2633
% SAT solver unit clauses: 901
% SAT solver binary clauses: 1232
% 
% Sat splits: 10
% Sat splitting refutations: 467
% 
% Memory used [KB]: 10106
% Time elapsed: 0.593 s
% ------------------------------
% ------------------------------
% Success in time 1.94 s
