Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% (2996)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 176
% Passive clauses: 1005
% Generated clauses: 12395
% Final active clauses: 160
% Final passive clauses: 679
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 633
% Trivial inequalities: 282
% Fw subsumption resolutions: 2063
% Fw demodulations: 1844
% Bw demodulations: 10
% 
% Simple tautologies: 351
% Equational tautologies: 21
% Forward subsumptions: 5317
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 193
% Forward superposition: 4320
% Backward superposition: 2649
% Self superposition: 89
% Equality factoring: 47
% Equality resolution: 50
% 
% Split clauses: 447
% Split components: 906
% SAT solver clauses: 1001
% SAT solver unit clauses: 7
% SAT solver binary clauses: 77
% 
% Sat splits: 447
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1001
% TWLsolver calls for satisfiability: 172
% 
% Memory used [KB]: 6140
% Time elapsed: 0.619 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_knownProof
% (3006)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 178
% Passive clauses: 518
% Generated clauses: 13289
% Final active clauses: 176
% Final passive clauses: 328
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 944
% Trivial inequalities: 294
% Fw subsumption resolutions: 1418
% Fw demodulations: 1738
% Bw demodulations: 13
% Condensations: 933
% 
% Simple tautologies: 286
% Equational tautologies: 19
% Forward subsumptions: 6278
% Fw demodulations to eq. taut.: 32
% 
% Binary resolution: 55
% Forward superposition: 4638
% Backward superposition: 2524
% Self superposition: 77
% Equality factoring: 34
% Equality resolution: 49
% 
% Split clauses: 66
% Split components: 141
% SAT solver clauses: 776
% SAT solver unit clauses: 72
% SAT solver binary clauses: 18
% 
% Sat splits: 66
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 776
% TWLsolver calls for satisfiability: 282
% 
% Memory used [KB]: 9466
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_knownProof
% (3034)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 372
% Passive clauses: 2668
% Generated clauses: 126333
% Final active clauses: 189
% Final passive clauses: 583
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 7173
% Trivial inequalities: 880
% Fw subsumption resolutions: 12339
% Fw demodulations: 8137
% Bw demodulations: 1390
% Condensations: 17012
% Global subsumptions: 765
% 
% Simple tautologies: 1233
% Equational tautologies: 1350
% Forward subsumptions: 73148
% Fw demodulations to eq. taut.: 101
% Bw demodulations to eq. taut.: 77
% 
% Unit resulting resolution: 5237
% Forward superposition: 32142
% Backward superposition: 29539
% Self superposition: 844
% Equality factoring: 453
% Equality resolution: 93
% 
% Split clauses: 238
% Split components: 496
% SAT solver clauses: 4984
% SAT solver unit clauses: 116
% SAT solver binary clauses: 182
% 
% Sat splits: 238
% Sat splitting refutations: 1073
% 
% Memory used [KB]: 63453
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f3869,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f157,f158,f177,f234,f238,f240,f243,f248,f250,f254,f257,f263,f266,f267,f272,f274,f278,f282,f283,f339,f341,f343,f347,f349,f353,f358,f360,f361,f364,f235,f367,f371,f381,f400,f412,f413,f430,f431,f446,f445,f444,f467,f466,f465,f483,f493,f494,f497,f510,f511,f512,f513,f514,f515,f534,f517,f518,f532,f519,f241,f547,f548,f550,f236,f653,f655,f656,f657,f660,f661,f694,f696,f697,f698,f701,f702,f745,f747,f748,f749,f752,f753,f767,f768,f799,f801,f796,f802,f946,f947,f949,f951,f952,f953,f954,f956,f958,f960,f961,f962,f966,f968,f970,f971,f973,f975,f977,f978,f981,f982,f984,f985,f988,f1001,f1003,f1004,f1005,f1008,f1009,f1011,f1013,f1014,f1016,f1018,f1019,f1058,f1060,f1061,f1062,f1065,f1066,f1078,f1079,f1081,f1083,f1084,f1085,f1088,f1089,f1127,f1128,f1125,f1129,f1133,f1135,f1136,f1137,f1140,f1141,f1145,f1147,f1148,f1150,f1152,f1153,f1219,f1221,f1222,f1224,f1226,f1227,f1259,f1261,f1262,f1263,f1266,f1267,f1281,f1282,f1321,f1323,f1324,f1325,f1328,f1329,f1343,f1344,f1390,f1392,f1393,f1394,f1397,f1398,f1412,f1413,f1438,f1439,f1460,f1462,f1463,f1464,f1467,f1468,f1474,f1610,f1611,f1612,f1617,f1619,f1613,f1627,f1628,f1632,f1634,f1422,f1635,f1636,f1637,f1353,f1638,f1639,f1640,f1291,f1641,f1642,f1215,f1644,f1107,f1647,f1056,f1648,f777,f1650,f1652,f1653,f1631,f1747,f1216,f1214,f1057,f1055,f159,f1801,f1803,f1804,f1806,f1808,f1809,f1978,f1986,f2045,f2047,f2048,f2050,f2052,f2053,f340,f2068,f162,f160,f2069,f166,f164,f2070,f2072,f2073,f2074,f2077,f2078,f674,f2096,f2095,f2201,f2202,f2203,f2205,f2206,f2204,f2211,f2213,f2214,f2215,f2218,f2219,f714,f2347,f1649,f2348,f1645,f1646,f2364,f2365,f2367,f2368,f2369,f2370,f2371,f2372,f2374,f2375,f2376,f2379,f2380,f1030,f2382,f2384,f2385,f2387,f2389,f2390,f1033,f2406,f2408,f2409,f2410,f2413,f2414,f765,f2439,f2441,f2442,f2443,f2446,f2447,f1098,f2468,f2470,f2471,f2472,f2475,f2476,f1076,f2497,f2499,f2500,f2501,f2504,f2505,f1175,f2509,f2511,f2512,f2514,f2516,f2517,f1179,f2545,f2547,f2548,f2549,f2552,f2553,f1279,f2571,f2574,f2582,f2583,f2586,f2588,f2589,f2591,f2593,f2594,f1249,f2638,f2640,f2641,f2642,f2645,f2646,f1341,f2667,f2573,f2668,f533,f2576,f2693,f2695,f2696,f2697,f2700,f2701,f1410,f2778,f2780,f2781,f2782,f2785,f2786,f2790,f2793,f2795,f2796,f2797,f2800,f2801,f2806,f2810,f2812,f2813,f2814,f2817,f2818,f2823,f2916,f2918,f2921,f2924,f2925,f2926,f2927,f2929,f2932,f2935,f2936,f2937,f2938,f2941,f2942,f2943,f2945,f2917,f2946,f2947,f2948,f2944,f2951,f2953,f2954,f2955,f2958,f2959,f2919,f2973,f2974,f2997,f2998,f3015,f3016,f3095,f3097,f3098,f3099,f3102,f3103,f3116,f3118,f3119,f3137,f3139,f3140,f3141,f3144,f3145,f3151,f3158,f3160,f3161,f3163,f3165,f3166,f3180,f3186,f3188,f3189,f3190,f3193,f3194,f3200,f3211,f3212,f3340,f3342,f3345,f3348,f3349,f3350,f3351,f3353,f3356,f3359,f3360,f3361,f3362,f3365,f3367,f3369,f1629,f3370,f3371,f3372,f3374,f3375,f3376,f3377,f3378,f3380,f3366,f3381,f3382,f3341,f3400,f3402,f3403,f3404,f3407,f3408,f3343,f3428,f3434,f3435,f3429,f3448,f3450,f3451,f3452,f3455,f3456,f3463,f3469,f3471,f3472,f3473,f3476,f3477,f3483,f3497,f3498,f3503,f342,f3504,f3524,f3525,f535,f3580,f3581,f3582,f3583,f3584,f3585,f3586,f3573,f3589,f3575,f3591,f3576,f3587,f3592,f3610,f3611,f3597,f3612,f3599,f3613,f3600,f242,f3614,f3615,f3617,f3618,f3436,f3661,f3368,f3677,f3679,f3680,f239,f3681,f264,f3684,f251,f3685,f265,f3691,f3836,f3838,f800,f3842,f1618,f3843,f1620,f3844,f3852,f3837,f3853,f3854,f1643,f3862,f3379,f3839,f3868])).
fof(f3868,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3867,f154])).
fof(f154,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd',aHL_TRUTH)).
fof(f3867,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3866,f161])).
fof(f161,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl2),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3866,plain,(
  ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3865,f2457])).
fof(f2457,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | $spl0),
  inference(forward_demodulation,[],[f2456,f121])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X17] : ! [X7,X9] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X17),X7))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X17),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.closSem.state'(X17),X7)))))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.closSem.closSem_65_prove_ltb')).
fof(f2456,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))))))) | $spl0),
  inference(subsumption_resolution,[],[f2450,f152])).
fof(f152,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd',aHL_FALSITY)).
fof(f2450,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))))))) | $spl0),
  inference(superposition,[],[f141,f160])).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X6,X7] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X7,X6,X2,X1,X0)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X6,X2,X1,X0)),s('type.closSem.v',X7))))))) & ((s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',sK6(X2,X1,X0))))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X6,X7] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X8] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X8))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X8),s('type.closSem.v',X7))))))) & (? [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(rectify,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4))))))) & (? [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) => ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) => ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))) <=> ! [X19,X20] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X20))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X19),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),X7))))) => ? [X21] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X21))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X19),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X21),s('type.closSem.v',X20)))))))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.clos_knownProof.state_globals_approx_def')).
fof(f3865,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f3864])).
fof(f3864,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f3690,f139])).
fof(f139,plain,(
  ( ! [X6,X2,X0,X7,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X7,X6,X2,X1,X0)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f3690,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3689,f152])).
fof(f3689,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3676,f160])).
fof(f3676,plain,(
  p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',sK6(X2,X1,X0))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f3675,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) | ($spl0 | $spl2)),
  inference(equality_resolution,[],[f2717])).
fof(f2717,plain,(
  ( ! [X0] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X0,sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',X0))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f2543,f2457])).
fof(f2543,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X1),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X0,X1,sK2,sK1,sK0)),s('type.closSem.v',X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2534,f156])).
fof(f156,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X10] : (s(X6,X10) = s(X6,X10) <=> p(s(bool,'T_0')))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.bool.REFL_CLAUSE')).
fof(f2534,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X0,X1,sK2,sK1,sK0)),s('type.closSem.v',X0)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X1),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1)))))) ) | $spl2),
  inference(superposition,[],[f140,f161])).
fof(f140,plain,(
  ( ! [X6,X2,X0,X7,X1] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X6,X2,X1,X0)),s('type.closSem.v',X7)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f3839,plain,(
  ( ! [X132,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X132) | s(bool,X131) != s(bool,X132)) ) | $spl216),
  inference(cnf_transformation,[],[f3839_D])).
fof(f3839_D,plain,(
  ( ! [X132,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X132) | s(bool,X131) != s(bool,X132)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f3379,plain,(
  ( ! [X608] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X608)))) ) | $spl174),
  inference(cnf_transformation,[],[f3379_D])).
fof(f3379_D,plain,(
  ( ! [X608] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X608)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f3862,plain,(
  ( ! [X263,X262,X264] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X262))) | s(bool,'T_0') = s(bool,X263) | s(bool,X264) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X263) != s(bool,X264)) ) | ($spl0 | $spl2 | $spl24 | $spl194)),
  inference(subsumption_resolution,[],[f3793,f2819])).
fof(f2819,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X41,sK1,sK0),sK5(X41,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X41,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X41)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f508])).
fof(f508,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(trivial_inequality_removal,[],[f502])).
fof(f502,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f340,f151])).
fof(f151,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd',aHL_BOOL_CASES)).
fof(f2802,plain,(
  ( ! [X36] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X36,sK1,sK0),sK5(X36,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X36,sK1,sK0))))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2776,f508])).
fof(f2776,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X0,sK1,sK0),sK5(X0,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X0,sK1,sK0))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) ) | $spl2),
  inference(equality_resolution,[],[f2628])).
fof(f2628,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X1))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X0,sK1,sK0)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X1,sK5(X0,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) ) | $spl2),
  inference(superposition,[],[f2543,f141])).
fof(f3793,plain,(
  ( ! [X263,X262,X264] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X262,sK1,sK0),sK5(X262,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X262,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X262))) | s(bool,'T_0') = s(bool,X263) | s(bool,X264) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X263) != s(bool,X264)) ) | ($spl0 | $spl2 | $spl24 | $spl194)),
  inference(superposition,[],[f3583,f3686])).
fof(f3686,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | ($spl0 | $spl2 | $spl24)),
  inference(resolution,[],[f3675,f508])).
fof(f1643,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl110),
  inference(cnf_transformation,[],[f1643_D])).
fof(f1643_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f3854,plain,(
  ( ! [X227,X225,X226] : (s(bool,X225) != s(bool,X226) | s(bool,'T_0') = s(bool,X227) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X225) != s(bool,X227)) ) | $spl218),
  inference(cnf_transformation,[],[f3854_D])).
fof(f3854_D,plain,(
  ( ! [X227,X225,X226] : (s(bool,X225) != s(bool,X226) | s(bool,'T_0') = s(bool,X227) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X225) != s(bool,X227)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f3853,plain,(
  ( ! [X229,X227,X225,X230,X228,X226,X224] : (s(bool,X225) != s(bool,X226) | s(bool,X225) != s(bool,X227) | s(bool,X224) != s(bool,X228) | s(bool,X229) != s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,X228) != s(bool,X229) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X227)) ) | ($spl0 | $spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3785,f1067])).
fof(f1067,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) != s(bool,X65) | s(bool,X62) != s(bool,X64) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,X63)) ) | ($spl24 | $spl40)),
  inference(resolution,[],[f765,f508])).
fof(f3785,plain,(
  ( ! [X229,X227,X225,X230,X228,X226,X224] : (s(bool,'T_0') != s(bool,X224) | s(bool,X225) != s(bool,X226) | s(bool,X225) != s(bool,X227) | s(bool,X224) != s(bool,X228) | s(bool,X229) != s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,X228) != s(bool,X229) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X227)) ) | ($spl0 | $spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f3104,f3686])).
fof(f3104,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X98,X104] : (s(bool,X101) != s(bool,X102) | s(bool,X98) != s(bool,X100) | s(bool,X98) != s(bool,X99) | s(bool,X102) != s(bool,X103) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,X105) | s(bool,X103) != s(bool,X104) | s(bool,X100) != s(bool,X101) | s(bool,'T_0') = s(bool,X99)) ) | ($spl24 | $spl78)),
  inference(resolution,[],[f1410,f508])).
fof(f3837,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X122)) ) | $spl214),
  inference(cnf_transformation,[],[f3837_D])).
fof(f3837_D,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'T_0') = s(bool,X122)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f3852,plain,(
  ( ! [X222,X223] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X223) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2 | $spl24 | $spl142)),
  inference(subsumption_resolution,[],[f3784,f2819])).
fof(f3784,plain,(
  ( ! [X222,X223] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X222,sK1,sK0),sK5(X222,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X222,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X223) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2 | $spl24 | $spl142)),
  inference(superposition,[],[f2944,f3686])).
fof(f3844,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X183) != s(bool,X185) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X187) | s(bool,X186) != s(bool,X187) | s(bool,X184) != s(bool,X188) | s(bool,'T_0') = s(bool,X188)) ) | ($spl0 | $spl2 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f3769,f1067])).
fof(f3769,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'T_0') != s(bool,X183) | s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X183) != s(bool,X185) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X187) | s(bool,X186) != s(bool,X187) | s(bool,X184) != s(bool,X188) | s(bool,'T_0') = s(bool,X188)) ) | ($spl0 | $spl2 | $spl24 | $spl74)),
  inference(superposition,[],[f1399,f3686])).
fof(f1399,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (s(bool,X91) != s(bool,X92) | s(bool,X89) != s(bool,X91) | s(bool,X92) != s(bool,X93) | s(bool,X93) != s(bool,X94) | s(bool,'F_0') = s(bool,X95) | s(bool,X94) != s(bool,X95) | s(bool,X89) != s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | ($spl24 | $spl74)),
  inference(resolution,[],[f1341,f508])).
fof(f1620,plain,(
  ( ! [X173,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'F_0') = s(bool,X173) | s(bool,X172) != s(bool,X173)) ) | $spl92),
  inference(cnf_transformation,[],[f1620_D])).
fof(f1620_D,plain,(
  ( ! [X173,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'F_0') = s(bool,X173) | s(bool,X172) != s(bool,X173)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3843,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,X164) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X163) != s(bool,X165) | s(bool,X165) != s(bool,X166) | s(bool,'F_0') = s(bool,X166) | s(bool,X164) != s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | ($spl0 | $spl2 | $spl24 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f3765,f754])).
fof(f754,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) != s(bool,X55) | s(bool,X53) != s(bool,X54) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X54)) ) | ($spl24 | $spl38)),
  inference(resolution,[],[f714,f508])).
fof(f3765,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,'T_0') != s(bool,X163) | s(bool,X164) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X163) != s(bool,X165) | s(bool,X165) != s(bool,X166) | s(bool,'F_0') = s(bool,X166) | s(bool,X164) != s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | ($spl0 | $spl2 | $spl24 | $spl70)),
  inference(superposition,[],[f1330,f3686])).
fof(f1330,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X82) != s(bool,X83) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | ($spl24 | $spl70)),
  inference(resolution,[],[f1279,f508])).
fof(f1618,plain,(
  ( ! [X151,X153] : (s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153)) ) | $spl90),
  inference(cnf_transformation,[],[f1618_D])).
fof(f1618_D,plain,(
  ( ! [X151,X153] : (s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3842,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X148) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X147) != s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) != s(bool,X150) | s(bool,'T_0') = s(bool,X150)) ) | ($spl0 | $spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f3761,f691])).
fof(f691,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(subsumption_resolution,[],[f683,f156])).
fof(f683,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(superposition,[],[f498,f151])).
fof(f498,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f340,f151])).
fof(f3761,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,'T_0') != s(bool,X147) | s(bool,X148) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X147) != s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) != s(bool,X150) | s(bool,'T_0') = s(bool,X150)) ) | ($spl0 | $spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1268,f3686])).
fof(f1268,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,X71) != s(bool,X73) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) != s(bool,X72) | s(bool,'T_0') = s(bool,X72)) ) | ($spl24 | $spl54)),
  inference(resolution,[],[f1076,f508])).
fof(f800,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) | $spl44),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3838,plain,(
  ( ! [X132,X130,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X130) | s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X132)) ) | ($spl0 | $spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f3755,f151])).
fof(f3755,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') != s(bool,X130) | s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X130) | s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X132)) ) | ($spl0 | $spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1067,f3686])).
fof(f3836,plain,(
  ( ! [X121,X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl0 | $spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f3750,f151])).
fof(f3750,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,X121) | s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl0 | $spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f754,f3686])).
fof(f3691,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | $spl212),
  inference(cnf_transformation,[],[f3691_D])).
fof(f3691_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f265,plain,(
  ( ! [X80] : (p(s(bool,X80))) ) | $spl22),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  ( ! [X80] : (p(s(bool,X80))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3685,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f155])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f112])).
fof(f112,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10,X15,X11,X16] : ((s(bool,X10) = s(bool,X15) & (p(s(bool,X15)) => s(bool,X11) = s(bool,X16))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X15)) => p(s(bool,X16)))))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.bool.IMP_CONG')).
fof(f251,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | ~p(s(bool,X45))) ) | $spl18),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | ~p(s(bool,X45))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3684,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X22) != s(bool,X23)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f264,plain,(
  ( ! [X81] : (~p(s(bool,X81))) ) | $spl20),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  ( ! [X81] : (~p(s(bool,X81))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3681,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f239,plain,(
  ( ! [X23,X25] : (s(bool,X23) != s(bool,X25) | p(s(bool,X23))) ) | $spl12),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  ( ! [X23,X25] : (s(bool,X23) != s(bool,X25) | p(s(bool,X23))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3680,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f113])).
fof(f113,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f3679,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,X7) != s(bool,X8)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f3677,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3675,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f99])).
fof(f3368,plain,(
  ( ! [X407,X401,X403,X404,X406,X402] : (s(bool,X406) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X401,X402,X403),sK5(X401,X402,X403),X404,X402,X403)),s('type.closSem.v',sK6(X401,X402,X403)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X401))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X404))) | s(bool,'T_0') = s(bool,X407) | s(bool,X406) != s(bool,X407)) ) | $spl168),
  inference(cnf_transformation,[],[f3368_D])).
fof(f3368_D,plain,(
  ( ! [X407,X401,X403,X404,X406,X402] : (s(bool,X406) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X401,X402,X403),sK5(X401,X402,X403),X404,X402,X403)),s('type.closSem.v',sK6(X401,X402,X403)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X401))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X404))) | s(bool,'T_0') = s(bool,X407) | s(bool,X406) != s(bool,X407)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f3661,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X18,X19,X20),sK5(X18,X19,X20),X21,X19,X20)),s('type.closSem.v',sK6(X18,X19,X20)))) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl2 | $spl24 | $spl34 | $spl194)),
  inference(subsumption_resolution,[],[f3653,f2819])).
fof(f3653,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X22,sK1,sK0),sK5(X22,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X22,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X18,X19,X20),sK5(X18,X19,X20),X21,X19,X20)),s('type.closSem.v',sK6(X18,X19,X20)))) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl24 | $spl34 | $spl194)),
  inference(superposition,[],[f3583,f3195])).
fof(f3195,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X79,X80,X81),sK5(X79,X80,X81),X82,X80,X81)),s('type.closSem.v',sK6(X79,X80,X81)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X81),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X82))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X81),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X79)))) ) | ($spl24 | $spl34)),
  inference(resolution,[],[f3168,f508])).
fof(f3168,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X81)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X81,X80,X79),sK5(X81,X80,X79),X82,X80,X79)),s('type.closSem.v',sK6(X81,X80,X79)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X82)))) ) | $spl34),
  inference(resolution,[],[f3154,f548])).
fof(f3154,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X5,X6,X7),sK5(X5,X6,X7),X8,X6,X7)),s('type.closSem.v',sK6(X5,X6,X7))))) )),
  inference(subsumption_resolution,[],[f3149,f152])).
fof(f3149,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'F_0')) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X5,X6,X7),sK5(X5,X6,X7),X8,X6,X7)),s('type.closSem.v',sK6(X5,X6,X7))))) )),
  inference(superposition,[],[f3135,f151])).
fof(f3135,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X0,X1,X2),sK5(X0,X1,X2),X3,X1,X2)),s('type.closSem.v',sK6(X0,X1,X2))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) )),
  inference(equality_resolution,[],[f1170])).
fof(f1170,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X5))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X3,X4)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X5,sK5(X2,X3,X4),X6,X3,X4)),s('type.closSem.v',X5)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X3),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(superposition,[],[f140,f141])).
fof(f3436,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | $spl180),
  inference(cnf_transformation,[],[f3436_D])).
fof(f3436_D,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f3618,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl24 | $spl34 | $spl38 | $spl110)),
  inference(subsumption_resolution,[],[f3608,f754])).
fof(f3608,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl24 | $spl34 | $spl110)),
  inference(superposition,[],[f1643,f3195])).
fof(f3617,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24 | $spl38 | $spl110)),
  inference(subsumption_resolution,[],[f3607,f754])).
fof(f3607,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24 | $spl110)),
  inference(superposition,[],[f1643,f2819])).
fof(f3615,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl38 | $spl110)),
  inference(subsumption_resolution,[],[f3601,f754])).
fof(f3601,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | $spl110),
  inference(superposition,[],[f1643,f151])).
fof(f3614,plain,(
  ( ! [X28,X26,X31,X29,X27,X32] : (s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X31) != s(bool,X32) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X26,X27,X28),sK5(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK6(X26,X27,X28))))) ) | $spl210),
  inference(cnf_transformation,[],[f3614_D])).
fof(f3614_D,plain,(
  ( ! [X28,X26,X31,X29,X27,X32] : (s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X31) != s(bool,X32) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X26,X27,X28),sK5(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK6(X26,X27,X28))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f242,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) | $spl16),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3600,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X26,X27,X28),sK5(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK6(X26,X27,X28)))) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl24 | $spl34 | $spl110)),
  inference(superposition,[],[f1643,f3195])).
fof(f3613,plain,(
  ( ! [X24,X25,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X22,sK1,sK0),sK5(X22,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X22,sK1,sK0))))) ) | $spl208),
  inference(cnf_transformation,[],[f3613_D])).
fof(f3613_D,plain,(
  ( ! [X24,X25,X22] : (s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X22,sK1,sK0),sK5(X22,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X22,sK1,sK0))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f3599,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X22,sK1,sK0),sK5(X22,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X22,sK1,sK0)))) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24 | $spl110)),
  inference(superposition,[],[f1643,f2819])).
fof(f3612,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl206),
  inference(cnf_transformation,[],[f3612_D])).
fof(f3612_D,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f3597,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X17) != s(bool,X18)) ) | ($spl2 | $spl110)),
  inference(superposition,[],[f1643,f161])).
fof(f3611,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | ($spl24 | $spl110)),
  inference(subsumption_resolution,[],[f3595,f691])).
fof(f3595,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') != s(bool,X11) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | $spl110),
  inference(superposition,[],[f1643,f151])).
fof(f3610,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | $spl110),
  inference(subsumption_resolution,[],[f3593,f1643])).
fof(f3593,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | $spl110),
  inference(superposition,[],[f1643,f151])).
fof(f3592,plain,(
  ( ! [X57,X64,X54,X62,X56,X55,X63] : (s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,X62) != s(bool,X63) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl204),
  inference(cnf_transformation,[],[f3592_D])).
fof(f3592_D,plain,(
  ( ! [X57,X64,X54,X62,X56,X55,X63] : (s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,X62) != s(bool,X63) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f3587,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X2) != s(bool,X3)) ) | $spl198),
  inference(cnf_transformation,[],[f3587_D])).
fof(f3587_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X2) != s(bool,X3)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f3576,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,'T_0') != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X58) != s(bool,X60) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl78)),
  inference(superposition,[],[f3104,f3195])).
fof(f3591,plain,(
  ( ! [X52,X46,X53,X51] : (s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53)) ) | $spl202),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  ( ! [X52,X46,X53,X51] : (s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f3575,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,'T_0') != s(bool,X47) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X47) != s(bool,X49) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f3104,f2819])).
fof(f3589,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) | $spl200),
  inference(cnf_transformation,[],[f3589_D])).
fof(f3589_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f3573,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X32) != s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f3104,f161])).
fof(f3586,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl54 | $spl78)),
  inference(subsumption_resolution,[],[f3569,f1269])).
fof(f1269,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) != s(bool,X78)) ) | $spl54),
  inference(resolution,[],[f1076,f152])).
fof(f3569,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl78)),
  inference(superposition,[],[f3104,f151])).
fof(f3585,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (s(bool,X59) != s(bool,X60) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,X59) != s(bool,X61)) ) | $spl196),
  inference(cnf_transformation,[],[f3585_D])).
fof(f3585_D,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (s(bool,X59) != s(bool,X60) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,X59) != s(bool,X61)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f3584,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X58) != s(bool,X62) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3568,f1067])).
fof(f3568,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,'T_0') != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X58) != s(bool,X62) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X54,X55,X56),sK5(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK6(X54,X55,X56)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl78)),
  inference(superposition,[],[f3104,f3195])).
fof(f3583,plain,(
  ( ! [X50,X48,X46,X49] : (s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50)) ) | $spl194),
  inference(cnf_transformation,[],[f3583_D])).
fof(f3583_D,plain,(
  ( ! [X50,X48,X46,X49] : (s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f3582,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X47) != s(bool,X51) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3567,f1067])).
fof(f3567,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,'T_0') != s(bool,X47) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X47) != s(bool,X51) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X46,sK1,sK0),sK5(X46,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X46,sK1,sK0)))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f3104,f2819])).
fof(f3581,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) | $spl192),
  inference(cnf_transformation,[],[f3581_D])).
fof(f3581_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f3580,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X32) != s(bool,X36) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3565,f1067])).
fof(f3565,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X32) != s(bool,X36) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f3104,f161])).
fof(f535,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl32),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3525,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) | $spl190),
  inference(cnf_transformation,[],[f3525_D])).
fof(f3525_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f3524,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl164)),
  inference(subsumption_resolution,[],[f3513,f508])).
fof(f3513,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5)) ) | $spl164),
  inference(superposition,[],[f3343,f151])).
fof(f3504,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | $spl188),
  inference(cnf_transformation,[],[f3504_D])).
fof(f3504_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f342,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26)) ) | $spl26),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3503,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl0 | $spl24 | $spl158)),
  inference(subsumption_resolution,[],[f3493,f340])).
fof(f3493,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0),X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl0 | $spl158)),
  inference(superposition,[],[f3200,f160])).
fof(f3498,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X8,X7,X6),sK5(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK6(X8,X7,X6))))) ) | $spl186),
  inference(cnf_transformation,[],[f3498_D])).
fof(f3498_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X8,X7,X6),sK5(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK6(X8,X7,X6))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f3497,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X8,X7,X6),sK5(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK6(X8,X7,X6)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl158)),
  inference(subsumption_resolution,[],[f3487,f508])).
fof(f3487,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X8,X7,X6),sK5(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK6(X8,X7,X6)))) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | $spl158),
  inference(superposition,[],[f3200,f151])).
fof(f3483,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X20,X21,X22),sK5(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK6(X20,X21,X22)))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) ) | $spl184),
  inference(cnf_transformation,[],[f3483_D])).
fof(f3483_D,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X20,X21,X22),sK5(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK6(X20,X21,X22)))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f3477,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X84,X85,X86),sK5(X84,X85,X86),X87,X85,X86)),s('type.closSem.v',sK6(X84,X85,X86)))) | s(bool,X88) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X86),X85),s('type.sptree.spt'('type.clos_known.val_approx'),X87))) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X89)) | p(s(bool,X90)) | s(bool,X90) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X86),X85),s('type.sptree.spt'('type.clos_known.val_approx'),X84)))) ) | $spl14),
  inference(resolution,[],[f3169,f155])).
fof(f3169,plain,(
  ( ! [X87,X85,X83,X86,X84] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X84),s('type.sptree.spt'('type.clos_known.val_approx'),X85)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X85,X84,X83),sK5(X85,X84,X83),X86,X84,X83)),s('type.closSem.v',sK6(X85,X84,X83)))) | s(bool,X87) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X84),s('type.sptree.spt'('type.clos_known.val_approx'),X86))) | s(bool,'F_0') = s(bool,X87)) ) | $spl14),
  inference(resolution,[],[f3154,f552])).
fof(f552,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f543,f156])).
fof(f543,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(superposition,[],[f241,f151])).
fof(f3476,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X76,X77,X78),sK5(X76,X77,X78),X79,X77,X78)),s('type.closSem.v',sK6(X76,X77,X78)))) | s(bool,X80) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X78),X77),s('type.sptree.spt'('type.clos_known.val_approx'),X79))) | s(bool,'F_0') = s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X78),X77),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) | s(bool,X82) != s(bool,X83)) ) | $spl14),
  inference(resolution,[],[f3169,f115])).
fof(f3473,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X52,X53,X54),sK5(X52,X53,X54),X55,X53,X54)),s('type.closSem.v',sK6(X52,X53,X54)))) | s(bool,X56) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X54),X53),s('type.sptree.spt'('type.clos_known.val_approx'),X55))) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X59) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X54),X53),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) ) | $spl14),
  inference(resolution,[],[f3169,f114])).
fof(f3472,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X44,X45,X46),sK5(X44,X45,X46),X47,X45,X46)),s('type.closSem.v',sK6(X44,X45,X46)))) | s(bool,X48) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X45),s('type.sptree.spt'('type.clos_known.val_approx'),X47))) | s(bool,'F_0') = s(bool,X48) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X45),s('type.sptree.spt'('type.clos_known.val_approx'),X44)))) ) | $spl14),
  inference(resolution,[],[f3169,f113])).
fof(f3471,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X36,X37,X38),sK5(X36,X37,X38),X39,X37,X38)),s('type.closSem.v',sK6(X36,X37,X38)))) | s(bool,X40) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X38),X37),s('type.sptree.spt'('type.clos_known.val_approx'),X39))) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X38),X37),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) | s(bool,X42) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f3169,f111])).
fof(f3469,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X20,X21,X22),sK5(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK6(X20,X21,X22)))) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) | $spl14),
  inference(resolution,[],[f3169,f109])).
fof(f3463,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))))) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20))),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | p(s(bool,X22))) ) | $spl182),
  inference(cnf_transformation,[],[f3463_D])).
fof(f3463_D,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))))) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20))),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | p(s(bool,X22))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f3456,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X73,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X74),s('type.closSem.state'(X75),X76)),X75)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X73,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X74),s('type.closSem.state'(X75),X76)),X75)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X75),X76))))) | p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X78) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X75),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X74),s('type.closSem.state'(X75),X76))),s('type.sptree.spt'('type.clos_known.val_approx'),X73)))) )),
  inference(resolution,[],[f1095,f155])).
fof(f1095,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.closSem.state'(X0),X2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.closSem.state'(X0),X2)),X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X3,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.closSem.state'(X0),X2)),X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) )),
  inference(superposition,[],[f141,f121])).
fof(f3455,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X66,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X68),X69)),X68)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X66,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X68),X69)),X68)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X68),X69))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X68),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X68),X69))),s('type.sptree.spt'('type.clos_known.val_approx'),X66))) | s(bool,X71) != s(bool,X72)) )),
  inference(resolution,[],[f1095,f115])).
fof(f3452,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X45,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X46),s('type.closSem.state'(X47),X48)),X47)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X45,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X46),s('type.closSem.state'(X47),X48)),X47)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X47),X48))))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X47),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X46),s('type.closSem.state'(X47),X48))),s('type.sptree.spt'('type.clos_known.val_approx'),X45)))) )),
  inference(resolution,[],[f1095,f114])).
fof(f3451,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X38,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X39),s('type.closSem.state'(X40),X41)),X40)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X38,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X39),s('type.closSem.state'(X40),X41)),X40)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X40),X41))))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X40),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X39),s('type.closSem.state'(X40),X41))),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) )),
  inference(resolution,[],[f1095,f113])).
fof(f3450,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X31,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X32),s('type.closSem.state'(X33),X34)),X33)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X31,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X32),s('type.closSem.state'(X33),X34)),X33)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X32),s('type.closSem.state'(X33),X34))),s('type.sptree.spt'('type.clos_known.val_approx'),X31))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f1095,f111])).
fof(f3448,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X17,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20)),X19)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X18),s('type.closSem.state'(X19),X20))),s('type.sptree.spt'('type.clos_known.val_approx'),X17)))) )),
  inference(resolution,[],[f1095,f109])).
fof(f3429,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,X2,X1),sK5(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK6(X4,X2,X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))))) ) | $spl178),
  inference(cnf_transformation,[],[f3429_D])).
fof(f3429_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,X2,X1),sK5(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK6(X4,X2,X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f3435,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X38,X36,X35),sK5(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK6(X38,X36,X35)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl10 | $spl24 | $spl34 | $spl156)),
  inference(subsumption_resolution,[],[f3422,f3328])).
fof(f3328,plain,(
  ( ! [X315,X316,X318,X317] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X315,X316,X317),sK5(X315,X316,X317),X318,X316,X317)),s('type.closSem.v',sK6(X315,X316,X317))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X318))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X315)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(trivial_inequality_removal,[],[f3277])).
fof(f3277,plain,(
  ( ! [X315,X316,X318,X317] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X315,X316,X317),sK5(X315,X316,X317),X318,X316,X317)),s('type.closSem.v',sK6(X315,X316,X317))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X318))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X315)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(superposition,[],[f236,f3195])).
fof(f3422,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X31,X32,X33),sK5(X31,X32,X33),X34,X32,X33)),s('type.closSem.v',sK6(X31,X32,X33))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X38,X36,X35),sK5(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK6(X38,X36,X35)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl24 | $spl34 | $spl156)),
  inference(superposition,[],[f3180,f3195])).
fof(f3434,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X30,X28,X27),sK5(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK6(X30,X28,X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl10 | $spl24 | $spl156)),
  inference(subsumption_resolution,[],[f3421,f2912])).
fof(f2912,plain,(
  ( ! [X157] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X157,sK1,sK0),sK5(X157,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X157,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X157)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(trivial_inequality_removal,[],[f2870])).
fof(f2870,plain,(
  ( ! [X157] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X157,sK1,sK0),sK5(X157,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X157,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X157)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(superposition,[],[f236,f2819])).
fof(f3421,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X26,sK1,sK0),sK5(X26,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X26,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X30,X28,X27),sK5(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK6(X30,X28,X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl24 | $spl156)),
  inference(superposition,[],[f3180,f2819])).
fof(f3428,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,X2,X1),sK5(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK6(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl34 | $spl156)),
  inference(subsumption_resolution,[],[f3417,f548])).
fof(f3417,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,X2,X1),sK5(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK6(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl156),
  inference(superposition,[],[f3180,f151])).
fof(f3343,plain,(
  ( ! [X83,X81,X86,X84,X82] : (s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X81,X82,X83),sK5(X81,X82,X83),X84,X82,X83)),s('type.closSem.v',sK6(X81,X82,X83)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X81))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X84))) | p(s(bool,X86))) ) | $spl164),
  inference(cnf_transformation,[],[f3343_D])).
fof(f3343_D,plain,(
  ( ! [X83,X81,X86,X84,X82] : (s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X81,X82,X83),sK5(X81,X82,X83),X84,X82,X83)),s('type.closSem.v',sK6(X81,X82,X83)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X81))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X84))) | p(s(bool,X86))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f3408,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63))) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X63,X61,X60),sK5(X63,X61,X60),X62,X61,X60)),s('type.closSem.v',sK6(X63,X61,X60))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f155])).
fof(f3407,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X55))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X56,X54,X53),sK5(X56,X54,X53),X55,X54,X53)),s('type.closSem.v',sK6(X56,X54,X53)))) | s(bool,X58) != s(bool,X59)) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f115])).
fof(f3404,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X35,X33,X32),sK5(X35,X33,X32),X34,X33,X32)),s('type.closSem.v',sK6(X35,X33,X32))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f114])).
fof(f3403,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,X26,X25),sK5(X28,X26,X25),X27,X26,X25)),s('type.closSem.v',sK6(X28,X26,X25))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f113])).
fof(f3402,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X21,X19,X18),sK5(X21,X19,X18),X20,X19,X18)),s('type.closSem.v',sK6(X21,X19,X18)))) | s(bool,X23) != s(bool,X24)) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f111])).
fof(f3400,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X7,X5,X4),sK5(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK6(X7,X5,X4))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3328,f109])).
fof(f3341,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X39,X40,X41),sK5(X39,X40,X41),X42,X40,X41)),s('type.closSem.v',sK6(X39,X40,X41))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X39))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X42)))) ) | $spl162),
  inference(cnf_transformation,[],[f3341_D])).
fof(f3341_D,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X39,X40,X41),sK5(X39,X40,X41),X42,X40,X41)),s('type.closSem.v',sK6(X39,X40,X41))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X39))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X42)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f3382,plain,(
  ( ! [X629,X631,X632,X630] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X632)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X629))))) ) | $spl176),
  inference(cnf_transformation,[],[f3382_D])).
fof(f3382_D,plain,(
  ( ! [X629,X631,X632,X630] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X632)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X629))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f3381,plain,(
  ( ! [X629,X631,X625,X627,X632,X628,X630,X626] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X632)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X625,X626,X627),sK5(X625,X626,X627),X628,X626,X627)),s('type.closSem.v',sK6(X625,X626,X627))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X629)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X627),X626),s('type.sptree.spt'('type.clos_known.val_approx'),X628))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X627),X626),s('type.sptree.spt'('type.clos_known.val_approx'),X625)))) ) | ($spl24 | $spl34 | $spl154)),
  inference(subsumption_resolution,[],[f3325,f3154])).
fof(f3325,plain,(
  ( ! [X629,X631,X625,X627,X632,X628,X630,X626] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X629,X630,X631),sK5(X629,X630,X631),X632,X630,X631)),s('type.closSem.v',sK6(X629,X630,X631)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X632)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X625,X626,X627),sK5(X625,X626,X627),X628,X626,X627)),s('type.closSem.v',sK6(X625,X626,X627))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X631),X630),s('type.sptree.spt'('type.clos_known.val_approx'),X629)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X627),X626),s('type.sptree.spt'('type.clos_known.val_approx'),X628))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X627),X626),s('type.sptree.spt'('type.clos_known.val_approx'),X625)))) ) | ($spl24 | $spl34 | $spl154)),
  inference(superposition,[],[f3151,f3195])).
fof(f3366,plain,(
  ( ! [X377,X372,X374,X373,X375] : (s(bool,X377) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X372,X373,X374),sK5(X372,X373,X374),X375,X373,X374)),s('type.closSem.v',sK6(X372,X373,X374)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X372))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X375))) | s(bool,'T_0') = s(bool,X377)) ) | $spl166),
  inference(cnf_transformation,[],[f3366_D])).
fof(f3366_D,plain,(
  ( ! [X377,X372,X374,X373,X375] : (s(bool,X377) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X372,X373,X374),sK5(X372,X373,X374),X375,X373,X374)),s('type.closSem.v',sK6(X372,X373,X374)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X372))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X375))) | s(bool,'T_0') = s(bool,X377)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f3380,plain,(
  ( ! [X613,X609,X611,X612,X614,X610] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X613))) | s(bool,'T_0') = s(bool,X614) | s(bool,X614) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X609,X610,X611),sK5(X609,X610,X611),X612,X610,X611)),s('type.closSem.v',sK6(X609,X610,X611)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X611),X610),s('type.sptree.spt'('type.clos_known.val_approx'),X612))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X611),X610),s('type.sptree.spt'('type.clos_known.val_approx'),X609)))) ) | ($spl2 | $spl24 | $spl34 | $spl142)),
  inference(subsumption_resolution,[],[f3323,f2819])).
fof(f3323,plain,(
  ( ! [X613,X609,X611,X612,X614,X610] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X613,sK1,sK0),sK5(X613,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X613,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X613))) | s(bool,'T_0') = s(bool,X614) | s(bool,X614) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X609,X610,X611),sK5(X609,X610,X611),X612,X610,X611)),s('type.closSem.v',sK6(X609,X610,X611)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X611),X610),s('type.sptree.spt'('type.clos_known.val_approx'),X612))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X611),X610),s('type.sptree.spt'('type.clos_known.val_approx'),X609)))) ) | ($spl24 | $spl34 | $spl142)),
  inference(superposition,[],[f2944,f3195])).
fof(f3378,plain,(
  ( ! [X606,X605,X607,X608,X604] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X608))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X604,X605,X606),sK5(X604,X605,X606),X607,X605,X606)),s('type.closSem.v',sK6(X604,X605,X606))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X606),X605),s('type.sptree.spt'('type.clos_known.val_approx'),X607))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X606),X605),s('type.sptree.spt'('type.clos_known.val_approx'),X604)))) ) | ($spl2 | $spl24 | $spl34 | $spl138)),
  inference(subsumption_resolution,[],[f3322,f2819])).
fof(f3322,plain,(
  ( ! [X606,X605,X607,X608,X604] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X608,sK1,sK0),sK5(X608,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X608,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X608))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X604,X605,X606),sK5(X604,X605,X606),X607,X605,X606)),s('type.closSem.v',sK6(X604,X605,X606))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X606),X605),s('type.sptree.spt'('type.clos_known.val_approx'),X607))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X606),X605),s('type.sptree.spt'('type.clos_known.val_approx'),X604)))) ) | ($spl24 | $spl34 | $spl138)),
  inference(superposition,[],[f2919,f3195])).
fof(f3377,plain,(
  ( ! [X598] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X598,sK1,sK0),sK5(X598,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X598,sK1,sK0))))) ) | $spl172),
  inference(cnf_transformation,[],[f3377_D])).
fof(f3377_D,plain,(
  ( ! [X598] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X598,sK1,sK0),sK5(X598,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X598,sK1,sK0))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f3376,plain,(
  ( ! [X596,X598,X594,X597,X595] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X594,X595,X596),sK5(X594,X595,X596),X597,X595,X596)),s('type.closSem.v',sK6(X594,X595,X596))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X598,sK1,sK0),sK5(X598,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X598,sK1,sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X596),X595),s('type.sptree.spt'('type.clos_known.val_approx'),X597))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X596),X595),s('type.sptree.spt'('type.clos_known.val_approx'),X594)))) ) | ($spl2 | $spl24 | $spl34 | $spl134)),
  inference(subsumption_resolution,[],[f3320,f2819])).
fof(f3320,plain,(
  ( ! [X596,X598,X594,X597,X595] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X598))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X594,X595,X596),sK5(X594,X595,X596),X597,X595,X596)),s('type.closSem.v',sK6(X594,X595,X596))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X598,sK1,sK0),sK5(X598,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X598,sK1,sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X596),X595),s('type.sptree.spt'('type.clos_known.val_approx'),X597))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X596),X595),s('type.sptree.spt'('type.clos_known.val_approx'),X594)))) ) | ($spl24 | $spl34 | $spl134)),
  inference(superposition,[],[f2823,f3195])).
fof(f3375,plain,(
  ( ! [X593] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X593))))) ) | $spl170),
  inference(cnf_transformation,[],[f3375_D])).
fof(f3375_D,plain,(
  ( ! [X593] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X593))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f3374,plain,(
  ( ! [X592,X590,X593,X589,X591] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X593)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X589,X590,X591),sK5(X589,X590,X591),X592,X590,X591)),s('type.closSem.v',sK6(X589,X590,X591))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X591),X590),s('type.sptree.spt'('type.clos_known.val_approx'),X592))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X591),X590),s('type.sptree.spt'('type.clos_known.val_approx'),X589)))) ) | ($spl2 | $spl24 | $spl34 | $spl132)),
  inference(subsumption_resolution,[],[f3319,f2802])).
fof(f3319,plain,(
  ( ! [X592,X590,X593,X589,X591] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X593,sK1,sK0),sK5(X593,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X593,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X593)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X589,X590,X591),sK5(X589,X590,X591),X592,X590,X591)),s('type.closSem.v',sK6(X589,X590,X591))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X591),X590),s('type.sptree.spt'('type.clos_known.val_approx'),X592))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X591),X590),s('type.sptree.spt'('type.clos_known.val_approx'),X589)))) ) | ($spl24 | $spl34 | $spl132)),
  inference(superposition,[],[f2806,f3195])).
fof(f3372,plain,(
  ( ! [X511,X516,X518,X512,X514,X510,X517,X519,X513,X515] : (s(bool,X515) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X510,X511,X512),sK5(X510,X511,X512),X513,X511,X512)),s('type.closSem.v',sK6(X510,X511,X512)))) | s(bool,X514) != s(bool,X516) | s(bool,X516) != s(bool,X517) | s(bool,'F_0') = s(bool,X518) | s(bool,X517) != s(bool,X518) | s(bool,X515) != s(bool,X519) | s(bool,'T_0') = s(bool,X519) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X512),X511),s('type.sptree.spt'('type.clos_known.val_approx'),X513))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X512),X511),s('type.sptree.spt'('type.clos_known.val_approx'),X510)))) ) | ($spl24 | $spl34 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f3308,f1067])).
fof(f3308,plain,(
  ( ! [X511,X516,X518,X512,X514,X510,X517,X519,X513,X515] : (s(bool,'T_0') != s(bool,X514) | s(bool,X515) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X510,X511,X512),sK5(X510,X511,X512),X513,X511,X512)),s('type.closSem.v',sK6(X510,X511,X512)))) | s(bool,X514) != s(bool,X516) | s(bool,X516) != s(bool,X517) | s(bool,'F_0') = s(bool,X518) | s(bool,X517) != s(bool,X518) | s(bool,X515) != s(bool,X519) | s(bool,'T_0') = s(bool,X519) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X512),X511),s('type.sptree.spt'('type.clos_known.val_approx'),X513))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X512),X511),s('type.sptree.spt'('type.clos_known.val_approx'),X510)))) ) | ($spl24 | $spl34 | $spl74)),
  inference(superposition,[],[f1399,f3195])).
fof(f3371,plain,(
  ( ! [X476,X478,X474,X480,X482,X475,X477,X479,X481] : (s(bool,X479) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X474,X475,X476),sK5(X474,X475,X476),X477,X475,X476)),s('type.closSem.v',sK6(X474,X475,X476)))) | s(bool,X478) != s(bool,X480) | s(bool,X480) != s(bool,X481) | s(bool,'F_0') = s(bool,X481) | s(bool,X479) != s(bool,X482) | s(bool,'T_0') = s(bool,X482) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X476),X475),s('type.sptree.spt'('type.clos_known.val_approx'),X477))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X476),X475),s('type.sptree.spt'('type.clos_known.val_approx'),X474)))) ) | ($spl24 | $spl34 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f3304,f754])).
fof(f3304,plain,(
  ( ! [X476,X478,X474,X480,X482,X475,X477,X479,X481] : (s(bool,'T_0') != s(bool,X478) | s(bool,X479) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X474,X475,X476),sK5(X474,X475,X476),X477,X475,X476)),s('type.closSem.v',sK6(X474,X475,X476)))) | s(bool,X478) != s(bool,X480) | s(bool,X480) != s(bool,X481) | s(bool,'F_0') = s(bool,X481) | s(bool,X479) != s(bool,X482) | s(bool,'T_0') = s(bool,X482) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X476),X475),s('type.sptree.spt'('type.clos_known.val_approx'),X477))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X476),X475),s('type.sptree.spt'('type.clos_known.val_approx'),X474)))) ) | ($spl24 | $spl34 | $spl70)),
  inference(superposition,[],[f1330,f3195])).
fof(f3370,plain,(
  ( ! [X445,X447,X443,X449,X444,X446,X442,X448] : (s(bool,X447) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X442,X443,X444),sK5(X442,X443,X444),X445,X443,X444)),s('type.closSem.v',sK6(X442,X443,X444)))) | s(bool,X446) != s(bool,X448) | s(bool,'F_0') = s(bool,X448) | s(bool,X447) != s(bool,X449) | s(bool,'T_0') = s(bool,X449) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X444),X443),s('type.sptree.spt'('type.clos_known.val_approx'),X445))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X444),X443),s('type.sptree.spt'('type.clos_known.val_approx'),X442)))) ) | ($spl24 | $spl34 | $spl54)),
  inference(subsumption_resolution,[],[f3300,f691])).
fof(f3300,plain,(
  ( ! [X445,X447,X443,X449,X444,X446,X442,X448] : (s(bool,'T_0') != s(bool,X446) | s(bool,X447) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X442,X443,X444),sK5(X442,X443,X444),X445,X443,X444)),s('type.closSem.v',sK6(X442,X443,X444)))) | s(bool,X446) != s(bool,X448) | s(bool,'F_0') = s(bool,X448) | s(bool,X447) != s(bool,X449) | s(bool,'T_0') = s(bool,X449) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X444),X443),s('type.sptree.spt'('type.clos_known.val_approx'),X445))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X444),X443),s('type.sptree.spt'('type.clos_known.val_approx'),X442)))) ) | ($spl24 | $spl34 | $spl54)),
  inference(superposition,[],[f1268,f3195])).
fof(f1629,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) | $spl94),
  inference(cnf_transformation,[],[f1629_D])).
fof(f1629_D,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3369,plain,(
  ( ! [X430,X432,X434,X429,X431,X433,X435] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X435,X434,X433)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X435,X434,X433)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X433),X434))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X429,X430,X431),sK5(X429,X430,X431),X432,X430,X431)),s('type.closSem.v',sK6(X429,X430,X431))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X431),X430),s('type.sptree.spt'('type.clos_known.val_approx'),X432))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X431),X430),s('type.sptree.spt'('type.clos_known.val_approx'),X429)))) ) | ($spl24 | $spl34 | $spl58)),
  inference(subsumption_resolution,[],[f3298,f1099])).
fof(f1099,plain,(
  ( ! [X4,X5,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X5,X4,X3)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X5,X4,X3)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X4),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) )),
  inference(subsumption_resolution,[],[f1094,f152])).
fof(f1094,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X5,X4,X3)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X5,X4,X3)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X4),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) )),
  inference(superposition,[],[f141,f151])).
fof(f3298,plain,(
  ( ! [X430,X432,X434,X429,X431,X433,X435] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X433),X434),s('type.sptree.spt'('type.clos_known.val_approx'),X435))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X435,X434,X433)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X435,X434,X433)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X433),X434))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X429,X430,X431),sK5(X429,X430,X431),X432,X430,X431)),s('type.closSem.v',sK6(X429,X430,X431))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X431),X430),s('type.sptree.spt'('type.clos_known.val_approx'),X432))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X431),X430),s('type.sptree.spt'('type.clos_known.val_approx'),X429)))) ) | ($spl24 | $spl34 | $spl58)),
  inference(superposition,[],[f1098,f3195])).
fof(f3367,plain,(
  ( ! [X405,X407,X401,X403,X404,X406,X402] : (s(bool,X406) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X401,X402,X403),sK5(X401,X402,X403),X404,X402,X403)),s('type.closSem.v',sK6(X401,X402,X403)))) | s(bool,'F_0') = s(bool,X405) | s(bool,X406) != s(bool,X407) | s(bool,'T_0') = s(bool,X407) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X404))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X401)))) ) | ($spl24 | $spl34 | $spl40)),
  inference(subsumption_resolution,[],[f3294,f151])).
fof(f3294,plain,(
  ( ! [X405,X407,X401,X403,X404,X406,X402] : (s(bool,'T_0') != s(bool,X405) | s(bool,X406) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X401,X402,X403),sK5(X401,X402,X403),X404,X402,X403)),s('type.closSem.v',sK6(X401,X402,X403)))) | s(bool,'F_0') = s(bool,X405) | s(bool,X406) != s(bool,X407) | s(bool,'T_0') = s(bool,X407) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X404))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X403),X402),s('type.sptree.spt'('type.clos_known.val_approx'),X401)))) ) | ($spl24 | $spl34 | $spl40)),
  inference(superposition,[],[f1067,f3195])).
fof(f3365,plain,(
  ( ! [X377,X372,X374,X376,X373,X375] : (s(bool,X377) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X372,X373,X374),sK5(X372,X373,X374),X375,X373,X374)),s('type.closSem.v',sK6(X372,X373,X374)))) | s(bool,'F_0') = s(bool,X376) | s(bool,'T_0') = s(bool,X377) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X375))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X372)))) ) | ($spl24 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f3289,f151])).
fof(f3289,plain,(
  ( ! [X377,X372,X374,X376,X373,X375] : (s(bool,'T_0') != s(bool,X376) | s(bool,X377) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X372,X373,X374),sK5(X372,X373,X374),X375,X373,X374)),s('type.closSem.v',sK6(X372,X373,X374)))) | s(bool,'F_0') = s(bool,X376) | s(bool,'T_0') = s(bool,X377) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X375))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X374),X373),s('type.sptree.spt'('type.clos_known.val_approx'),X372)))) ) | ($spl24 | $spl34 | $spl38)),
  inference(superposition,[],[f754,f3195])).
fof(f3362,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,'T_0') != s(bool,X323) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X319,X320,X321),sK5(X319,X320,X321),X322,X320,X321)),s('type.closSem.v',sK6(X319,X320,X321))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X321),X320),s('type.sptree.spt'('type.clos_known.val_approx'),X322))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X321),X320),s('type.sptree.spt'('type.clos_known.val_approx'),X319)))) ) | ($spl10 | $spl14 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3278,f236])).
fof(f3278,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,'T_0') != s(bool,X323) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X319,X320,X321),sK5(X319,X320,X321),X322,X320,X321)),s('type.closSem.v',sK6(X319,X320,X321))))) | ~p(s(bool,X323)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X321),X320),s('type.sptree.spt'('type.clos_known.val_approx'),X322))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X321),X320),s('type.sptree.spt'('type.clos_known.val_approx'),X319)))) ) | ($spl14 | $spl24 | $spl34)),
  inference(superposition,[],[f241,f3195])).
fof(f3361,plain,(
  ( ! [X313,X310,X312,X314,X309,X311] : (p(s(bool,X313)) | p(s(bool,X314)) | s(bool,X314) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X309,X310,X311),sK5(X309,X310,X311),X312,X310,X311)),s('type.closSem.v',sK6(X309,X310,X311)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X311),X310),s('type.sptree.spt'('type.clos_known.val_approx'),X312))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X311),X310),s('type.sptree.spt'('type.clos_known.val_approx'),X309)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3276,f156])).
fof(f3276,plain,(
  ( ! [X313,X310,X312,X314,X309,X311] : (~p(s(bool,'T_0')) | p(s(bool,X313)) | p(s(bool,X314)) | s(bool,X314) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X309,X310,X311),sK5(X309,X310,X311),X312,X310,X311)),s('type.closSem.v',sK6(X309,X310,X311)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X311),X310),s('type.sptree.spt'('type.clos_known.val_approx'),X312))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X311),X310),s('type.sptree.spt'('type.clos_known.val_approx'),X309)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f155,f3195])).
fof(f3360,plain,(
  ( ! [X292,X294,X296,X293,X295,X291] : (s(bool,'T_0') != s(bool,X295) | p(s(bool,X296)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X291,X292,X293),sK5(X291,X292,X293),X294,X292,X293)),s('type.closSem.v',sK6(X291,X292,X293))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X293),X292),s('type.sptree.spt'('type.clos_known.val_approx'),X294))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X293),X292),s('type.sptree.spt'('type.clos_known.val_approx'),X291)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3273,f236])).
fof(f3273,plain,(
  ( ! [X292,X294,X296,X293,X295,X291] : (s(bool,'T_0') != s(bool,X295) | p(s(bool,X296)) | ~p(s(bool,X295)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X291,X292,X293),sK5(X291,X292,X293),X294,X292,X293)),s('type.closSem.v',sK6(X291,X292,X293))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X293),X292),s('type.sptree.spt'('type.clos_known.val_approx'),X294))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X293),X292),s('type.sptree.spt'('type.clos_known.val_approx'),X291)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f155,f3195])).
fof(f3359,plain,(
  ( ! [X271,X273,X275,X270,X276,X272,X274] : (p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X274) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X270,X271,X272),sK5(X270,X271,X272),X273,X271,X272)),s('type.closSem.v',sK6(X270,X271,X272)))) | s(bool,X275) != s(bool,X276) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X272),X271),s('type.sptree.spt'('type.clos_known.val_approx'),X273))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X272),X271),s('type.sptree.spt'('type.clos_known.val_approx'),X270)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3270,f156])).
fof(f3270,plain,(
  ( ! [X271,X273,X275,X270,X276,X272,X274] : (~p(s(bool,'T_0')) | p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X274) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X270,X271,X272),sK5(X270,X271,X272),X273,X271,X272)),s('type.closSem.v',sK6(X270,X271,X272)))) | s(bool,X275) != s(bool,X276) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X272),X271),s('type.sptree.spt'('type.clos_known.val_approx'),X273))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X272),X271),s('type.sptree.spt'('type.clos_known.val_approx'),X270)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f115,f3195])).
fof(f3356,plain,(
  ( ! [X243,X246,X244,X242,X248,X247,X245] : (s(bool,'T_0') != s(bool,X246) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X242,X243,X244),sK5(X242,X243,X244),X245,X243,X244)),s('type.closSem.v',sK6(X242,X243,X244))))) | ~p(s(bool,X247)) | s(bool,X247) != s(bool,X248) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X244),X243),s('type.sptree.spt'('type.clos_known.val_approx'),X245))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X244),X243),s('type.sptree.spt'('type.clos_known.val_approx'),X242)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3266,f236])).
fof(f3266,plain,(
  ( ! [X243,X246,X244,X242,X248,X247,X245] : (s(bool,'T_0') != s(bool,X246) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X242,X243,X244),sK5(X242,X243,X244),X245,X243,X244)),s('type.closSem.v',sK6(X242,X243,X244))))) | ~p(s(bool,X247)) | ~p(s(bool,X246)) | s(bool,X247) != s(bool,X248) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X244),X243),s('type.sptree.spt'('type.clos_known.val_approx'),X245))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X244),X243),s('type.sptree.spt'('type.clos_known.val_approx'),X242)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f115,f3195])).
fof(f3353,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X214] : (p(s(bool,X218)) | ~p(s(bool,X219)) | p(s(bool,X220)) | s(bool,X220) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X214,X215,X216),sK5(X214,X215,X216),X217,X215,X216)),s('type.closSem.v',sK6(X214,X215,X216)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X216),X215),s('type.sptree.spt'('type.clos_known.val_approx'),X217))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X216),X215),s('type.sptree.spt'('type.clos_known.val_approx'),X214)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3262,f156])).
fof(f3262,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X214] : (~p(s(bool,'T_0')) | p(s(bool,X218)) | ~p(s(bool,X219)) | p(s(bool,X220)) | s(bool,X220) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X214,X215,X216),sK5(X214,X215,X216),X217,X215,X216)),s('type.closSem.v',sK6(X214,X215,X216)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X216),X215),s('type.sptree.spt'('type.clos_known.val_approx'),X217))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X216),X215),s('type.sptree.spt'('type.clos_known.val_approx'),X214)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f114,f3195])).
fof(f3351,plain,(
  ( ! [X198,X196,X194,X199,X197,X195,X193] : (s(bool,'T_0') != s(bool,X197) | p(s(bool,X198)) | ~p(s(bool,X199)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X193,X194,X195),sK5(X193,X194,X195),X196,X194,X195)),s('type.closSem.v',sK6(X193,X194,X195))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X195),X194),s('type.sptree.spt'('type.clos_known.val_approx'),X196))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X195),X194),s('type.sptree.spt'('type.clos_known.val_approx'),X193)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3259,f236])).
fof(f3259,plain,(
  ( ! [X198,X196,X194,X199,X197,X195,X193] : (s(bool,'T_0') != s(bool,X197) | p(s(bool,X198)) | ~p(s(bool,X197)) | ~p(s(bool,X199)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X193,X194,X195),sK5(X193,X194,X195),X196,X194,X195)),s('type.closSem.v',sK6(X193,X194,X195))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X195),X194),s('type.sptree.spt'('type.clos_known.val_approx'),X196))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X195),X194),s('type.sptree.spt'('type.clos_known.val_approx'),X193)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f114,f3195])).
fof(f3350,plain,(
  ( ! [X185,X182,X180,X184,X183,X181,X179] : (p(s(bool,X183)) | p(s(bool,X184)) | s(bool,X183) != s(bool,X185) | s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X179,X180,X181),sK5(X179,X180,X181),X182,X180,X181)),s('type.closSem.v',sK6(X179,X180,X181)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X181),X180),s('type.sptree.spt'('type.clos_known.val_approx'),X182))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X181),X180),s('type.sptree.spt'('type.clos_known.val_approx'),X179)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3257,f156])).
fof(f3257,plain,(
  ( ! [X185,X182,X180,X184,X183,X181,X179] : (~p(s(bool,'T_0')) | p(s(bool,X183)) | p(s(bool,X184)) | s(bool,X183) != s(bool,X185) | s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X179,X180,X181),sK5(X179,X180,X181),X182,X180,X181)),s('type.closSem.v',sK6(X179,X180,X181)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X181),X180),s('type.sptree.spt'('type.clos_known.val_approx'),X182))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X181),X180),s('type.sptree.spt'('type.clos_known.val_approx'),X179)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f113,f3195])).
fof(f3349,plain,(
  ( ! [X146,X144,X149,X147,X145,X150,X148] : (s(bool,'T_0') != s(bool,X148) | p(s(bool,X149)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X144,X145,X146),sK5(X144,X145,X146),X147,X145,X146)),s('type.closSem.v',sK6(X144,X145,X146))))) | s(bool,X149) != s(bool,X150) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X146),X145),s('type.sptree.spt'('type.clos_known.val_approx'),X147))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X146),X145),s('type.sptree.spt'('type.clos_known.val_approx'),X144)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3252,f236])).
fof(f3252,plain,(
  ( ! [X146,X144,X149,X147,X145,X150,X148] : (s(bool,'T_0') != s(bool,X148) | p(s(bool,X149)) | ~p(s(bool,X148)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X144,X145,X146),sK5(X144,X145,X146),X147,X145,X146)),s('type.closSem.v',sK6(X144,X145,X146))))) | s(bool,X149) != s(bool,X150) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X146),X145),s('type.sptree.spt'('type.clos_known.val_approx'),X147))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X146),X145),s('type.sptree.spt'('type.clos_known.val_approx'),X144)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f113,f3195])).
fof(f3348,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131] : (p(s(bool,X134)) | ~p(s(bool,X135)) | s(bool,X134) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X130,X131,X132),sK5(X130,X131,X132),X133,X131,X132)),s('type.closSem.v',sK6(X130,X131,X132)))) | s(bool,X135) != s(bool,X136) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X132),X131),s('type.sptree.spt'('type.clos_known.val_approx'),X133))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X132),X131),s('type.sptree.spt'('type.clos_known.val_approx'),X130)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3250,f156])).
fof(f3250,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131] : (~p(s(bool,'T_0')) | p(s(bool,X134)) | ~p(s(bool,X135)) | s(bool,X134) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X130,X131,X132),sK5(X130,X131,X132),X133,X131,X132)),s('type.closSem.v',sK6(X130,X131,X132)))) | s(bool,X135) != s(bool,X136) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X132),X131),s('type.sptree.spt'('type.clos_known.val_approx'),X133))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X132),X131),s('type.sptree.spt'('type.clos_known.val_approx'),X130)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f111,f3195])).
fof(f3345,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (s(bool,'T_0') != s(bool,X113) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X109,X110,X111),sK5(X109,X110,X111),X112,X110,X111)),s('type.closSem.v',sK6(X109,X110,X111))))) | ~p(s(bool,X114)) | s(bool,X114) != s(bool,X115) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X111),X110),s('type.sptree.spt'('type.clos_known.val_approx'),X112))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X111),X110),s('type.sptree.spt'('type.clos_known.val_approx'),X109)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3247,f236])).
fof(f3247,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (s(bool,'T_0') != s(bool,X113) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X109,X110,X111),sK5(X109,X110,X111),X112,X110,X111)),s('type.closSem.v',sK6(X109,X110,X111))))) | ~p(s(bool,X114)) | ~p(s(bool,X113)) | s(bool,X114) != s(bool,X115) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X111),X110),s('type.sptree.spt'('type.clos_known.val_approx'),X112))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X111),X110),s('type.sptree.spt'('type.clos_known.val_approx'),X109)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f111,f3195])).
fof(f3342,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X85) != s(bool,X87) | s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X81,X82,X83),sK5(X81,X82,X83),X84,X82,X83)),s('type.closSem.v',sK6(X81,X82,X83)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X84))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X81)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3243,f156])).
fof(f3243,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X85) != s(bool,X87) | s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X81,X82,X83),sK5(X81,X82,X83),X84,X82,X83)),s('type.closSem.v',sK6(X81,X82,X83)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X84))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X82),s('type.sptree.spt'('type.clos_known.val_approx'),X81)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f109,f3195])).
fof(f3340,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X44)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X39,X40,X41),sK5(X39,X40,X41),X42,X40,X41)),s('type.closSem.v',sK6(X39,X40,X41))))) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X39)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3237,f236])).
fof(f3237,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X44)) | ~p(s(bool,X43)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X39,X40,X41),sK5(X39,X40,X41),X42,X40,X41)),s('type.closSem.v',sK6(X39,X40,X41))))) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X40),s('type.sptree.spt'('type.clos_known.val_approx'),X39)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f109,f3195])).
fof(f3212,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))))) ) | $spl160),
  inference(cnf_transformation,[],[f3212_D])).
fof(f3212_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f3211,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))))) ) | ($spl10 | $spl154)),
  inference(subsumption_resolution,[],[f3204,f236])).
fof(f3204,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,X7,X8),sK5(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK6(X6,X7,X8)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl154),
  inference(superposition,[],[f3151,f151])).
fof(f3200,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X17,X18,X19),sK5(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK6(X17,X18,X19)))) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) | $spl158),
  inference(cnf_transformation,[],[f3200_D])).
fof(f3200_D,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X17,X18,X19),sK5(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK6(X17,X18,X19)))) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f3194,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X73,X74,X75),sK5(X73,X74,X75),X76,X74,X75)),s('type.closSem.v',sK6(X73,X74,X75)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X75),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) | p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X78) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X75),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X73)))) ) | $spl34),
  inference(resolution,[],[f3168,f155])).
fof(f3193,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X66,X67,X68),sK5(X66,X67,X68),X69,X67,X68)),s('type.closSem.v',sK6(X66,X67,X68)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X68),X67),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X68),X67),s('type.sptree.spt'('type.clos_known.val_approx'),X66))) | s(bool,X71) != s(bool,X72)) ) | $spl34),
  inference(resolution,[],[f3168,f115])).
fof(f3190,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X45,X46,X47),sK5(X45,X46,X47),X48,X46,X47)),s('type.closSem.v',sK6(X45,X46,X47)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X47),X46),s('type.sptree.spt'('type.clos_known.val_approx'),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X47),X46),s('type.sptree.spt'('type.clos_known.val_approx'),X45)))) ) | $spl34),
  inference(resolution,[],[f3168,f114])).
fof(f3189,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X38,X39,X40),sK5(X38,X39,X40),X41,X39,X40)),s('type.closSem.v',sK6(X38,X39,X40)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X40),X39),s('type.sptree.spt'('type.clos_known.val_approx'),X41))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X40),X39),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) | $spl34),
  inference(resolution,[],[f3168,f113])).
fof(f3188,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X31,X32,X33),sK5(X31,X32,X33),X34,X32,X33)),s('type.closSem.v',sK6(X31,X32,X33)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31))) | s(bool,X36) != s(bool,X37)) ) | $spl34),
  inference(resolution,[],[f3168,f111])).
fof(f3186,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X17,X18,X19),sK5(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK6(X17,X18,X19)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17)))) ) | $spl34),
  inference(resolution,[],[f3168,f109])).
fof(f3180,plain,(
  ( ! [X14,X12,X13,X11,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | ~p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X13,X12,X11),sK5(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK6(X13,X12,X11))))) ) | $spl156),
  inference(cnf_transformation,[],[f3180_D])).
fof(f3180_D,plain,(
  ( ! [X14,X12,X13,X11,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | ~p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X13,X12,X11),sK5(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK6(X13,X12,X11))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f3166,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X67),X68),s('type.sptree.spt'('type.clos_known.val_approx'),X69)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X69,X68,X67),sK5(X69,X68,X67),X70,X68,X67)),s('type.closSem.v',sK6(X69,X68,X67)))) | p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X67),X68),s('type.sptree.spt'('type.clos_known.val_approx'),X70)))) )),
  inference(resolution,[],[f3154,f155])).
fof(f3165,plain,(
  ( ! [X61,X66,X64,X62,X60,X65,X63] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X62,X61,X60),sK5(X62,X61,X60),X63,X61,X60)),s('type.closSem.v',sK6(X62,X61,X60)))) | ~p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63))) | s(bool,X64) != s(bool,X66)) )),
  inference(resolution,[],[f3154,f115])).
fof(f3163,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X47),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X48,X47,X46),sK5(X48,X47,X46),X49,X47,X46)),s('type.closSem.v',sK6(X48,X47,X46)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X47),s('type.sptree.spt'('type.clos_known.val_approx'),X49)))) )),
  inference(resolution,[],[f3154,f114])).
fof(f3161,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X34,X33,X32),sK5(X34,X33,X32),X35,X33,X32)),s('type.closSem.v',sK6(X34,X33,X32)))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) )),
  inference(resolution,[],[f3154,f113])).
fof(f3160,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X27,X26,X25),sK5(X27,X26,X25),X28,X26,X25)),s('type.closSem.v',sK6(X27,X26,X25)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f3154,f111])).
fof(f3158,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X13,X12,X11),sK5(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK6(X13,X12,X11)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14)))) )),
  inference(resolution,[],[f3154,f109])).
fof(f3151,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X7,X5,X4),sK5(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK6(X7,X5,X4)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,X9)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))))) ) | $spl154),
  inference(cnf_transformation,[],[f3151_D])).
fof(f3151_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X7,X5,X4),sK5(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK6(X7,X5,X4)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,X9)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f3145,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63)))) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X63,X61,X60),sK5(X63,X61,X60),X62,X61,X60)),s('type.closSem.v',sK6(X63,X61,X60))))) )),
  inference(resolution,[],[f3135,f155])).
fof(f3144,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X55)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X56,X54,X53),sK5(X56,X54,X53),X55,X54,X53)),s('type.closSem.v',sK6(X56,X54,X53)))) | s(bool,X58) != s(bool,X59)) )),
  inference(resolution,[],[f3135,f115])).
fof(f3141,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X35,X33,X32),sK5(X35,X33,X32),X34,X33,X32)),s('type.closSem.v',sK6(X35,X33,X32))))) )),
  inference(resolution,[],[f3135,f114])).
fof(f3140,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28)))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,X26,X25),sK5(X28,X26,X25),X27,X26,X25)),s('type.closSem.v',sK6(X28,X26,X25))))) )),
  inference(resolution,[],[f3135,f113])).
fof(f3139,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X21,X19,X18),sK5(X21,X19,X18),X20,X19,X18)),s('type.closSem.v',sK6(X21,X19,X18)))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f3135,f111])).
fof(f3137,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7)))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X7,X5,X4),sK5(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK6(X7,X5,X4))))) )),
  inference(resolution,[],[f3135,f109])).
fof(f3119,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X28) != s(bool,X29) | s(bool,X25) != s(bool,X26)) ) | $spl152),
  inference(cnf_transformation,[],[f3119_D])).
fof(f3119_D,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X28) != s(bool,X29) | s(bool,X25) != s(bool,X26)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f3118,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X24)) | s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) != s(bool,X30) | s(bool,X26) != s(bool,X27)) ) | ($spl10 | $spl78)),
  inference(subsumption_resolution,[],[f3112,f236])).
fof(f3112,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X24)) | s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) != s(bool,X30) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X24)) ) | $spl78),
  inference(superposition,[],[f1410,f151])).
fof(f3116,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X9) | p(s(bool,X9)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl150),
  inference(cnf_transformation,[],[f3116_D])).
fof(f3116_D,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X9) | p(s(bool,X9)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f3103,plain,(
  ( ! [X94,X92,X90,X88,X97,X95,X93,X91,X89,X96] : (s(bool,X88) != s(bool,X89) | s(bool,X88) != s(bool,X90) | s(bool,X91) != s(bool,X92) | s(bool,X92) != s(bool,X93) | s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,X93) != s(bool,X94) | s(bool,X90) != s(bool,X91) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X89) != s(bool,X97)) ) | $spl78),
  inference(resolution,[],[f1410,f155])).
fof(f3102,plain,(
  ( ! [X80,X78,X87,X85,X83,X81,X79,X77,X86,X84,X82] : (s(bool,X77) != s(bool,X78) | s(bool,X77) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X81) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,'F_0') = s(bool,X84) | s(bool,X82) != s(bool,X83) | s(bool,X79) != s(bool,X80) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X78) != s(bool,X85) | s(bool,X86) != s(bool,X87)) ) | $spl78),
  inference(resolution,[],[f1410,f115])).
fof(f3099,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X44,X53,X51,X49] : (s(bool,X44) != s(bool,X45) | s(bool,X44) != s(bool,X46) | s(bool,X47) != s(bool,X48) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) != s(bool,X50) | s(bool,X46) != s(bool,X47) | p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X45) != s(bool,X54)) ) | $spl78),
  inference(resolution,[],[f1410,f114])).
fof(f3098,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X42,X40] : (s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X38) != s(bool,X39) | s(bool,X35) != s(bool,X36) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X34) != s(bool,X42)) ) | $spl78),
  inference(resolution,[],[f1410,f113])).
fof(f3097,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32,X22] : (s(bool,X22) != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,X25) != s(bool,X26) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,X27) != s(bool,X28) | s(bool,X24) != s(bool,X25) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X23) != s(bool,X30) | s(bool,X31) != s(bool,X32)) ) | $spl78),
  inference(resolution,[],[f1410,f111])).
fof(f3095,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X1) != s(bool,X9)) ) | $spl78),
  inference(resolution,[],[f1410,f109])).
fof(f3016,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | $spl148),
  inference(cnf_transformation,[],[f3016_D])).
fof(f3016_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3015,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl138)),
  inference(subsumption_resolution,[],[f3008,f508])).
fof(f3008,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl138),
  inference(superposition,[],[f2919,f151])).
fof(f2998,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0))))) ) | $spl146),
  inference(cnf_transformation,[],[f2998_D])).
fof(f2998_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f2997,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0))))) ) | ($spl10 | $spl134)),
  inference(subsumption_resolution,[],[f2988,f236])).
fof(f2988,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl134),
  inference(superposition,[],[f2823,f151])).
fof(f2974,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))))) ) | $spl144),
  inference(cnf_transformation,[],[f2974_D])).
fof(f2974_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f2973,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,X2))) ) | ($spl10 | $spl132)),
  inference(subsumption_resolution,[],[f2967,f236])).
fof(f2967,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X3,sK1,sK0),sK5(X3,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X3,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl132),
  inference(superposition,[],[f2806,f151])).
fof(f2919,plain,(
  ( ! [X30,X28] : (s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,sK1,sK0),sK5(X28,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X28,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | p(s(bool,X30))) ) | $spl138),
  inference(cnf_transformation,[],[f2919_D])).
fof(f2919_D,plain,(
  ( ! [X30,X28] : (s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,sK1,sK0),sK5(X28,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X28,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | p(s(bool,X30))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2959,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X33))) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X33,sK1,sK0),sK5(X33,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X33,sK1,sK0))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f155])).
fof(f2958,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X29,sK1,sK0),sK5(X29,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X29,sK1,sK0)))) | s(bool,X31) != s(bool,X32)) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f115])).
fof(f2955,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X17,sK1,sK0),sK5(X17,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X17,sK1,sK0))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f114])).
fof(f2954,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X13,sK1,sK0),sK5(X13,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X13,sK1,sK0))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f113])).
fof(f2953,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X9,sK1,sK0),sK5(X9,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X9,sK1,sK0)))) | s(bool,X11) != s(bool,X12)) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f111])).
fof(f2951,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X1,sK1,sK0),sK5(X1,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X1,sK1,sK0))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2912,f109])).
fof(f2944,plain,(
  ( ! [X194,X192,X195] : (s(bool,X194) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X192,sK1,sK0),sK5(X192,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X192,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X192))) | s(bool,'T_0') = s(bool,X195) | s(bool,X194) != s(bool,X195)) ) | $spl142),
  inference(cnf_transformation,[],[f2944_D])).
fof(f2944_D,plain,(
  ( ! [X194,X192,X195] : (s(bool,X194) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X192,sK1,sK0),sK5(X192,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X192,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X192))) | s(bool,'T_0') = s(bool,X195) | s(bool,X194) != s(bool,X195)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f2948,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264] : (s(bool,X261) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X259,sK1,sK0),sK5(X259,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X259,sK1,sK0)))) | s(bool,X260) != s(bool,X262) | s(bool,X262) != s(bool,X263) | s(bool,'F_0') = s(bool,X264) | s(bool,X263) != s(bool,X264) | s(bool,X261) != s(bool,X265) | s(bool,'T_0') = s(bool,X265) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X259)))) ) | ($spl2 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f2901,f1067])).
fof(f2901,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264] : (s(bool,'T_0') != s(bool,X260) | s(bool,X261) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X259,sK1,sK0),sK5(X259,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X259,sK1,sK0)))) | s(bool,X260) != s(bool,X262) | s(bool,X262) != s(bool,X263) | s(bool,'F_0') = s(bool,X264) | s(bool,X263) != s(bool,X264) | s(bool,X261) != s(bool,X265) | s(bool,'T_0') = s(bool,X265) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X259)))) ) | ($spl2 | $spl24 | $spl74)),
  inference(superposition,[],[f1399,f2819])).
fof(f2947,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (s(bool,X237) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X235,sK1,sK0),sK5(X235,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X235,sK1,sK0)))) | s(bool,X236) != s(bool,X238) | s(bool,X238) != s(bool,X239) | s(bool,'F_0') = s(bool,X239) | s(bool,X237) != s(bool,X240) | s(bool,'T_0') = s(bool,X240) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X235)))) ) | ($spl2 | $spl24 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f2897,f754])).
fof(f2897,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (s(bool,'T_0') != s(bool,X236) | s(bool,X237) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X235,sK1,sK0),sK5(X235,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X235,sK1,sK0)))) | s(bool,X236) != s(bool,X238) | s(bool,X238) != s(bool,X239) | s(bool,'F_0') = s(bool,X239) | s(bool,X237) != s(bool,X240) | s(bool,'T_0') = s(bool,X240) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X235)))) ) | ($spl2 | $spl24 | $spl70)),
  inference(superposition,[],[f1330,f2819])).
fof(f2946,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,X217) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X215,sK1,sK0),sK5(X215,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X215,sK1,sK0)))) | s(bool,X216) != s(bool,X218) | s(bool,'F_0') = s(bool,X218) | s(bool,X217) != s(bool,X219) | s(bool,'T_0') = s(bool,X219) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X215)))) ) | ($spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f2893,f691])).
fof(f2893,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,'T_0') != s(bool,X216) | s(bool,X217) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X215,sK1,sK0),sK5(X215,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X215,sK1,sK0)))) | s(bool,X216) != s(bool,X218) | s(bool,'F_0') = s(bool,X218) | s(bool,X217) != s(bool,X219) | s(bool,'T_0') = s(bool,X219) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X215)))) ) | ($spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1268,f2819])).
fof(f2917,plain,(
  ( ! [X4] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,sK1,sK0),sK5(X4,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X4,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) ) | $spl136),
  inference(cnf_transformation,[],[f2917_D])).
fof(f2917_D,plain,(
  ( ! [X4] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,sK1,sK0),sK5(X4,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X4,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f2945,plain,(
  ( ! [X210,X208,X211,X209] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X211,X210,X209)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X211,X210,X209)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X209),X210))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X208,sK1,sK0),sK5(X208,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X208,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X208)))) ) | ($spl2 | $spl24 | $spl58)),
  inference(subsumption_resolution,[],[f2891,f1099])).
fof(f2891,plain,(
  ( ! [X210,X208,X211,X209] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X209),X210),s('type.sptree.spt'('type.clos_known.val_approx'),X211))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X211,X210,X209)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X211,X210,X209)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X209),X210))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X208,sK1,sK0),sK5(X208,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X208,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X208)))) ) | ($spl2 | $spl24 | $spl58)),
  inference(superposition,[],[f1098,f2819])).
fof(f2943,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,X194) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X192,sK1,sK0),sK5(X192,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X192,sK1,sK0)))) | s(bool,'F_0') = s(bool,X193) | s(bool,X194) != s(bool,X195) | s(bool,'T_0') = s(bool,X195) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X192)))) ) | ($spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f2887,f151])).
fof(f2887,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,'T_0') != s(bool,X193) | s(bool,X194) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X192,sK1,sK0),sK5(X192,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X192,sK1,sK0)))) | s(bool,'F_0') = s(bool,X193) | s(bool,X194) != s(bool,X195) | s(bool,'T_0') = s(bool,X195) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X192)))) ) | ($spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1067,f2819])).
fof(f2942,plain,(
  ( ! [X180,X178] : (s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X178,sK1,sK0),sK5(X178,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X178,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X178))) | s(bool,'T_0') = s(bool,X180)) ) | $spl140),
  inference(cnf_transformation,[],[f2942_D])).
fof(f2942_D,plain,(
  ( ! [X180,X178] : (s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X178,sK1,sK0),sK5(X178,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X178,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X178))) | s(bool,'T_0') = s(bool,X180)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2941,plain,(
  ( ! [X180,X178,X179] : (s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X178,sK1,sK0),sK5(X178,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X178,sK1,sK0)))) | s(bool,'F_0') = s(bool,X179) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X178)))) ) | ($spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f2882,f151])).
fof(f2882,plain,(
  ( ! [X180,X178,X179] : (s(bool,'T_0') != s(bool,X179) | s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X178,sK1,sK0),sK5(X178,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X178,sK1,sK0)))) | s(bool,'F_0') = s(bool,X179) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X178)))) ) | ($spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f754,f2819])).
fof(f2938,plain,(
  ( ! [X158,X159] : (s(bool,'T_0') != s(bool,X159) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X158,sK1,sK0),sK5(X158,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X158,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X158)))) ) | ($spl2 | $spl10 | $spl14 | $spl24)),
  inference(subsumption_resolution,[],[f2871,f236])).
fof(f2871,plain,(
  ( ! [X158,X159] : (s(bool,'T_0') != s(bool,X159) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X158,sK1,sK0),sK5(X158,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X158,sK1,sK0))))) | ~p(s(bool,X159)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X158)))) ) | ($spl2 | $spl14 | $spl24)),
  inference(superposition,[],[f241,f2819])).
fof(f2937,plain,(
  ( ! [X156,X154,X155] : (p(s(bool,X155)) | p(s(bool,X156)) | s(bool,X156) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X154,sK1,sK0),sK5(X154,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X154,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X154)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2869,f156])).
fof(f2869,plain,(
  ( ! [X156,X154,X155] : (~p(s(bool,'T_0')) | p(s(bool,X155)) | p(s(bool,X156)) | s(bool,X156) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X154,sK1,sK0),sK5(X154,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X154,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X154)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f155,f2819])).
fof(f2936,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X147)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X145,sK1,sK0),sK5(X145,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X145,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X145)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2866,f236])).
fof(f2866,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X147)) | ~p(s(bool,X146)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X145,sK1,sK0),sK5(X145,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X145,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X145)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f155,f2819])).
fof(f2935,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X136,sK1,sK0),sK5(X136,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X136,sK1,sK0)))) | s(bool,X138) != s(bool,X139) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X136)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2863,f156])).
fof(f2863,plain,(
  ( ! [X138,X136,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X136,sK1,sK0),sK5(X136,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X136,sK1,sK0)))) | s(bool,X138) != s(bool,X139) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X136)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f115,f2819])).
fof(f2932,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') != s(bool,X121) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X120,sK1,sK0),sK5(X120,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X120,sK1,sK0))))) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,X123) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X120)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2859,f236])).
fof(f2859,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') != s(bool,X121) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X120,sK1,sK0),sK5(X120,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X120,sK1,sK0))))) | ~p(s(bool,X122)) | ~p(s(bool,X121)) | s(bool,X122) != s(bool,X123) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X120)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f115,f2819])).
fof(f2929,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X105)) | ~p(s(bool,X106)) | p(s(bool,X107)) | s(bool,X107) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X104,sK1,sK0),sK5(X104,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X104,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X104)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2855,f156])).
fof(f2855,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X105)) | ~p(s(bool,X106)) | p(s(bool,X107)) | s(bool,X107) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X104,sK1,sK0),sK5(X104,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X104,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X104)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f114,f2819])).
fof(f2927,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X92,sK1,sK0),sK5(X92,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X92,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X92)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2852,f236])).
fof(f2852,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,X93)) | ~p(s(bool,X95)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X92,sK1,sK0),sK5(X92,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X92,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X92)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f114,f2819])).
fof(f2926,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X85) != s(bool,X87) | s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X84,sK1,sK0),sK5(X84,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X84,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X84)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2850,f156])).
fof(f2850,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X85) != s(bool,X87) | s(bool,X86) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X84,sK1,sK0),sK5(X84,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X84,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X84)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f113,f2819])).
fof(f2925,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X65) | p(s(bool,X66)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X64,sK1,sK0),sK5(X64,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X64,sK1,sK0))))) | s(bool,X66) != s(bool,X67) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X64)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2845,f236])).
fof(f2845,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X65) | p(s(bool,X66)) | ~p(s(bool,X65)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X64,sK1,sK0),sK5(X64,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X64,sK1,sK0))))) | s(bool,X66) != s(bool,X67) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X64)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f113,f2819])).
fof(f2924,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X56,sK1,sK0),sK5(X56,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X56,sK1,sK0)))) | s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X56)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2843,f156])).
fof(f2843,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X56,sK1,sK0),sK5(X56,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X56,sK1,sK0)))) | s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X56)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f111,f2819])).
fof(f2921,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') != s(bool,X45) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X44,sK1,sK0),sK5(X44,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X44,sK1,sK0))))) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X44)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2840,f236])).
fof(f2840,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') != s(bool,X45) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X44,sK1,sK0),sK5(X44,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X44,sK1,sK0))))) | ~p(s(bool,X46)) | ~p(s(bool,X45)) | s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X44)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f111,f2819])).
fof(f2918,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,sK1,sK0),sK5(X28,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X28,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X28)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2836,f156])).
fof(f2836,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X28,sK1,sK0),sK5(X28,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X28,sK1,sK0)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X28)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f109,f2819])).
fof(f2916,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,sK1,sK0),sK5(X4,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X4,sK1,sK0))))) | s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2830,f236])).
fof(f2830,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X4,sK1,sK0),sK5(X4,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X4,sK1,sK0))))) | s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f109,f2819])).
fof(f2823,plain,(
  ( ! [X6,X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,sK1,sK0),sK5(X6,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X6,sK1,sK0))))) ) | $spl134),
  inference(cnf_transformation,[],[f2823_D])).
fof(f2823_D,plain,(
  ( ! [X6,X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,sK1,sK0),sK5(X6,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X6,sK1,sK0))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2818,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X38,sK1,sK0),sK5(X38,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X38,sK1,sK0)))) | p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f155])).
fof(f2817,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X34,sK1,sK0),sK5(X34,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X34,sK1,sK0)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,X36) != s(bool,X37)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f115])).
fof(f2814,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X22,sK1,sK0),sK5(X22,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X22,sK1,sK0)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f114])).
fof(f2813,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X18,sK1,sK0),sK5(X18,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X18,sK1,sK0)))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f113])).
fof(f2812,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X14,sK1,sK0),sK5(X14,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X14,sK1,sK0)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | s(bool,X16) != s(bool,X17)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f111])).
fof(f2810,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X6,sK1,sK0),sK5(X6,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X6,sK1,sK0)))) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2802,f109])).
fof(f2806,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X1,sK1,sK0),sK5(X1,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X1,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X3))) ) | $spl132),
  inference(cnf_transformation,[],[f2806_D])).
fof(f2806_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X1,sK1,sK0),sK5(X1,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X1,sK1,sK0)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X3))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f2801,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X33)))) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X33,sK1,sK0),sK5(X33,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X33,sK1,sK0))))) ) | $spl2),
  inference(resolution,[],[f2776,f155])).
fof(f2800,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X29,sK1,sK0),sK5(X29,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X29,sK1,sK0)))) | s(bool,X31) != s(bool,X32)) ) | $spl2),
  inference(resolution,[],[f2776,f115])).
fof(f2797,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X17)))) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X17,sK1,sK0),sK5(X17,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X17,sK1,sK0))))) ) | $spl2),
  inference(resolution,[],[f2776,f114])).
fof(f2796,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X13,sK1,sK0),sK5(X13,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X13,sK1,sK0))))) ) | $spl2),
  inference(resolution,[],[f2776,f113])).
fof(f2795,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X9,sK1,sK0),sK5(X9,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X9,sK1,sK0)))) | s(bool,X11) != s(bool,X12)) ) | $spl2),
  inference(resolution,[],[f2776,f111])).
fof(f2793,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(X1,sK1,sK0),sK5(X1,sK1,sK0),sK2,sK1,sK0)),s('type.closSem.v',sK6(X1,sK1,sK0))))) ) | $spl2),
  inference(resolution,[],[f2776,f109])).
fof(f2790,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) ) | $spl130),
  inference(cnf_transformation,[],[f2790_D])).
fof(f2790_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2786,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X71))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X72),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X73),X74))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X73),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X75))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X71,X72,X75,X74,X73)),s('type.closSem.v',X71)))) )),
  inference(resolution,[],[f1184,f155])).
fof(f1184,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X6,X5,X4)),s('type.closSem.v',X7)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X4),X5))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) )),
  inference(subsumption_resolution,[],[f1173,f156])).
fof(f1173,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X6,X5,X4)),s('type.closSem.v',X7)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X4),X5))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) )),
  inference(superposition,[],[f140,f151])).
fof(f2785,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X63))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X64),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X65),X66))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X65),X66),s('type.sptree.spt'('type.clos_known.val_approx'),X67))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X63,X64,X67,X66,X65)),s('type.closSem.v',X63))) | s(bool,X69) != s(bool,X70)) )),
  inference(resolution,[],[f1184,f115])).
fof(f2782,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X39))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X40),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X41),X42))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X42),s('type.sptree.spt'('type.clos_known.val_approx'),X43))) | p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X39,X40,X43,X42,X41)),s('type.closSem.v',X39)))) )),
  inference(resolution,[],[f1184,f114])).
fof(f2781,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X31))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X32),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X34),s('type.sptree.spt'('type.clos_known.val_approx'),X35))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X31,X32,X35,X34,X33)),s('type.closSem.v',X31)))) )),
  inference(resolution,[],[f1184,f113])).
fof(f2780,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X23))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X24),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X25),X26))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X23,X24,X27,X26,X25)),s('type.closSem.v',X23))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f1184,f111])).
fof(f2778,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7)))) )),
  inference(resolution,[],[f1184,f109])).
fof(f1410,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,X8)) | s(bool,X1) != s(bool,X8) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f1410_D])).
fof(f1410_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,X8)) | s(bool,X1) != s(bool,X8) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2701,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X86) | s(bool,X85) != s(bool,X86) | s(bool,X82) != s(bool,X83) | p(s(bool,X87)) | p(s(bool,X88)) | s(bool,X81) != s(bool,X88)) ) | $spl74),
  inference(resolution,[],[f1341,f155])).
fof(f2700,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X70) != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X73) != s(bool,X74) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) != s(bool,X76) | s(bool,X72) != s(bool,X73) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X71) != s(bool,X77) | s(bool,X78) != s(bool,X79)) ) | $spl74),
  inference(resolution,[],[f1341,f115])).
fof(f2697,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) != s(bool,X46) | s(bool,X42) != s(bool,X43) | p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X41) != s(bool,X49)) ) | $spl74),
  inference(resolution,[],[f1341,f114])).
fof(f2696,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) != s(bool,X36) | s(bool,X32) != s(bool,X33) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X31) != s(bool,X38)) ) | $spl74),
  inference(resolution,[],[f1341,f113])).
fof(f2695,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X21) != s(bool,X27) | s(bool,X28) != s(bool,X29)) ) | $spl74),
  inference(resolution,[],[f1341,f111])).
fof(f2693,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X1) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f1341,f109])).
fof(f2576,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | $spl128),
  inference(cnf_transformation,[],[f2576_D])).
fof(f2576_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f533,plain,(
  ( ! [X4,X7] : (s(bool,X4) != s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) | $spl30),
  inference(cnf_transformation,[],[f533_D])).
fof(f533_D,plain,(
  ( ! [X4,X7] : (s(bool,X4) != s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2668,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl24 | $spl88)),
  inference(subsumption_resolution,[],[f2665,f691])).
fof(f2665,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | $spl88),
  inference(superposition,[],[f1613,f151])).
fof(f2573,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | $spl127),
  inference(cnf_transformation,[],[f2573_D])).
fof(f2573_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f2667,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl48 | $spl88)),
  inference(subsumption_resolution,[],[f2663,f802])).
fof(f2663,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl88),
  inference(superposition,[],[f1613,f151])).
fof(f1341,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl74),
  inference(cnf_transformation,[],[f1341_D])).
fof(f1341_D,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2646,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,X72) != s(bool,X74) | s(bool,X75) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X74) != s(bool,X75) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X73) != s(bool,X79)) ) | $spl70),
  inference(resolution,[],[f1279,f155])).
fof(f2645,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl70),
  inference(resolution,[],[f1279,f115])).
fof(f2642,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl70),
  inference(resolution,[],[f1279,f114])).
fof(f2641,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl70),
  inference(resolution,[],[f1279,f113])).
fof(f2640,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl70),
  inference(resolution,[],[f1279,f111])).
fof(f2638,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl70),
  inference(resolution,[],[f1279,f109])).
fof(f1249,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | ~p(s(bool,X17)) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X13,X11,X12,X14,X15))))) ) | $spl68),
  inference(cnf_transformation,[],[f1249_D])).
fof(f1249_D,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | ~p(s(bool,X17)) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X13,X11,X12,X14,X15))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2594,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X75),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X77,X75,X76,X78,X79)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X77))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X75),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X76)))) )),
  inference(resolution,[],[f139,f155])).
fof(f2593,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X67),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X69,X67,X68,X70,X71)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X69))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X67),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f139,f115])).
fof(f2591,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X51),s('type.sptree.spt'('type.clos_known.val_approx'),X52))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X53,X51,X52,X54,X55)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X53))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X51),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) )),
  inference(resolution,[],[f139,f114])).
fof(f2589,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X35),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X37,X35,X36,X38,X39)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X37))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X35),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) )),
  inference(resolution,[],[f139,f113])).
fof(f2588,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X27),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X29,X27,X28,X30,X31)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X29))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X27),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f139,f111])).
fof(f2586,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X13,X11,X12,X14,X15)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12)))) )),
  inference(resolution,[],[f139,f109])).
fof(f2583,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | $spl86),
  inference(subsumption_resolution,[],[f2580,f151])).
fof(f2580,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | $spl86),
  inference(superposition,[],[f1611,f151])).
fof(f2582,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl86),
  inference(duplicate_literal_removal,[],[f2578])).
fof(f2578,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl86),
  inference(superposition,[],[f1611,f151])).
fof(f2574,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl24 | $spl84)),
  inference(subsumption_resolution,[],[f2569,f340])).
fof(f2569,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | $spl84),
  inference(superposition,[],[f1474,f151])).
fof(f2571,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | p(s(bool,X1))) ) | ($spl10 | $spl84)),
  inference(subsumption_resolution,[],[f2567,f236])).
fof(f2567,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl84),
  inference(superposition,[],[f1474,f151])).
fof(f1279,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) | $spl70),
  inference(cnf_transformation,[],[f1279_D])).
fof(f1279_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2553,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,X64) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X67) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X65) != s(bool,X70)) ) | $spl54),
  inference(resolution,[],[f1076,f155])).
fof(f2552,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl54),
  inference(resolution,[],[f1076,f115])).
fof(f2549,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl54),
  inference(resolution,[],[f1076,f114])).
fof(f2548,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl54),
  inference(resolution,[],[f1076,f113])).
fof(f2547,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f1076,f111])).
fof(f2545,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl54),
  inference(resolution,[],[f1076,f109])).
fof(f1179,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | ~p(s(bool,X17)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X11,X12,X13,X14,X15)),s('type.closSem.v',X11))))) ) | $spl66),
  inference(cnf_transformation,[],[f1179_D])).
fof(f1179_D,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | ~p(s(bool,X17)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X11,X12,X13,X14,X15)),s('type.closSem.v',X11))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2517,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X75,X76,X77,X78,X79)),s('type.closSem.v',X75)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X75))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X76),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) )),
  inference(resolution,[],[f140,f155])).
fof(f2516,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X67,X68,X69,X70,X71)),s('type.closSem.v',X67)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X67))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X68),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f140,f115])).
fof(f2514,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X51,X52,X53,X54,X55)),s('type.closSem.v',X51)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X51))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X52),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X53)))) )),
  inference(resolution,[],[f140,f114])).
fof(f2512,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X35,X36,X37,X38,X39)),s('type.closSem.v',X35)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X35))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X36),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) )),
  inference(resolution,[],[f140,f113])).
fof(f2511,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X27,X28,X29,X30,X31)),s('type.closSem.v',X27)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X27))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X28),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f140,f111])).
fof(f2509,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X11,X12,X13,X14,X15)),s('type.closSem.v',X11)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) )),
  inference(resolution,[],[f140,f109])).
fof(f1175,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11))))) ) | $spl64),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2505,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X71))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X72),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X73),X74))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X73),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X75)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X71,X72,X75,X74,X73)),s('type.closSem.v',X71)))) )),
  inference(resolution,[],[f140,f155])).
fof(f2504,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X63))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X64),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X65),X66))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X65),X66),s('type.sptree.spt'('type.clos_known.val_approx'),X67)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X63,X64,X67,X66,X65)),s('type.closSem.v',X63))) | s(bool,X69) != s(bool,X70)) )),
  inference(resolution,[],[f140,f115])).
fof(f2501,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X39))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X40),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X41),X42))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X42),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X39,X40,X43,X42,X41)),s('type.closSem.v',X39)))) )),
  inference(resolution,[],[f140,f114])).
fof(f2500,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X31))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X32),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X34),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X31,X32,X35,X34,X33)),s('type.closSem.v',X31)))) )),
  inference(resolution,[],[f140,f113])).
fof(f2499,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X23))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X24),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X25),X26))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X23,X24,X27,X26,X25)),s('type.closSem.v',X23))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f140,f111])).
fof(f2497,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7)))) )),
  inference(resolution,[],[f140,f109])).
fof(f1076,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl54),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2476,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl40),
  inference(resolution,[],[f765,f155])).
fof(f2475,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f765,f115])).
fof(f2472,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f765,f114])).
fof(f2471,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f765,f113])).
fof(f2470,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f765,f111])).
fof(f2468,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f765,f109])).
fof(f1098,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X4))) ) | $spl58),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X4))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2447,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X48,X49,X50)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X48,X49,X50)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X50),X49))))) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X50),X49),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) )),
  inference(resolution,[],[f141,f155])).
fof(f2446,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X42,X43,X44)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X42,X43,X44)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X44),X43))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X44),X43),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f141,f115])).
fof(f2443,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X24,X25,X26)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X24,X25,X26)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X26),X25))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X26),X25),s('type.sptree.spt'('type.clos_known.val_approx'),X24)))) )),
  inference(resolution,[],[f141,f114])).
fof(f2442,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X18,X19,X20)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X18,X19,X20)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X20),X19))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) )),
  inference(resolution,[],[f141,f113])).
fof(f2441,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X12,X13,X14)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X12,X13,X14)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X14),X13))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X13),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f141,f111])).
fof(f2439,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0)))) )),
  inference(resolution,[],[f141,f109])).
fof(f765,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2414,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl38),
  inference(resolution,[],[f714,f155])).
fof(f2413,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f714,f115])).
fof(f2410,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f714,f114])).
fof(f2409,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f714,f113])).
fof(f2408,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f714,f111])).
fof(f2406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f714,f109])).
fof(f1033,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK6(X2,X1,X0)))) | ~p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) ) | $spl52),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK6(X2,X1,X0)))) | ~p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2390,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X57),s('type.sptree.spt'('type.clos_known.val_approx'),X58)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X59))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X58,X57,X56)),s('type.sptree.spt'('type.clos_known.val_approx'),X58))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X59),s('type.closSem.v',sK6(X58,X57,X56))))) )),
  inference(resolution,[],[f142,f155])).
fof(f2389,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X49),X50),s('type.sptree.spt'('type.clos_known.val_approx'),X51)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X52))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X51,X50,X49)),s('type.sptree.spt'('type.clos_known.val_approx'),X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X52),s('type.closSem.v',sK6(X51,X50,X49)))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f142,f115])).
fof(f2387,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X38))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X37,X36,X35)),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X38),s('type.closSem.v',sK6(X37,X36,X35))))) )),
  inference(resolution,[],[f142,f114])).
fof(f2385,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X21),X22),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X24))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X23,X22,X21)),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X24),s('type.closSem.v',sK6(X23,X22,X21))))) )),
  inference(resolution,[],[f142,f113])).
fof(f2384,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X15),s('type.sptree.spt'('type.clos_known.val_approx'),X16)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X17))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X16,X15,X14)),s('type.sptree.spt'('type.clos_known.val_approx'),X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X17),s('type.closSem.v',sK6(X16,X15,X14)))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f142,f111])).
fof(f2382,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK6(X2,X1,X0))))) )),
  inference(resolution,[],[f142,f109])).
fof(f1030,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X5)) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK6(X1,X2,X3)))))) ) | $spl50),
  inference(cnf_transformation,[],[f1030_D])).
fof(f1030_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X5)) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK6(X1,X2,X3)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2380,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X56),s('type.closSem.v',sK6(X57,X58,X59))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X56))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X57,X58,X59)),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X59),X58),s('type.sptree.spt'('type.clos_known.val_approx'),X57)))) )),
  inference(resolution,[],[f142,f155])).
fof(f2379,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X49),s('type.closSem.v',sK6(X50,X51,X52))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X49))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X50,X51,X52)),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X52),X51),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f142,f115])).
fof(f2376,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X28),s('type.closSem.v',sK6(X29,X30,X31))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X28))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X29,X30,X31)),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) )),
  inference(resolution,[],[f142,f114])).
fof(f2375,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X21),s('type.closSem.v',sK6(X22,X23,X24))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X21))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X22,X23,X24)),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X24),X23),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) )),
  inference(resolution,[],[f142,f113])).
fof(f2374,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X14),s('type.closSem.v',sK6(X15,X16,X17))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X14))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X15,X16,X17)),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X17),X16),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f142,f111])).
fof(f2372,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK6(X1,X2,X3))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) )),
  inference(resolution,[],[f142,f109])).
fof(f2371,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl124),
  inference(cnf_transformation,[],[f2371_D])).
fof(f2371_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2370,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl112)),
  inference(subsumption_resolution,[],[f2362,f151])).
fof(f2362,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl112)),
  inference(superposition,[],[f1645,f160])).
fof(f2369,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | $spl112),
  inference(subsumption_resolution,[],[f2360,f1645])).
fof(f2360,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl112),
  inference(superposition,[],[f1645,f151])).
fof(f2368,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl112),
  inference(subsumption_resolution,[],[f2359,f1645])).
fof(f2359,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl112),
  inference(superposition,[],[f1645,f151])).
fof(f2367,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(subsumption_resolution,[],[f2358,f151])).
fof(f2358,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(superposition,[],[f1645,f151])).
fof(f2365,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) | $spl122),
  inference(cnf_transformation,[],[f2365_D])).
fof(f2365_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2364,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(duplicate_literal_removal,[],[f2352])).
fof(f2352,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(superposition,[],[f1645,f151])).
fof(f1646,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X13)) ) | $spl114),
  inference(cnf_transformation,[],[f1646_D])).
fof(f1646_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1645,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4)) ) | $spl112),
  inference(cnf_transformation,[],[f1645_D])).
fof(f1645_D,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2348,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | s(bool,X153) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X153) != s(bool,X154) | s(bool,'F_0') = s(bool,X154) | s(bool,X152) != s(bool,X155) | s(bool,'T_0') = s(bool,X155)) ) | ($spl0 | $spl24 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f2303,f754])).
fof(f2303,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,'F_0') != s(bool,X151) | s(bool,X151) != s(bool,X152) | s(bool,X153) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X153) != s(bool,X154) | s(bool,'F_0') = s(bool,X154) | s(bool,X152) != s(bool,X155) | s(bool,'T_0') = s(bool,X155)) ) | ($spl0 | $spl24 | $spl70)),
  inference(superposition,[],[f1330,f160])).
fof(f1649,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X8)) ) | $spl116),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2347,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,X139) != s(bool,X140) | s(bool,X141) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X141) | s(bool,X140) != s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | ($spl0 | $spl24 | $spl38 | $spl54)),
  inference(subsumption_resolution,[],[f2300,f754])).
fof(f2300,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'F_0') != s(bool,X139) | s(bool,X139) != s(bool,X140) | s(bool,X141) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X141) | s(bool,X140) != s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | ($spl0 | $spl24 | $spl54)),
  inference(superposition,[],[f1268,f160])).
fof(f714,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f714_D])).
fof(f714_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2219,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f552,f155])).
fof(f2218,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f552,f115])).
fof(f2215,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl14),
  inference(resolution,[],[f552,f114])).
fof(f2214,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f552,f113])).
fof(f2213,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f552,f111])).
fof(f2211,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f552,f109])).
fof(f2204,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) != s(bool,X126)) ) | $spl120),
  inference(cnf_transformation,[],[f2204_D])).
fof(f2204_D,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) != s(bool,X126)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2206,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X145) != s(bool,X147) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,X148) | s(bool,X146) != s(bool,X149) | s(bool,'T_0') = s(bool,X149)) ) | ($spl2 | $spl24 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f2160,f754])).
fof(f2160,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X145) | s(bool,X146) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X145) != s(bool,X147) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,X148) | s(bool,X146) != s(bool,X149) | s(bool,'T_0') = s(bool,X149)) ) | ($spl2 | $spl24 | $spl70)),
  inference(superposition,[],[f1330,f161])).
fof(f2205,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X135) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,X136) | s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f2157,f691])).
fof(f2157,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X134) | s(bool,X135) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,X136) | s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1268,f161])).
fof(f2203,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,X126)) ) | ($spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f2153,f151])).
fof(f2153,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,X126)) ) | ($spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1067,f161])).
fof(f2202,plain,(
  ( ! [X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X119)) ) | $spl118),
  inference(cnf_transformation,[],[f2202_D])).
fof(f2202_D,plain,(
  ( ! [X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X119)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f2201,plain,(
  ( ! [X118,X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X119)) ) | ($spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f2149,f151])).
fof(f2149,plain,(
  ( ! [X118,X119] : (s(bool,'T_0') != s(bool,X118) | s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X119)) ) | ($spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f754,f161])).
fof(f2095,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f2087])).
fof(f2087,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(superposition,[],[f533,f151])).
fof(f2096,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(subsumption_resolution,[],[f2086,f533])).
fof(f2086,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f533,f151])).
fof(f674,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f674_D])).
fof(f674_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2078,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f548,f155])).
fof(f2077,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f548,f115])).
fof(f2074,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl34),
  inference(resolution,[],[f548,f114])).
fof(f2073,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f548,f113])).
fof(f2072,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f548,f111])).
fof(f2070,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f548,f109])).
fof(f164,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl4),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f166,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl7),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f2069,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
fof(f107,plain,(
  s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f91])).
fof(f91,plain,(
  ? [X0,X1,X2,X3] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.closSem.state'(X0),X1))),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.closSem.state'(X0),X1))),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.closSem.state'(X0),X1))),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.closSem.state'(X6),X7))),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.closSem.state'(X6),X7))),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.clos_knownProof.state_globals_approx_clock_fupd')).
fof(f160,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl0),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f162,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl3),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2068,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
fof(f340,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | ~p(s(bool,X18))) ) | $spl24),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | ~p(s(bool,X18))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2053,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f152,f155])).
fof(f2052,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f152,f115])).
fof(f2050,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f152,f114])).
fof(f2048,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f152,f113])).
fof(f2047,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f152,f111])).
fof(f2045,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f152,f109])).
fof(f1986,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f1818,f154])).
fof(f1818,plain,(
  ~p(s(bool,'T_0')) | $spl32),
  inference(backward_demodulation,[],[f535,f152])).
fof(f1978,plain,(
  $false | ($spl7 | $spl32)),
  inference(trivial_inequality_removal,[],[f1970])).
fof(f1970,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl7 | $spl32)),
  inference(superposition,[],[f178,f535])).
fof(f178,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(subsumption_resolution,[],[f174,f166])).
fof(f174,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl7),
  inference(superposition,[],[f166,f151])).
fof(f1809,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f152,f155])).
fof(f1808,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f152,f115])).
fof(f1806,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f152,f114])).
fof(f1804,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f152,f113])).
fof(f1803,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f152,f111])).
fof(f1801,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f152,f109])).
fof(f159,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl1),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1055,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl46),
  inference(duplicate_literal_removal,[],[f1048])).
fof(f1048,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl46),
  inference(superposition,[],[f801,f151])).
fof(f1057,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl46),
  inference(subsumption_resolution,[],[f1052,f151])).
fof(f1052,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl46),
  inference(superposition,[],[f801,f151])).
fof(f1214,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl48 | $spl60)),
  inference(subsumption_resolution,[],[f1207,f802])).
fof(f1207,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl60),
  inference(superposition,[],[f1128,f151])).
fof(f1216,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | ($spl24 | $spl60)),
  inference(subsumption_resolution,[],[f1211,f691])).
fof(f1211,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl60),
  inference(superposition,[],[f1128,f151])).
fof(f1747,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f1656,f154])).
fof(f1656,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f1655,f152])).
fof(f1655,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f160,f164])).
fof(f1631,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl97),
  inference(cnf_transformation,[],[f1631_D])).
fof(f1631_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f1653,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f501,f340])).
fof(f501,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f340,f151])).
fof(f1652,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f558,f236])).
fof(f558,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f236,f151])).
fof(f1650,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(subsumption_resolution,[],[f773,f755])).
fof(f755,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X56) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(resolution,[],[f714,f152])).
fof(f773,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(superposition,[],[f755,f151])).
fof(f777,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl4 | $spl38)),
  inference(superposition,[],[f755,f164])).
fof(f1648,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | $spl48),
  inference(subsumption_resolution,[],[f993,f802])).
fof(f993,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(superposition,[],[f802,f151])).
fof(f1056,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f1055,f164])).
fof(f1647,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl40),
  inference(subsumption_resolution,[],[f1103,f1068])).
fof(f1068,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') != s(bool,X66) | s(bool,X66) != s(bool,X67) | s(bool,'F_0') = s(bool,X68) | s(bool,X67) != s(bool,X68)) ) | $spl40),
  inference(resolution,[],[f765,f152])).
fof(f1103,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl40),
  inference(superposition,[],[f1068,f151])).
fof(f1107,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13)) ) | ($spl4 | $spl40)),
  inference(superposition,[],[f1068,f164])).
fof(f1644,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5)) ) | $spl62),
  inference(subsumption_resolution,[],[f1188,f1129])).
fof(f1188,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl62),
  inference(superposition,[],[f1129,f151])).
fof(f1215,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl4 | $spl48 | $spl60)),
  inference(forward_demodulation,[],[f1214,f164])).
fof(f1642,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(subsumption_resolution,[],[f1287,f1269])).
fof(f1287,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(superposition,[],[f1269,f151])).
fof(f1641,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X16) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) | $spl108),
  inference(cnf_transformation,[],[f1641_D])).
fof(f1641_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X16) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1291,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,X16) != s(bool,X17)) ) | ($spl4 | $spl54)),
  inference(superposition,[],[f1269,f164])).
fof(f1640,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) | $spl106),
  inference(cnf_transformation,[],[f1640_D])).
fof(f1640_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1639,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl70),
  inference(subsumption_resolution,[],[f1349,f1331])).
fof(f1331,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,'F_0') != s(bool,X86) | s(bool,X86) != s(bool,X87) | s(bool,X88) != s(bool,X89) | s(bool,X89) != s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,X87) != s(bool,X88)) ) | $spl70),
  inference(resolution,[],[f1279,f152])).
fof(f1349,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl70),
  inference(superposition,[],[f1331,f151])).
fof(f1638,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23)) ) | $spl104),
  inference(cnf_transformation,[],[f1638_D])).
fof(f1638_D,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1353,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21)) ) | ($spl4 | $spl70)),
  inference(superposition,[],[f1331,f164])).
fof(f1637,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl102),
  inference(cnf_transformation,[],[f1637_D])).
fof(f1637_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1636,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(subsumption_resolution,[],[f1418,f1400])).
fof(f1400,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,'F_0') != s(bool,X96) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,X99) | s(bool,X99) != s(bool,X100) | s(bool,'F_0') = s(bool,X101) | s(bool,X100) != s(bool,X101) | s(bool,X97) != s(bool,X98)) ) | $spl74),
  inference(resolution,[],[f1341,f152])).
fof(f1418,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f1400,f151])).
fof(f1635,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27)) ) | $spl100),
  inference(cnf_transformation,[],[f1635_D])).
fof(f1635_D,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1422,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X25) != s(bool,X26) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X28) | s(bool,X27) != s(bool,X28) | s(bool,X24) != s(bool,X25)) ) | ($spl4 | $spl74)),
  inference(superposition,[],[f1400,f164])).
fof(f1634,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | $spl98),
  inference(cnf_transformation,[],[f1634_D])).
fof(f1634_D,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1632,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') != s(bool,'F_0') | ($spl4 | $spl82)),
  inference(forward_demodulation,[],[f1445,f121])).
fof(f1445,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))))))) | ($spl4 | $spl82)),
  inference(superposition,[],[f1439,f164])).
fof(f1628,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) | $spl82),
  inference(subsumption_resolution,[],[f1446,f1439])).
fof(f1446,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) ) | $spl82),
  inference(superposition,[],[f1439,f151])).
fof(f1627,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1626,f154])).
fof(f1626,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1625,f164])).
fof(f1625,plain,(
  ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1624,f334])).
fof(f334,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | $spl6),
  inference(subsumption_resolution,[],[f292,f152])).
fof(f292,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | $spl6),
  inference(superposition,[],[f141,f165])).
fof(f165,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl6),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1624,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1623,f121])).
fof(f1623,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f1622])).
fof(f1622,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f1473,f139])).
fof(f1473,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1472,f152])).
fof(f1472,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1459,f165])).
fof(f1459,plain,(
  p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(sK2,sK1,sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f142])).
fof(f1458,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) | ($spl4 | $spl6)),
  inference(equality_resolution,[],[f1455])).
fof(f1455,plain,(
  ( ! [X0] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(sK2,sK1,sK0)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X0,sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',X0))))) ) | ($spl4 | $spl6)),
  inference(superposition,[],[f620,f334])).
fof(f620,plain,(
  ( ! [X2,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK1))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X2,X3,sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f619,f121])).
fof(f619,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X2,X3,sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f568,f156])).
fof(f568,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X2,X3,sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)))))))) ) | $spl4),
  inference(superposition,[],[f140,f164])).
fof(f1613,plain,(
  ( ! [X132,X133] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X133) | s(bool,X132) != s(bool,X133)) ) | $spl88),
  inference(cnf_transformation,[],[f1613_D])).
fof(f1613_D,plain,(
  ( ! [X132,X133] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X133) | s(bool,X132) != s(bool,X133)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1619,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(bool,X171) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X170) != s(bool,X172) | s(bool,X172) != s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,X171) != s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl4 | $spl6 | $spl24 | $spl38 | $spl70)),
  inference(subsumption_resolution,[],[f1564,f754])).
fof(f1564,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(bool,'T_0') != s(bool,X170) | s(bool,X171) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X170) != s(bool,X172) | s(bool,X172) != s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,X171) != s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl4 | $spl6 | $spl24 | $spl70)),
  inference(superposition,[],[f1330,f1469])).
fof(f1469,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f1458,f508])).
fof(f1617,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X152) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153) | s(bool,X152) != s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | ($spl4 | $spl6 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f1559,f691])).
fof(f1559,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,'T_0') != s(bool,X151) | s(bool,X152) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153) | s(bool,X152) != s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | ($spl4 | $spl6 | $spl24 | $spl54)),
  inference(superposition,[],[f1268,f1469])).
fof(f1612,plain,(
  ( ! [X132,X133,X131] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,X133) | s(bool,'T_0') = s(bool,X133)) ) | ($spl4 | $spl6 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f1551,f151])).
fof(f1551,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') != s(bool,X131) | s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,X133) | s(bool,'T_0') = s(bool,X133)) ) | ($spl4 | $spl6 | $spl24 | $spl40)),
  inference(superposition,[],[f1067,f1469])).
fof(f1611,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X122)) ) | $spl86),
  inference(cnf_transformation,[],[f1611_D])).
fof(f1611_D,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X122)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1610,plain,(
  ( ! [X121,X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl4 | $spl6 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f1544,f151])).
fof(f1544,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,X121) | s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl4 | $spl6 | $spl24 | $spl38)),
  inference(superposition,[],[f754,f1469])).
fof(f1474,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | p(s(bool,X1))) ) | $spl84),
  inference(cnf_transformation,[],[f1474_D])).
fof(f1474_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | p(s(bool,X1))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1468,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f155])).
fof(f1467,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X22) != s(bool,X23)) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f115])).
fof(f1464,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f114])).
fof(f1463,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f113])).
fof(f1462,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0)))) | s(bool,X7) != s(bool,X8)) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f111])).
fof(f1460,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(sK6(sK2,sK1,sK0),sK5(sK2,sK1,sK0),sK2,'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1)),sK0)),s('type.closSem.v',sK6(sK2,sK1,sK0))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1458,f109])).
fof(f1439,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6)))))) ) | $spl82),
  inference(cnf_transformation,[],[f1439_D])).
fof(f1439_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1438,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6))))) | p(s(bool,X4))) ) | ($spl10 | $spl58)),
  inference(subsumption_resolution,[],[f1429,f236])).
fof(f1429,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6))))) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl58),
  inference(superposition,[],[f1098,f151])).
fof(f1413,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) | $spl80),
  inference(cnf_transformation,[],[f1413_D])).
fof(f1413_D,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1412,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24)) ) | ($spl10 | $spl74)),
  inference(subsumption_resolution,[],[f1407,f236])).
fof(f1407,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X21)) ) | $spl74),
  inference(superposition,[],[f1341,f151])).
fof(f1398,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X86) | s(bool,X85) != s(bool,X86) | s(bool,X82) != s(bool,X83) | p(s(bool,X87)) | p(s(bool,X88)) | s(bool,X81) != s(bool,X88)) ) | $spl74),
  inference(resolution,[],[f1341,f155])).
fof(f1397,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X70) != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X73) != s(bool,X74) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) != s(bool,X76) | s(bool,X72) != s(bool,X73) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X71) != s(bool,X77) | s(bool,X78) != s(bool,X79)) ) | $spl74),
  inference(resolution,[],[f1341,f115])).
fof(f1394,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) != s(bool,X46) | s(bool,X42) != s(bool,X43) | p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X41) != s(bool,X49)) ) | $spl74),
  inference(resolution,[],[f1341,f114])).
fof(f1393,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) != s(bool,X36) | s(bool,X32) != s(bool,X33) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X31) != s(bool,X38)) ) | $spl74),
  inference(resolution,[],[f1341,f113])).
fof(f1392,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X21) != s(bool,X27) | s(bool,X28) != s(bool,X29)) ) | $spl74),
  inference(resolution,[],[f1341,f111])).
fof(f1390,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X1) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f1341,f109])).
fof(f1344,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23) | s(bool,X19) != s(bool,X20)) ) | $spl76),
  inference(cnf_transformation,[],[f1344_D])).
fof(f1344_D,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23) | s(bool,X19) != s(bool,X20)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1343,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21)) ) | ($spl10 | $spl70)),
  inference(subsumption_resolution,[],[f1338,f236])).
fof(f1338,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X18)) ) | $spl70),
  inference(superposition,[],[f1279,f151])).
fof(f1329,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,X72) != s(bool,X74) | s(bool,X75) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X74) != s(bool,X75) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X73) != s(bool,X79)) ) | $spl70),
  inference(resolution,[],[f1279,f155])).
fof(f1328,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl70),
  inference(resolution,[],[f1279,f115])).
fof(f1325,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl70),
  inference(resolution,[],[f1279,f114])).
fof(f1324,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl70),
  inference(resolution,[],[f1279,f113])).
fof(f1323,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl70),
  inference(resolution,[],[f1279,f111])).
fof(f1321,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl70),
  inference(resolution,[],[f1279,f109])).
fof(f1282,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X16) != s(bool,X17)) ) | $spl72),
  inference(cnf_transformation,[],[f1282_D])).
fof(f1282_D,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X16) != s(bool,X17)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1281,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl10 | $spl54)),
  inference(subsumption_resolution,[],[f1276,f236])).
fof(f1276,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X15)) ) | $spl54),
  inference(superposition,[],[f1076,f151])).
fof(f1267,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,X64) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X67) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X65) != s(bool,X70)) ) | $spl54),
  inference(resolution,[],[f1076,f155])).
fof(f1266,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl54),
  inference(resolution,[],[f1076,f115])).
fof(f1263,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl54),
  inference(resolution,[],[f1076,f114])).
fof(f1262,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl54),
  inference(resolution,[],[f1076,f113])).
fof(f1261,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f1076,f111])).
fof(f1259,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl54),
  inference(resolution,[],[f1076,f109])).
fof(f1227,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X75),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X77,X75,X76,X78,X79)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X77))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X75),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X76)))) )),
  inference(resolution,[],[f139,f155])).
fof(f1226,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X67),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X69,X67,X68,X70,X71)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X69))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X67),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f139,f115])).
fof(f1224,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X51),s('type.sptree.spt'('type.clos_known.val_approx'),X52))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X53,X51,X52,X54,X55)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X53))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X51),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) )),
  inference(resolution,[],[f139,f114])).
fof(f1222,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X35),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X37,X35,X36,X38,X39)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X37))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X35),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) )),
  inference(resolution,[],[f139,f113])).
fof(f1221,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X27),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X29,X27,X28,X30,X31)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X29))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X27),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f139,f111])).
fof(f1219,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK4(X13,X11,X12,X14,X15)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12)))) )),
  inference(resolution,[],[f139,f109])).
fof(f1153,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X75,X76,X77,X78,X79)),s('type.closSem.v',X75)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X75))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X76),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) )),
  inference(resolution,[],[f140,f155])).
fof(f1152,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X67,X68,X69,X70,X71)),s('type.closSem.v',X67)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X67))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X68),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f140,f115])).
fof(f1150,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X51,X52,X53,X54,X55)),s('type.closSem.v',X51)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X51))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X52),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X53)))) )),
  inference(resolution,[],[f140,f114])).
fof(f1148,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X35,X36,X37,X38,X39)),s('type.closSem.v',X35)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X35))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X36),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) )),
  inference(resolution,[],[f140,f113])).
fof(f1147,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X27,X28,X29,X30,X31)),s('type.closSem.v',X27)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X27))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X28),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f140,f111])).
fof(f1145,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X11,X12,X13,X14,X15)),s('type.closSem.v',X11)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) )),
  inference(resolution,[],[f140,f109])).
fof(f1141,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X71))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X72),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X73),X74))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X73),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X75)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X71,X72,X75,X74,X73)),s('type.closSem.v',X71)))) )),
  inference(resolution,[],[f140,f155])).
fof(f1140,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X63))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X64),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X65),X66))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X65),X66),s('type.sptree.spt'('type.clos_known.val_approx'),X67)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X63,X64,X67,X66,X65)),s('type.closSem.v',X63))) | s(bool,X69) != s(bool,X70)) )),
  inference(resolution,[],[f140,f115])).
fof(f1137,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X39))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X40),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X41),X42))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X42),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X39,X40,X43,X42,X41)),s('type.closSem.v',X39)))) )),
  inference(resolution,[],[f140,f114])).
fof(f1136,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X31))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X32),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X34),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X31,X32,X35,X34,X33)),s('type.closSem.v',X31)))) )),
  inference(resolution,[],[f140,f113])).
fof(f1135,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X23))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X24),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X25),X26))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X23,X24,X27,X26,X25)),s('type.closSem.v',X23))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f140,f111])).
fof(f1133,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK4(X7,X8,X11,X10,X9)),s('type.closSem.v',X7)))) )),
  inference(resolution,[],[f140,f109])).
fof(f1129,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) | $spl62),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1125,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl40)),
  inference(duplicate_literal_removal,[],[f1118])).
fof(f1118,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl40)),
  inference(superposition,[],[f1067,f151])).
fof(f1128,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) != s(bool,X21)) ) | $spl60),
  inference(cnf_transformation,[],[f1128_D])).
fof(f1128_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) != s(bool,X21)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1127,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl4 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f1117,f151])).
fof(f1117,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X19) | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl4 | $spl24 | $spl40)),
  inference(superposition,[],[f1067,f164])).
fof(f1089,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X48,X49,X50)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X48,X49,X50)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X50),X49))))) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X50),X49),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) )),
  inference(resolution,[],[f141,f155])).
fof(f1088,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X42,X43,X44)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X42,X43,X44)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X44),X43))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X44),X43),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f141,f115])).
fof(f1085,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X24,X25,X26)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X24,X25,X26)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X26),X25))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X26),X25),s('type.sptree.spt'('type.clos_known.val_approx'),X24)))) )),
  inference(resolution,[],[f141,f114])).
fof(f1084,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X18,X19,X20)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X18,X19,X20)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X20),X19))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) )),
  inference(resolution,[],[f141,f113])).
fof(f1083,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X12,X13,X14)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X12,X13,X14)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X14),X13))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X13),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f141,f111])).
fof(f1081,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK6(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK5(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0)))) )),
  inference(resolution,[],[f141,f109])).
fof(f1079,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl56),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s(bool,X13) != s(bool,X14)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1078,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | ($spl10 | $spl40)),
  inference(subsumption_resolution,[],[f1073,f236])).
fof(f1073,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,X12)) ) | $spl40),
  inference(superposition,[],[f765,f151])).
fof(f1066,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl40),
  inference(resolution,[],[f765,f155])).
fof(f1065,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f765,f115])).
fof(f1062,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f765,f114])).
fof(f1061,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f765,f113])).
fof(f1060,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f765,f111])).
fof(f1058,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f765,f109])).
fof(f1019,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X57),s('type.sptree.spt'('type.clos_known.val_approx'),X58)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X59))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X58,X57,X56)),s('type.sptree.spt'('type.clos_known.val_approx'),X58))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X59),s('type.closSem.v',sK6(X58,X57,X56))))) )),
  inference(resolution,[],[f142,f155])).
fof(f1018,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X49),X50),s('type.sptree.spt'('type.clos_known.val_approx'),X51)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X52))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X51,X50,X49)),s('type.sptree.spt'('type.clos_known.val_approx'),X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X52),s('type.closSem.v',sK6(X51,X50,X49)))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f142,f115])).
fof(f1016,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X38))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X37,X36,X35)),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X38),s('type.closSem.v',sK6(X37,X36,X35))))) )),
  inference(resolution,[],[f142,f114])).
fof(f1014,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X21),X22),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X24))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X23,X22,X21)),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X24),s('type.closSem.v',sK6(X23,X22,X21))))) )),
  inference(resolution,[],[f142,f113])).
fof(f1013,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X15),s('type.sptree.spt'('type.clos_known.val_approx'),X16)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X17))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X16,X15,X14)),s('type.sptree.spt'('type.clos_known.val_approx'),X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X17),s('type.closSem.v',sK6(X16,X15,X14)))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f142,f111])).
fof(f1011,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK6(X2,X1,X0))))) )),
  inference(resolution,[],[f142,f109])).
fof(f1009,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X56),s('type.closSem.v',sK6(X57,X58,X59))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X56))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X57,X58,X59)),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X59),X58),s('type.sptree.spt'('type.clos_known.val_approx'),X57)))) )),
  inference(resolution,[],[f142,f155])).
fof(f1008,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X49),s('type.closSem.v',sK6(X50,X51,X52))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X49))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X50,X51,X52)),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X52),X51),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f142,f115])).
fof(f1005,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X28),s('type.closSem.v',sK6(X29,X30,X31))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X28))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X29,X30,X31)),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) )),
  inference(resolution,[],[f142,f114])).
fof(f1004,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X21),s('type.closSem.v',sK6(X22,X23,X24))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X21))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X22,X23,X24)),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X24),X23),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) )),
  inference(resolution,[],[f142,f113])).
fof(f1003,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X14),s('type.closSem.v',sK6(X15,X16,X17))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X14))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X15,X16,X17)),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X17),X16),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f142,f111])).
fof(f1001,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK6(X1,X2,X3))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK5(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) )),
  inference(resolution,[],[f142,f109])).
fof(f988,plain,(
  ( ! [X161,X162] : (~p(s(bool,X161)) | s(bool,'T_0') = s(bool,X162)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f921,f151])).
fof(f921,plain,(
  ( ! [X161,X162] : (s(bool,'F_0') != s(bool,X162) | ~p(s(bool,X161)) | s(bool,'T_0') = s(bool,X162)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f498,f800])).
fof(f985,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X144) != s(bool,X146)) ) | $spl44),
  inference(subsumption_resolution,[],[f913,f152])).
fof(f913,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,'F_0')) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X144) != s(bool,X146)) ) | $spl44),
  inference(superposition,[],[f155,f800])).
fof(f984,plain,(
  ( ! [X140,X138] : (p(s(bool,X140)) | ~p(s(bool,X138))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f983,f800])).
fof(f983,plain,(
  ( ! [X140,X138,X139] : (s(bool,'F_0') != s(bool,X139) | p(s(bool,X140)) | ~p(s(bool,X138))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f911,f340])).
fof(f911,plain,(
  ( ! [X140,X138,X139] : (s(bool,'F_0') != s(bool,X139) | p(s(bool,X140)) | ~p(s(bool,X138)) | p(s(bool,X139))) ) | $spl44),
  inference(superposition,[],[f155,f800])).
fof(f982,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl44),
  inference(subsumption_resolution,[],[f907,f152])).
fof(f907,plain,(
  ( ! [X127,X125,X126,X124] : (p(s(bool,'F_0')) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl44),
  inference(superposition,[],[f115,f800])).
fof(f981,plain,(
  ( ! [X123,X122,X120] : (~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f980,f800])).
fof(f980,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,X121) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f906,f340])).
fof(f906,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,X121) | p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | $spl44),
  inference(superposition,[],[f115,f800])).
fof(f978,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl44),
  inference(subsumption_resolution,[],[f899,f152])).
fof(f899,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'F_0')) | p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl44),
  inference(superposition,[],[f114,f800])).
fof(f977,plain,(
  ( ! [X87,X86,X84] : (p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f976,f800])).
fof(f976,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f897,f340])).
fof(f897,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87)) | p(s(bool,X85))) ) | $spl44),
  inference(superposition,[],[f114,f800])).
fof(f975,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl44),
  inference(subsumption_resolution,[],[f894,f152])).
fof(f894,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'F_0')) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl44),
  inference(superposition,[],[f113,f800])).
fof(f973,plain,(
  ( ! [X59,X58,X56] : (p(s(bool,X58)) | ~p(s(bool,X56)) | s(bool,X58) != s(bool,X59)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f972,f800])).
fof(f972,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | s(bool,X58) != s(bool,X59)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f890,f340])).
fof(f890,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59)) ) | $spl44),
  inference(superposition,[],[f113,f800])).
fof(f971,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl44),
  inference(subsumption_resolution,[],[f887,f152])).
fof(f887,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'F_0')) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl44),
  inference(superposition,[],[f111,f800])).
fof(f970,plain,(
  ( ! [X39,X38,X36] : (~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f969,f800])).
fof(f969,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f885,f340])).
fof(f885,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | $spl44),
  inference(superposition,[],[f111,f800])).
fof(f968,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl44),
  inference(subsumption_resolution,[],[f880,f152])).
fof(f880,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'F_0')) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl44),
  inference(superposition,[],[f109,f800])).
fof(f966,plain,(
  ( ! [X6,X4,X7] : (p(s(bool,X6)) | ~p(s(bool,X4)) | s(bool,X6) != s(bool,X7)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f965,f800])).
fof(f965,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X4)) | s(bool,X6) != s(bool,X7)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f877,f340])).
fof(f877,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) ) | $spl44),
  inference(superposition,[],[f109,f800])).
fof(f962,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,X159)) | p(s(bool,X160)) | p(s(bool,X161))) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f860,f548])).
fof(f860,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,X159)) | p(s(bool,X160)) | p(s(bool,X161)) | s(bool,'F_0') != s(bool,X161)) ) | $spl44),
  inference(superposition,[],[f155,f800])).
fof(f961,plain,(
  ( ! [X152,X151,X150] : (s(bool,X150) != s(bool,X151) | p(s(bool,X152)) | ~p(s(bool,X151))) ) | $spl44),
  inference(subsumption_resolution,[],[f857,f152])).
fof(f857,plain,(
  ( ! [X152,X151,X150] : (s(bool,X150) != s(bool,X151) | p(s(bool,X152)) | ~p(s(bool,X151)) | p(s(bool,'F_0'))) ) | $spl44),
  inference(superposition,[],[f155,f800])).
fof(f960,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X143) != s(bool,X144)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f854,f548])).
fof(f854,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,'F_0') != s(bool,X142) | s(bool,X143) != s(bool,X144)) ) | $spl44),
  inference(superposition,[],[f115,f800])).
fof(f958,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X125) != s(bool,X126) | ~p(s(bool,X127)) | ~p(s(bool,X126)) | s(bool,X127) != s(bool,X128)) ) | $spl44),
  inference(subsumption_resolution,[],[f850,f152])).
fof(f850,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X125) != s(bool,X126) | p(s(bool,'F_0')) | ~p(s(bool,X127)) | ~p(s(bool,X126)) | s(bool,X127) != s(bool,X128)) ) | $spl44),
  inference(superposition,[],[f115,f800])).
fof(f956,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | p(s(bool,X112))) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f846,f548])).
fof(f846,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | p(s(bool,X112)) | s(bool,'F_0') != s(bool,X112)) ) | $spl44),
  inference(superposition,[],[f114,f800])).
fof(f954,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | ~p(s(bool,X100))) ) | $spl44),
  inference(subsumption_resolution,[],[f843,f152])).
fof(f843,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | ~p(s(bool,X100)) | p(s(bool,'F_0'))) ) | $spl44),
  inference(superposition,[],[f114,f800])).
fof(f953,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X90) != s(bool,X92)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f841,f548])).
fof(f841,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X90) != s(bool,X92) | s(bool,'F_0') != s(bool,X91)) ) | $spl44),
  inference(superposition,[],[f113,f800])).
fof(f952,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X69) != s(bool,X70) | p(s(bool,X71)) | ~p(s(bool,X70)) | s(bool,X71) != s(bool,X72)) ) | $spl44),
  inference(subsumption_resolution,[],[f836,f152])).
fof(f836,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X69) != s(bool,X70) | p(s(bool,X71)) | ~p(s(bool,X70)) | p(s(bool,'F_0')) | s(bool,X71) != s(bool,X72)) ) | $spl44),
  inference(superposition,[],[f113,f800])).
fof(f951,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,X64)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f834,f548])).
fof(f834,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,'F_0') != s(bool,X62) | s(bool,X63) != s(bool,X64)) ) | $spl44),
  inference(superposition,[],[f111,f800])).
fof(f949,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X49) != s(bool,X50) | ~p(s(bool,X51)) | ~p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl44),
  inference(subsumption_resolution,[],[f831,f152])).
fof(f831,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,'F_0')) | ~p(s(bool,X51)) | ~p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl44),
  inference(superposition,[],[f111,f800])).
fof(f947,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f827,f548])).
fof(f827,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,'F_0') != s(bool,X35)) ) | $spl44),
  inference(superposition,[],[f109,f800])).
fof(f946,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | ~p(s(bool,X10)) | s(bool,X11) != s(bool,X12)) ) | $spl44),
  inference(subsumption_resolution,[],[f821,f152])).
fof(f821,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,'F_0')) | s(bool,X11) != s(bool,X12)) ) | $spl44),
  inference(superposition,[],[f109,f800])).
fof(f802,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f802_D])).
fof(f802_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f796,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl38)),
  inference(duplicate_literal_removal,[],[f789])).
fof(f789,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl38)),
  inference(superposition,[],[f754,f151])).
fof(f801,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X15)) ) | $spl46),
  inference(cnf_transformation,[],[f801_D])).
fof(f801_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f799,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f788,f151])).
fof(f788,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl24 | $spl38)),
  inference(superposition,[],[f754,f164])).
fof(f768,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl42),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f767,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | ($spl10 | $spl38)),
  inference(subsumption_resolution,[],[f762,f236])).
fof(f762,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl38),
  inference(superposition,[],[f714,f151])).
fof(f753,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl38),
  inference(resolution,[],[f714,f155])).
fof(f752,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f714,f115])).
fof(f749,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f714,f114])).
fof(f748,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f714,f113])).
fof(f747,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f714,f111])).
fof(f745,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f714,f109])).
fof(f702,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f552,f155])).
fof(f701,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f552,f115])).
fof(f698,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl14),
  inference(resolution,[],[f552,f114])).
fof(f697,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f552,f113])).
fof(f696,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f552,f111])).
fof(f694,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f552,f109])).
fof(f661,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f548,f155])).
fof(f660,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f548,f115])).
fof(f657,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl34),
  inference(resolution,[],[f548,f114])).
fof(f656,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f548,f113])).
fof(f655,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f548,f111])).
fof(f653,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f548,f109])).
fof(f236,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8))) ) | $spl10),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f550,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f541,f340])).
fof(f541,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f241,f151])).
fof(f548,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f547,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f536,f508])).
fof(f536,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f241,f151])).
fof(f241,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl14),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X27)) | ~p(s(bool,X28))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f519,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl24),
  inference(resolution,[],[f508,f155])).
fof(f532,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f532_D])).
fof(f532_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f518,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl24),
  inference(resolution,[],[f508,f155])).
fof(f517,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f508,f115])).
fof(f534,plain,(
  ( ! [X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl24),
  inference(subsumption_resolution,[],[f516,f155])).
fof(f516,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl24),
  inference(resolution,[],[f508,f114])).
fof(f515,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f508,f114])).
fof(f514,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl24),
  inference(resolution,[],[f508,f113])).
fof(f513,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f508,f113])).
fof(f512,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f508,f111])).
fof(f511,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl24),
  inference(resolution,[],[f508,f109])).
fof(f510,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f508,f109])).
fof(f497,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f160,f287])).
fof(f287,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl6),
  inference(backward_demodulation,[],[f165,f107])).
fof(f494,plain,(
  $false | ($spl6 | $spl26)),
  inference(trivial_inequality_removal,[],[f492])).
fof(f492,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl26)),
  inference(superposition,[],[f342,f165])).
fof(f493,plain,(
  $false | $spl26),
  inference(equality_resolution,[],[f342])).
fof(f483,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f242])).
fof(f465,plain,(
  $false | $spl18),
  inference(resolution,[],[f429,f154])).
fof(f429,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(equality_resolution,[],[f251])).
fof(f466,plain,(
  $false | $spl18),
  inference(resolution,[],[f429,f153])).
fof(f153,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.state_globals_approx_clock_fupd','thm.bool.TRUTH')).
fof(f467,plain,(
  $false | $spl18),
  inference(resolution,[],[f429,f156])).
fof(f444,plain,(
  $false | $spl20),
  inference(resolution,[],[f264,f154])).
fof(f445,plain,(
  $false | $spl20),
  inference(resolution,[],[f264,f153])).
fof(f446,plain,(
  $false | $spl20),
  inference(resolution,[],[f264,f156])).
fof(f431,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl18),
  inference(subsumption_resolution,[],[f420,f251])).
fof(f420,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(superposition,[],[f251,f151])).
fof(f430,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) ) | $spl18),
  inference(subsumption_resolution,[],[f419,f251])).
fof(f419,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f251,f151])).
fof(f413,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2))) ) | $spl12),
  inference(subsumption_resolution,[],[f402,f239])).
fof(f402,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(superposition,[],[f239,f151])).
fof(f412,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f401,f239])).
fof(f401,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f239,f151])).
fof(f400,plain,(
  $false | $spl8),
  inference(equality_resolution,[],[f235])).
fof(f381,plain,(
  $false | $spl22),
  inference(resolution,[],[f265,f152])).
fof(f371,plain,(
  ( ! [X111,X112] : (s(bool,'F_0') != s(bool,X112) | p(s(bool,X111)) | ~p(s(bool,X112))) ) | $spl6),
  inference(forward_demodulation,[],[f370,f165])).
fof(f370,plain,(
  ( ! [X111,X112] : (p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f331,f152])).
fof(f331,plain,(
  ( ! [X111,X112] : (p(s(bool,'F_0')) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(superposition,[],[f155,f165])).
fof(f367,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X99) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(forward_demodulation,[],[f366,f165])).
fof(f366,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(subsumption_resolution,[],[f325,f152])).
fof(f325,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,'F_0')) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(superposition,[],[f115,f165])).
fof(f235,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10)) ) | $spl8),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f364,plain,(
  ( ! [X94,X92,X93] : (s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(subsumption_resolution,[],[f363,f152])).
fof(f363,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(subsumption_resolution,[],[f362,f155])).
fof(f362,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(forward_demodulation,[],[f323,f165])).
fof(f323,plain,(
  ( ! [X94,X92,X93] : (s(bool,'F_0') != s(bool,X92) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(superposition,[],[f115,f165])).
fof(f361,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl6),
  inference(subsumption_resolution,[],[f322,f113])).
fof(f322,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl6),
  inference(superposition,[],[f115,f165])).
fof(f360,plain,(
  ( ! [X88,X87,X86] : (s(bool,'F_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl6),
  inference(subsumption_resolution,[],[f321,f155])).
fof(f321,plain,(
  ( ! [X88,X87,X86] : (s(bool,'F_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl6),
  inference(superposition,[],[f115,f165])).
fof(f358,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl6),
  inference(forward_demodulation,[],[f357,f165])).
fof(f357,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f317,f152])).
fof(f317,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'F_0')) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(superposition,[],[f114,f165])).
fof(f353,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,X60) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61)) ) | $spl6),
  inference(forward_demodulation,[],[f352,f165])).
fof(f352,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f312,f152])).
fof(f312,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,'F_0')) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(superposition,[],[f113,f165])).
fof(f349,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(forward_demodulation,[],[f348,f165])).
fof(f348,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f305,f152])).
fof(f305,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'F_0')) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(superposition,[],[f111,f165])).
fof(f347,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(subsumption_resolution,[],[f346,f152])).
fof(f346,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(subsumption_resolution,[],[f345,f155])).
fof(f345,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X36)) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(forward_demodulation,[],[f304,f165])).
fof(f304,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X36)) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f111,f165])).
fof(f343,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl6),
  inference(subsumption_resolution,[],[f302,f155])).
fof(f302,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl6),
  inference(superposition,[],[f111,f165])).
fof(f341,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl6),
  inference(subsumption_resolution,[],[f301,f113])).
fof(f301,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f111,f165])).
fof(f339,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19)) ) | $spl6),
  inference(forward_demodulation,[],[f338,f165])).
fof(f338,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f298,f152])).
fof(f298,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'F_0')) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl6),
  inference(superposition,[],[f109,f165])).
fof(f283,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
fof(f282,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
fof(f278,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') != s(bool,X116) | p(s(bool,X115)) | p(s(bool,X116))) ) | $spl2),
  inference(forward_demodulation,[],[f277,f161])).
fof(f277,plain,(
  ( ! [X116,X115] : (p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f228,f156])).
fof(f228,plain,(
  ( ! [X116,X115] : (~p(s(bool,'T_0')) | p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f155,f161])).
fof(f274,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(forward_demodulation,[],[f273,f161])).
fof(f273,plain,(
  ( ! [X105,X106,X104] : (p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(subsumption_resolution,[],[f222,f156])).
fof(f222,plain,(
  ( ! [X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(superposition,[],[f115,f161])).
fof(f272,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X103) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102)) ) | $spl2),
  inference(forward_demodulation,[],[f271,f161])).
fof(f271,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102) | s(bool,X103) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f221,f155])).
fof(f221,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102) | s(bool,X103) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f115,f161])).
fof(f267,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl2),
  inference(subsumption_resolution,[],[f217,f113])).
fof(f217,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl2),
  inference(superposition,[],[f115,f161])).
fof(f266,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl2),
  inference(subsumption_resolution,[],[f216,f155])).
fof(f216,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl2),
  inference(superposition,[],[f115,f161])).
fof(f263,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,X82) | p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82))) ) | $spl2),
  inference(forward_demodulation,[],[f262,f161])).
fof(f262,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f214,f156])).
fof(f214,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f114,f161])).
fof(f257,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,X66) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67)) ) | $spl2),
  inference(forward_demodulation,[],[f256,f161])).
fof(f256,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,X66) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f209,f156])).
fof(f209,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,X66) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f113,f161])).
fof(f254,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,X47) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(subsumption_resolution,[],[f253,f156])).
fof(f253,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X47) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(subsumption_resolution,[],[f252,f115])).
fof(f252,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X47) | p(s(bool,X48)) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(forward_demodulation,[],[f203,f161])).
fof(f203,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,X47) | p(s(bool,X48)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(superposition,[],[f113,f161])).
fof(f250,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(forward_demodulation,[],[f249,f161])).
fof(f249,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(subsumption_resolution,[],[f202,f156])).
fof(f202,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(superposition,[],[f111,f161])).
fof(f248,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(forward_demodulation,[],[f247,f161])).
fof(f247,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f201,f155])).
fof(f201,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f111,f161])).
fof(f243,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(subsumption_resolution,[],[f197,f155])).
fof(f197,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(superposition,[],[f111,f161])).
fof(f240,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl2),
  inference(subsumption_resolution,[],[f196,f113])).
fof(f196,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl2),
  inference(superposition,[],[f111,f161])).
fof(f238,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25)) ) | $spl2),
  inference(forward_demodulation,[],[f237,f161])).
fof(f237,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(subsumption_resolution,[],[f195,f156])).
fof(f195,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) ) | $spl2),
  inference(superposition,[],[f109,f161])).
fof(f234,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f233,f156])).
fof(f233,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f232,f115])).
fof(f232,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X9)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(forward_demodulation,[],[f190,f161])).
fof(f190,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(superposition,[],[f109,f161])).
fof(f177,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f176,f162])).
fof(f176,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl7),
  inference(trivial_inequality_removal,[],[f175])).
fof(f175,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | $spl7),
  inference(superposition,[],[f166,f151])).
fof(f158,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
fof(f157,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),sK1),s('type.sptree.spt'('type.clos_known.val_approx'),sK2))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),sK3),s('type.closSem.state'(sK0),sK1))),s('type.sptree.spt'('type.clos_known.val_approx'),sK2)))),
  inference(superposition,[],[f107,f151])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 214
% Passive clauses: 602
% Generated clauses: 3721
% Final active clauses: 142
% Final passive clauses: 278
% Input formulas: 36
% Initial clauses: 48
% 
% Selected by SInE selection: 31
% SInE iterations: 1
% 
% Duplicate literals: 47
% Trivial inequalities: 114
% Fw subsumption resolutions: 610
% Fw demodulations: 208
% Bw demodulations: 32
% 
% Simple tautologies: 90
% Equational tautologies: 6
% Forward subsumptions: 1613
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 579
% Forward superposition: 922
% Backward superposition: 979
% Self superposition: 4
% Equality resolution: 68
% 
% Split clauses: 452
% Split components: 974
% SAT solver clauses: 479
% SAT solver unit clauses: 11
% SAT solver binary clauses: 107
% 
% Sat splits: 452
% Sat splitting refutations: 18
% 
% Memory used [KB]: 3837
% Time elapsed: 0.289 s
% ------------------------------
% ------------------------------
% Success in time 12.073 s
