Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% (8644)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 165
% Passive clauses: 2875
% Generated clauses: 15031
% Final active clauses: 109
% Final passive clauses: 2311
% Input formulas: 32
% Initial clauses: 62
% 
% Duplicate literals: 590
% Trivial inequalities: 377
% Fw subsumption resolutions: 2029
% Fw demodulations: 1464
% Bw demodulations: 41
% 
% Simple tautologies: 330
% Equational tautologies: 45
% Forward subsumptions: 6734
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 429
% Factoring: 1
% Forward superposition: 4931
% Backward superposition: 4652
% Self superposition: 161
% Equality factoring: 92
% Equality resolution: 34
% 
% Split clauses: 288
% Split components: 579
% SAT solver clauses: 725
% SAT solver unit clauses: 26
% SAT solver binary clauses: 189
% 
% Sat splits: 288
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 722
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 7164
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (8666)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 159
% Passive clauses: 1083
% Generated clauses: 11190
% Final active clauses: 95
% Final passive clauses: 541
% Input formulas: 32
% Initial clauses: 62
% 
% Duplicate literals: 398
% Trivial inequalities: 191
% Fw subsumption resolutions: 903
% Fw demodulations: 910
% Bw demodulations: 411
% Condensations: 712
% Global subsumptions: 9
% 
% Simple tautologies: 178
% Equational tautologies: 35
% Forward subsumptions: 5225
% Fw demodulations to eq. taut.: 65
% Bw demodulations to eq. taut.: 11
% 
% Binary resolution: 200
% Factoring: 1
% Forward superposition: 2862
% Backward superposition: 3140
% Self superposition: 102
% Equality factoring: 50
% Equality resolution: 36
% 
% Split clauses: 53
% Split components: 108
% SAT solver clauses: 1250
% SAT solver unit clauses: 74
% SAT solver binary clauses: 65
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 2
% 
% Sat splits: 53
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1213
% TWLsolver calls for satisfiability: 326
% 
% Memory used [KB]: 9594
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f4873,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f118,f182,f181,f117,f184,f116,f186,f180,f887,f890,f889,f888,f976,f982,f2519,f2524,f980,f3672,f3675,f3674,f3673,f3727,f3730,f3732,f3734,f3738,f3776,f3783,f3994,f4054,f4052,f4790,f4793,f4798,f4799,f4800,f4801,f4802,f4803,f4804,f4805,f4809,f4810,f4866])).
fof(f4866,plain,(
  $false | ($spl0 | $spl13 | $spl17)),
  inference(subsumption_resolution,[],[f4865,f119])).
fof(f119,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.TRUTH')).
fof(f4865,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl13 | $spl17)),
  inference(forward_demodulation,[],[f4864,f4823])).
fof(f4823,plain,(
  s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | $spl0),
  inference(unit_resulting_resolution,[],[f180,f129])).
fof(f129,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f4864,plain,(
  ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | ($spl13 | $spl17)),
  inference(backward_demodulation,[],[f4845,f2524])).
fof(f4845,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(X12,fun(sK1,bool)),X13),s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X14))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X12,fun(sK1,bool)),X13),s(X12,X14))),s(sK1,sK6(sK3,sK1))))) ) | $spl13),
  inference(superposition,[],[f175,f4660])).
fof(f4660,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(trivial_inequality_removal,[],[f4659])).
fof(f4659,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(superposition,[],[f3854,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X30] : (? [X18] : s('type.semanticPrimitives.error_result'(X6),X30) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X18))) | ? [X18] : s('type.semanticPrimitives.error_result'(X6),X30) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X18))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.semanticPrimitives.error_result_nchotomy')).
fof(f3854,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f979,f162])).
fof(f162,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X18,X29] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X18))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X29))) <=> s(X6,X18) = s(X6,X29))),
  file('thm.evalProps.exc_rel_raise|49|','thm.semanticPrimitives.error_result_|49||49||split|0')).
fof(f979,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) | $spl13),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f175,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X3))),s(X1,X2))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X3))),s(X1,X2))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X3))),s(X1,X2))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7] : ! [X19,X20,X10] : s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X20))),s(X7,X19))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X20))),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X19)))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.evalProps.exc_rel_def|split|0')).
fof(f4810,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865,f3866,f3858,f3857,f4808,f185])).
fof(f185,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | $spl7),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f4808,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f3791,f979])).
fof(f3791,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) | $spl14),
  inference(superposition,[],[f153,f982])).
fof(f3857,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | ($spl13 | $spl14)),
  inference(forward_demodulation,[],[f3852,f982])).
fof(f3852,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(unit_resulting_resolution,[],[f979,f153])).
fof(f3858,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | ($spl13 | $spl14)),
  inference(forward_demodulation,[],[f3856,f982])).
fof(f3856,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(trivial_inequality_removal,[],[f3855])).
fof(f3855,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),sK3) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(superposition,[],[f979,f153])).
fof(f3866,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | ($spl13 | $spl14)),
  inference(forward_demodulation,[],[f3860,f982])).
fof(f3860,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl13),
  inference(unit_resulting_resolution,[],[f3851,f153])).
fof(f3865,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X0) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1))))) ) | $spl13),
  inference(superposition,[],[f3851,f153])).
fof(f4661,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | ($spl13 | $spl14)),
  inference(forward_demodulation,[],[f4660,f982])).
fof(f4515,plain,(
  ( ! [X2,X1] : (s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),X1,sK1)) | s(bool,X2) = s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f4509,f119])).
fof(f4509,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),X1,sK1)) | s(bool,X2) = s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(superposition,[],[f3818,f202])).
fof(f4506,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),X0),s(sK1,sK5)))) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),sK6('const.semanticPrimitives.Rraise_1'(s(fun(sK1,bool),X0)),fun(sK1,bool)),sK1))) ) | $spl14),
  inference(superposition,[],[f3818,f1323])).
fof(f4267,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),X0))) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1))))) ) | $spl13),
  inference(superposition,[],[f3862,f153])).
fof(f4264,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)))))))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f3862,f162])).
fof(f3862,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)))))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f3851,f162])).
fof(f3851,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f979,f934])).
fof(f3818,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK1,bool),X6),s(sK1,sK5)))) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),X6,sK1))) ) | $spl14),
  inference(forward_demodulation,[],[f3797,f982])).
fof(f3797,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK1,bool),X6),s(sK1,sK5)))) | s(sK1,sK6(sK3,sK1)) = s(sK1,sK8(sK6(sK3,sK1),X6,sK1))) ) | $spl14),
  inference(superposition,[],[f165,f982])).
fof(f3532,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6)),X6)))) | s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8)))) )),
  inference(forward_demodulation,[],[f3526,f1323])).
fof(f3526,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) | s('type.semanticPrimitives.error_result'(X6),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6)),X6))))) )),
  inference(superposition,[],[f934,f1323])).
fof(f3531,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(forward_demodulation,[],[f3522,f1323])).
fof(f3522,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(unit_resulting_resolution,[],[f1323,f934])).
fof(f2727,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) != s('type.semanticPrimitives.error_result'(X0),X2) | s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X2,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f925,f153])).
fof(f2738,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6)),X6)))) | s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8)))) )),
  inference(forward_demodulation,[],[f2697,f1323])).
fof(f2697,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))) | s('type.semanticPrimitives.error_result'(X6),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7)),'type.semanticPrimitives.error_result'(X6)),X6))))) )),
  inference(superposition,[],[f925,f1323])).
fof(f2737,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(forward_demodulation,[],[f2693,f1323])).
fof(f2693,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1)))),'type.semanticPrimitives.error_result'(X0))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(unit_resulting_resolution,[],[f1323,f925])).
fof(f2580,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),X1))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f699,f153])).
fof(f2556,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),X1))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f699,f153])).
fof(f2553,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))))) )),
  inference(unit_resulting_resolution,[],[f699,f162])).
fof(f1876,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,happ(s(fun(X127,bool),happ(s(fun(X126,fun(X127,bool)),X128),s(X126,X129))),s(X127,X130))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X126,fun(X127,bool)),X128),s('type.semanticPrimitives.error_result'(X126),'const.semanticPrimitives.Rraise_1'(s(X126,X129))),s('type.semanticPrimitives.error_result'(X127),'const.semanticPrimitives.Rraise_1'(s(X127,X130)))),X131,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X131),s(bool,happ(s(fun(X127,bool),happ(s(fun(X126,fun(X127,bool)),X128),s(X126,X129))),s(X127,X130))))))) )),
  inference(forward_demodulation,[],[f1849,f175])).
fof(f1849,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (~p(s(bool,happ(s(fun(bool,bool),X131),s(bool,happ(s(fun(X127,bool),happ(s(fun(X126,fun(X127,bool)),X128),s(X126,X129))),s(X127,X130)))))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X126,fun(X127,bool)),X128),s('type.semanticPrimitives.error_result'(X126),'const.semanticPrimitives.Rraise_1'(s(X126,X129))),s('type.semanticPrimitives.error_result'(X127),'const.semanticPrimitives.Rraise_1'(s(X127,X130))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X126,fun(X127,bool)),X128),s('type.semanticPrimitives.error_result'(X126),'const.semanticPrimitives.Rraise_1'(s(X126,X129))),s('type.semanticPrimitives.error_result'(X127),'const.semanticPrimitives.Rraise_1'(s(X127,X130)))),X131,bool))) )),
  inference(superposition,[],[f165,f175])).
fof(f1844,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X96,bool),happ(s(fun(X95,fun(X96,bool)),X97),s(X95,X98))),s(X96,X99))))) | s(bool,X100) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X95,fun(X96,bool)),X97),s('type.semanticPrimitives.error_result'(X95),'const.semanticPrimitives.Rraise_1'(s(X95,X98))),s('type.semanticPrimitives.error_result'(X96),'const.semanticPrimitives.Rraise_1'(s(X96,X99)))))) )),
  inference(superposition,[],[f162,f175])).
fof(f1841,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X80,bool),happ(s(fun(X79,fun(X80,bool)),X81),s(X79,X82))),s(X80,X83)))),bool)) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X79,fun(X80,bool)),X81),s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rraise_1'(s(X79,X82))),s('type.semanticPrimitives.error_result'(X80),'const.semanticPrimitives.Rraise_1'(s(X80,X83)))))) )),
  inference(superposition,[],[f1419,f175])).
fof(f1840,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'F_0') != s(bool,happ(s(fun(X74,bool),happ(s(fun(X73,fun(X74,bool)),X75),s(X73,X76))),s(X74,X77))) | s(bool,X78) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X73,fun(X74,bool)),X75),s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rraise_1'(s(X73,X76))),s('type.semanticPrimitives.error_result'(X74),'const.semanticPrimitives.Rraise_1'(s(X74,X77))))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f649,f175])).
fof(f1826,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),happ(s(fun(X34,fun(X35,bool)),X36),s(X34,X37))),s(X35,X38))) | s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X34,fun(X35,bool)),X36),s('type.semanticPrimitives.error_result'(X34),'const.semanticPrimitives.Rraise_1'(s(X34,X37))),s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,X38)))))) )),
  inference(superposition,[],[f175,f126])).
fof(f1824,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X27) = s(bool,happ(s(fun(X23,bool),happ(s(fun(X22,fun(X23,bool)),X24),s(X22,X25))),s(X23,X26))) | s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X22,fun(X23,bool)),X24),s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rraise_1'(s(X22,X25))),s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rraise_1'(s(X23,X26))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f175,f202])).
fof(f1822,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),happ(s(fun(X10,fun(X11,bool)),X12),s(X10,X13))),s(X11,X14))) | s(bool,X15) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X10,fun(X11,bool)),X12),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X13))),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f175,f202])).
fof(f1820,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(X2,fun(X0,bool)),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4))),s('type.semanticPrimitives.error_result'(X0),X1))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X2,fun(X0,bool)),X3),s(X2,X4))),s(X0,sK6(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f175,f153])).
fof(f1797,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X2,bool)),X3),s('type.semanticPrimitives.error_result'(X0),X1),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4))))) = s(bool,happ(s(fun(X2,bool),happ(s(fun(X0,fun(X2,bool)),X3),s(X0,sK6(X1,X0)))),s(X2,X4))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f175,f153])).
fof(f1607,plain,(
  ( ! [X37,X36] : (s(bool,X37) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool)) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f1595,f119])).
fof(f1595,plain,(
  ( ! [X37,X36] : (~p(s(bool,'T_0')) | s(bool,X37) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool)) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X36)) )),
  inference(superposition,[],[f1557,f1419])).
fof(f1598,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1583,f119])).
fof(f1583,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1557,f202])).
fof(f1572,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))))) )),
  inference(forward_demodulation,[],[f1566,f1323])).
fof(f1566,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X3)),'type.semanticPrimitives.abort')))))))) )),
  inference(unit_resulting_resolution,[],[f1323,f177])).
fof(f1562,plain,(
  ( ! [X35,X34] : (s(bool,X34) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool),X35,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(forward_demodulation,[],[f1532,f1323])).
fof(f1532,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool)) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool),X35,bool)) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f165,f1419])).
fof(f1561,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | ~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X33)))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(forward_demodulation,[],[f1531,f1323])).
fof(f1531,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X33)))) | s(bool,X33) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X31)),bool)) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f164,f1419])).
fof(f1559,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f1529,f1323])).
fof(f1529,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X28) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f163,f1419])).
fof(f1527,plain,(
  ( ! [X24,X23] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) | s(bool,X24) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X23)),bool)) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f162,f1419])).
fof(f1558,plain,(
  ( ! [X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1526,f167])).
fof(f1526,plain,(
  ( ! [X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool)))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f153,f1419])).
fof(f1557,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1556,f1323])).
fof(f1556,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1520,f119])).
fof(f1520,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f171,f1419])).
fof(f1555,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X3)),bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1516,f119])).
fof(f1516,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X3)),bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f132,f1419])).
fof(f1450,plain,(
  ( ! [X85,X83,X84] : (s('type.semanticPrimitives.error_result'(X84),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X83))) != s('type.semanticPrimitives.error_result'(X84),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X85))) | s('type.semanticPrimitives.abort',X85) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X83)),'type.semanticPrimitives.abort'))) )),
  inference(superposition,[],[f160,f1323])).
fof(f1446,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),X66),s(X65,X67)))) | s(X65,X67) = s(X65,sK8(X67,sK6('const.semanticPrimitives.Rraise_1'(s(fun(X65,bool),X66)),fun(X65,bool)),X65))) )),
  inference(superposition,[],[f165,f1323])).
fof(f1486,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(X62,bool),X63),s(X62,X64)))) | p(s(bool,happ(s(fun(X62,bool),X63),s(X62,sK8(X64,sK6('const.semanticPrimitives.Rraise_1'(s(fun(X62,bool),X63)),fun(X62,bool)),X62)))))) )),
  inference(forward_demodulation,[],[f1445,f1323])).
fof(f1445,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(X62,bool),X63),s(X62,sK8(X64,sK6('const.semanticPrimitives.Rraise_1'(s(fun(X62,bool),X63)),fun(X62,bool)),X62))))) | ~p(s(bool,happ(s(fun(X62,bool),sK6('const.semanticPrimitives.Rraise_1'(s(fun(X62,bool),X63)),fun(X62,bool))),s(X62,X64))))) )),
  inference(superposition,[],[f166,f1323])).
fof(f1442,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X55)),bool)) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f649,f1323])).
fof(f1476,plain,(
  ( ! [X19,X20,X18] : (s(X18,X19) = s(X18,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(X18,X19)),X18),X20,X18)) | ~p(s(bool,happ(s(fun(X18,bool),X20),s(X18,X19))))) )),
  inference(forward_demodulation,[],[f1425,f1323])).
fof(f1425,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),X20),s(X18,X19)))) | s(X18,sK6('const.semanticPrimitives.Rraise_1'(s(X18,X19)),X18)) = s(X18,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(X18,X19)),X18),X20,X18))) )),
  inference(superposition,[],[f165,f1323])).
fof(f1419,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool)) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f1323,f126])).
fof(f1417,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X6)),bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1323,f202])).
fof(f1415,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1323,f202])).
fof(f1388,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4) | s(bool,X5) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1381,f119])).
fof(f1381,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4) | s(bool,X5) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f176,f202])).
fof(f1380,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X2,fun(X0,bool)),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X0),X1)))) | s('type.semanticPrimitives.abort',X4) = s('type.semanticPrimitives.abort',sK7(X1,X0)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f176,f153])).
fof(f1378,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X2,bool)),X3),s('type.semanticPrimitives.error_result'(X0),X1),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))) | s('type.semanticPrimitives.abort',X4) = s('type.semanticPrimitives.abort',sK7(X1,X0)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f176,f153])).
fof(f1350,plain,(
  ( ! [X14,X12,X13] : (s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X13))) != s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X14))) | s(X12,X14) = s(X12,sK6('const.semanticPrimitives.Rraise_1'(s(X12,X13)),X12))) )),
  inference(superposition,[],[f162,f916])).
fof(f1347,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool))) )),
  inference(superposition,[],[f916,f126])).
fof(f1345,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f916,f202])).
fof(f1343,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,X1) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f916,f202])).
fof(f1323,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK6('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0))) )),
  inference(unit_resulting_resolution,[],[f916,f162])).
fof(f1319,plain,(
  ( ! [X52,X50,X51] : (s('type.semanticPrimitives.abort',X50) = s('type.semanticPrimitives.abort',sK8(sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X50)),X51),X52,'type.semanticPrimitives.abort')) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X52),s('type.semanticPrimitives.abort',X50))))) )),
  inference(forward_demodulation,[],[f1310,f1261])).
fof(f1310,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X52),s('type.semanticPrimitives.abort',X50)))) | s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X50)),X51)) = s('type.semanticPrimitives.abort',sK8(sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X50)),X51),X52,'type.semanticPrimitives.abort'))) )),
  inference(superposition,[],[f165,f1261])).
fof(f1305,plain,(
  ( ! [X35,X36,X34] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X34))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X36))) | s('type.semanticPrimitives.abort',X36) = s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X34)),X35))) )),
  inference(superposition,[],[f162,f1261])).
fof(f1298,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) != s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X5))) | s('type.semanticPrimitives.abort',X5) = s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2)),X3))) )),
  inference(superposition,[],[f160,f1261])).
fof(f1292,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK9(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK9(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f172,f126])).
fof(f1290,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK9(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK9(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f172,f202])).
fof(f1288,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK9(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK9(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f172,f202])).
fof(f1287,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK9(X16,X17,X18,bool))) )),
  inference(superposition,[],[f172,f126])).
fof(f1285,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK9(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f172,f202])).
fof(f1283,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK9(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f172,f202])).
fof(f1265,plain,(
  ( ! [X14,X12,X13] : (s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X13))) != s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X14))) | s('type.semanticPrimitives.abort',X14) = s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X13)),X12))) )),
  inference(superposition,[],[f160,f915])).
fof(f1261,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.abort',X0) = s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)),X1))) )),
  inference(unit_resulting_resolution,[],[f915,f160])).
fof(f1242,plain,(
  ( ! [X4,X2,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))) | s(X1,X3) != s(X1,X4) | s(bool,X5) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1230,f119])).
fof(f1230,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X4)))) | s(X1,X3) != s(X1,X4) | s(bool,X5) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f164,f202])).
fof(f1222,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X37)))) | s(bool,'F_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f164,f126])).
fof(f1220,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X29)))) | s(bool,'T_0') != s(bool,X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f164,f202])).
fof(f1219,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f164,f126])).
fof(f1217,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17)))) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X19)))) | s(bool,X16) != s(bool,X19) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f164,f202])).
fof(f1215,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | s(bool,X8) != s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f164,f202])).
fof(f1214,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X4),bool),X6),s('type.semanticPrimitives.error_result'(X4),X5)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X4),bool),X6),s('type.semanticPrimitives.error_result'(X4),X7)))) | s('type.semanticPrimitives.error_result'(X4),X7) != s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X5,X4)))) | s('type.semanticPrimitives.error_result'(X4),X5) = s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,sK6(X5,X4))))) )),
  inference(superposition,[],[f164,f153])).
fof(f1213,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X0),bool),X2),s('type.semanticPrimitives.error_result'(X0),X1)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X0),bool),X2),s('type.semanticPrimitives.error_result'(X0),X3)))) | s('type.semanticPrimitives.error_result'(X0),X3) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f164,f153])).
fof(f1185,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rraise_1'(s(X103,X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X107)))),X108,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f1169,f174])).
fof(f1169,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'F_0')))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rraise_1'(s(X103,X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X107))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rraise_1'(s(X103,X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X107)))),X108,bool))) )),
  inference(superposition,[],[f165,f174])).
fof(f1165,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X84))) | s(bool,X84) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X79,fun(X80,bool)),X81),s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rraise_1'(s(X79,X82))),s('type.semanticPrimitives.error_result'(X80),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X83)))))) )),
  inference(superposition,[],[f162,f174])).
fof(f1178,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,X10) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1148,f188])).
fof(f1148,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X10) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f174,f202])).
fof(f1147,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X2,fun(X0,bool)),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4))),s('type.semanticPrimitives.error_result'(X0),X1))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f174,f153])).
fof(f1146,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X2,bool)),X3),s('type.semanticPrimitives.error_result'(X0),X1),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f174,f153])).
fof(f1140,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) = s(bool,'const.evalProps.exc_rel_3'(s(fun(bool,fun(X45,bool)),X46),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X44))),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X47))))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f174,f126])).
fof(f1122,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1005,f162])).
fof(f1118,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rraise_1'(s(X104,X107)))),X108,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f1103,f173])).
fof(f1103,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'F_0')))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rraise_1'(s(X104,X107))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X103,fun(X104,bool)),X105),s('type.semanticPrimitives.error_result'(X103),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X106))),s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rraise_1'(s(X104,X107)))),X108,bool))) )),
  inference(superposition,[],[f165,f173])).
fof(f1099,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X84))) | s(bool,X84) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X79,fun(X80,bool)),X81),s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X82))),s('type.semanticPrimitives.error_result'(X80),'const.semanticPrimitives.Rraise_1'(s(X80,X83)))))) )),
  inference(superposition,[],[f162,f173])).
fof(f1111,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,X10) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1082,f188])).
fof(f1082,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X10) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f173,f202])).
fof(f1081,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X2,fun(X0,bool)),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X0),X1))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f173,f153])).
fof(f1076,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X45,fun(bool,bool)),X46),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X47))),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X44))))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f173,f126])).
fof(f1066,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X2,bool)),X3),s('type.semanticPrimitives.error_result'(X0),X1),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f173,f153])).
fof(f1053,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f694,f162])).
fof(f1005,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f984,f162])).
fof(f984,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))))) | $spl7),
  inference(unit_resulting_resolution,[],[f185,f162])).
fof(f694,plain,(
  s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f188,f162])).
fof(f720,plain,(
  ( ! [X15,X16] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f162,f126])).
fof(f808,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,X15) = s(bool,sK8(X15,X16,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f165,f126])).
fof(f878,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),X15),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | s(bool,'T_0') = s(bool,sK8(X14,X15,bool))) )),
  inference(superposition,[],[f166,f126])).
fof(f922,plain,(
  ( ! [X9] : (s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X9,bool)))) | s(bool,'T_0') = s(bool,sK6(X9,bool))) )),
  inference(superposition,[],[f153,f126])).
fof(f940,plain,(
  ( ! [X23,X21,X22] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),X22))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),X23))) | s('type.semanticPrimitives.error_result'(X21),X23) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,sK6(X22,X21)))) | s('type.semanticPrimitives.error_result'(X21),X22) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X22,X21))))) )),
  inference(superposition,[],[f162,f153])).
fof(f938,plain,(
  ( ! [X17,X15,X16] : (s('type.semanticPrimitives.error_result'(X15),X16) != s('type.semanticPrimitives.error_result'(X15),X17) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X15)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X15),X17))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X15)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rraise_1'(s(X15,sK6(X16,X15)))))) | s('type.semanticPrimitives.error_result'(X15),X16) = s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X16,X15))))) )),
  inference(superposition,[],[f163,f153])).
fof(f937,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X12),bool),X14),s('type.semanticPrimitives.error_result'(X12),X13)))) | s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,sK6(X13,X12)))) = s('type.semanticPrimitives.error_result'(X12),sK8('const.semanticPrimitives.Rraise_1'(s(X12,sK6(X13,X12))),X14,'type.semanticPrimitives.error_result'(X12))) | s('type.semanticPrimitives.error_result'(X12),X13) = s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X13,X12))))) )),
  inference(superposition,[],[f165,f153])).
fof(f934,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) | s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X7,X6))))) )),
  inference(superposition,[],[f167,f153])).
fof(f932,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK6(X1,X0)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X1,X0))))) )),
  inference(superposition,[],[f162,f153])).
fof(f929,plain,(
  ( ! [X19,X20,X18] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X18),X19))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X18),X20))) | s('type.semanticPrimitives.error_result'(X18),X20) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X19,X18)))) | s('type.semanticPrimitives.error_result'(X18),X19) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rraise_1'(s(X18,sK6(X19,X18))))) )),
  inference(superposition,[],[f162,f153])).
fof(f927,plain,(
  ( ! [X14,X12,X13] : (s('type.semanticPrimitives.error_result'(X12),X13) != s('type.semanticPrimitives.error_result'(X12),X14) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X12)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X12),X14))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X12)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X13,X12)))))) | s('type.semanticPrimitives.error_result'(X12),X13) = s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,sK6(X13,X12))))) )),
  inference(superposition,[],[f163,f153])).
fof(f926,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X9),bool),X11),s('type.semanticPrimitives.error_result'(X9),X10)))) | s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X10,X9)))) = s('type.semanticPrimitives.error_result'(X9),sK8('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X10,X9))),X11,'type.semanticPrimitives.error_result'(X9))) | s('type.semanticPrimitives.error_result'(X9),X10) = s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,sK6(X10,X9))))) )),
  inference(superposition,[],[f165,f153])).
fof(f925,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))) | s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,sK6(X7,X6))))) )),
  inference(superposition,[],[f167,f153])).
fof(f923,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) | s('type.semanticPrimitives.abort',X2) = s('type.semanticPrimitives.abort',sK7(X1,X0)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) )),
  inference(superposition,[],[f160,f153])).
fof(f919,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X4,bool)))) | s(bool,'T_0') = s(bool,sK6(X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f153,f202])).
fof(f917,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X0,bool)))) | s(bool,X1) = s(bool,sK6(X0,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f153,f202])).
fof(f915,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f167,f153])).
fof(f916,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f167,f153])).
fof(f875,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'T_0') = s(bool,sK8(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f166,f202])).
fof(f873,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK8(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f166,f202])).
fof(f828,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f816,f119])).
fof(f816,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f165,f202])).
fof(f811,plain,(
  ( ! [X24,X23] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X23)))) | s(bool,'F_0') = s(bool,sK8('F_0',X24,bool)) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f165,f126])).
fof(f809,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17)))) | s(bool,'T_0') = s(bool,sK8('T_0',X18,bool)) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f165,f202])).
fof(f805,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X8)))) | s(bool,X7) = s(bool,sK8(X7,X9,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f165,f202])).
fof(f803,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) | s(bool,X1) = s(bool,sK8(X1,X2,bool)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f165,f202])).
fof(f762,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f163,f126])).
fof(f760,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X18))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f163,f202])).
fof(f759,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,X16) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f163,f126])).
fof(f756,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f163,f202])).
fof(f754,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X2) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f163,f202])).
fof(f709,plain,(
  ( ! [X24,X23] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X23))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f162,f126])).
fof(f703,plain,(
  ( ! [X8,X7,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f162,f202])).
fof(f701,plain,(
  ( ! [X2,X3,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f162,f202])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))) )),
  inference(unit_resulting_resolution,[],[f167,f162])).
fof(f660,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f649,f126])).
fof(f650,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f458,f119])).
fof(f458,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f171,f202])).
fof(f649,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f454,f119])).
fof(f454,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f132,f202])).
fof(f441,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f128,f202])).
fof(f375,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f202,f202])).
fof(f259,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f251,f187])).
fof(f251,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f168,f126])).
fof(f202,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f126,f126])).
fof(f188,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f187,f128])).
fof(f176,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))) | s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))) | s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4)) & (s('type.semanticPrimitives.abort',X3) != s('type.semanticPrimitives.abort',X4) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))) <=> s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))) <=> s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',X4))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : ! [X21,X22,X23] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X21),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X23))),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X22)))))) <=> s('type.semanticPrimitives.abort',X22) = s('type.semanticPrimitives.abort',X23))),
  file('thm.evalProps.exc_rel_raise|49|','thm.evalProps.exc_rel_def|split|1')).
fof(f177,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.abort',X3) != s('type.semanticPrimitives.abort',X4) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f174,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X24,X25,X26] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X25),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X24))),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X26)))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.evalProps.exc_rel_def|split|2')).
fof(f173,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : ! [X25,X27,X28] : s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X25),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X28))),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X27)))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.evalProps.exc_rel_def|split|3')).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.exc_rel_raise|49|',aHL_EXT)).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.IMP_ANTISYM_AX')).
fof(f171,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f82])).
fof(f167,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X29,X18] : ~s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X29))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X18)))),
  file('thm.evalProps.exc_rel_raise|49|','thm.semanticPrimitives.error_result_distinct')).
fof(f164,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK8(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X17,X18] : (? [X15] : (s(X6,X15) = s(X6,X18) & p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X15))))) <=> p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X18)))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.UNWIND_THM2')).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK8(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f113])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f163,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f160,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2)) & (s('type.semanticPrimitives.abort',X1) != s('type.semanticPrimitives.abort',X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2)))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) <=> s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) <=> s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X18,X29] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X18))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X29))) <=> s('type.semanticPrimitives.abort',X18) = s('type.semanticPrimitives.abort',X29))),
  file('thm.evalProps.exc_rel_raise|49|','thm.semanticPrimitives.error_result_|49||49||split|1')).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.abort',X1) != s('type.semanticPrimitives.abort',X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f133,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f128,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f89])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.exc_rel_raise|49|',aHL_BOOL_CASES)).
fof(f187,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f122,f119])).
fof(f122,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.evalProps.exc_rel_raise|49|','thm.bool.NOT_CLAUSES|split|2')).
fof(f4809,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865,f3866,f3858,f3857,f4808])).
fof(f4805,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865,f3866,f3858,f3857])).
fof(f4804,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865,f3866,f3858])).
fof(f4803,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865,f3866])).
fof(f4802,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661,f3865])).
fof(f4801,plain,(
  $false | ($spl7 | $spl13 | $spl14)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f185,f984,f1005,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1122,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1607,f1797,f1820,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f2553,f2556,f2580,f2737,f2738,f2727,f3531,f3532,f982,f979,f3818,f3854,f3851,f3862,f4264,f4267,f4506,f4515,f4661])).
fof(f4800,plain,(
  $false | ($spl0 | $spl6 | $spl14 | $spl17)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f713,f754,f756,f759,f760,f762,f822,f803,f805,f809,f811,f828,f873,f875,f887,f890,f913,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f936,f937,f938,f940,f972,f978,f976,f922,f878,f808,f720,f694,f990,f1026,f1049,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1605,f1607,f1619,f1620,f1737,f1739,f1741,f1745,f1656,f1748,f1749,f1750,f1751,f1752,f1686,f1767,f1769,f1770,f1797,f1820,f1862,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f1899,f1902,f1926,f1947,f2014,f2016,f2018,f2019,f2026,f2027,f2029,f2030,f2088,f2089,f2108,f2110,f2111,f2112,f2113,f2437,f2444,f2148,f2150,f2152,f2153,f2160,f2161,f2163,f2164,f2487,f2490,f2524,f2533,f2532,f2547,f2553,f2556,f2580,f2594,f2679,f2596,f2597,f2598,f2680,f2602,f2603,f2606,f2608,f2610,f2681,f2682,f2683,f2684,f2685,f2617,f2687,f2619,f2620,f2688,f2691,f2692,f2656,f2658,f2660,f2661,f2737,f2738,f2727,f2782,f2783,f2785,f3347,f3340,f3326,f3376,f2840,f2846,f2847,f2873,f3392,f2900,f3404,f2965,f2967,f2979,f2980,f2981,f2984,f3469,f3470,f3475,f3476,f3033,f3035,f3048,f3049,f3052,f3055,f3056,f3057,f3081,f3085,f3086,f3091,f3129,f3131,f3143,f3144,f3145,f3148,f3515,f3516,f3172,f3520,f3521,f3207,f3195,f3531,f3532,f3572,f3578,f3583,f3586,f3587,f3589,f3590,f3592,f3660,f3668,f3622,f3625,f3672,f3675,f3702,f3704,f3705,f3706,f3707,f3708,f3744,f186,f982,f3812,f3818,f3824,f3825,f3826,f3827,f3828,f3850,f3885,f3886,f4062,f4100,f4084,f4108,f4212,f4252,f4235,f4260,f4410,f4411,f4413,f4421,f4465,f4472,f4505,f4506,f4515,f4549,f4550,f4551,f4599,f4556,f4600,f4558,f4601,f4574,f4621,f4663,f4666,f4667,f4668,f4705,f4706,f4713,f4716,f4717,f4719,f4720,f4721,f4722,f4723,f4725,f4730,f4731,f4732,f4733,f4737,f4739,f4741,f4742,f4753,f4754,f4755,f4756,f4757,f4759,f4764,f4765,f4766,f4767,f4771,f4773,f4775,f4776,f4797,f180])).
fof(f4797,plain,(
  ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | ($spl6 | $spl14 | $spl17)),
  inference(forward_demodulation,[],[f4796,f3827])).
fof(f4796,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl14 | $spl17)),
  inference(forward_demodulation,[],[f2524,f982])).
fof(f4776,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,happ(s(fun(sK1,bool),happ(s(fun(X88,fun(sK1,bool)),X89),s(X88,X90))),s(sK1,sK5)))))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X88,fun(sK1,bool)),X89),s('type.semanticPrimitives.error_result'(X88),'const.semanticPrimitives.Rraise_1'(s(X88,X90))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X88,fun(sK1,bool)),X89),s('type.semanticPrimitives.error_result'(X88),'const.semanticPrimitives.Rraise_1'(s(X88,X90))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X91,bool))) ) | $spl6),
  inference(superposition,[],[f165,f3827])).
fof(f4775,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,happ(s(fun(sK1,bool),happ(s(fun(X83,fun(sK1,bool)),X84),s(X83,X85))),s(sK1,sK5)))))) | p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X87)))) | s(bool,X87) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X83,fun(sK1,bool)),X84),s('type.semanticPrimitives.error_result'(X83),'const.semanticPrimitives.Rraise_1'(s(X83,X85))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f164,f3827])).
fof(f4773,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X78) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(X75,fun(sK1,bool)),X76),s(X75,X77))),s(sK1,sK5))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X78))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(X75,fun(sK1,bool)),X76),s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,X77))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) ) | $spl6),
  inference(superposition,[],[f163,f3827])).
fof(f4771,plain,(
  ( ! [X70,X68,X69,X67] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X70))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X67,fun(sK1,bool)),X68),s(X67,X69))),s(sK1,sK5))))) | s(bool,X70) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X67,fun(sK1,bool)),X68),s('type.semanticPrimitives.error_result'(X67),'const.semanticPrimitives.Rraise_1'(s(X67,X69))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f162,f3827])).
fof(f4767,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X54,fun(sK1,bool)),X55),s(X54,X56))),s(sK1,sK5)))) | s(bool,X57) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X54,fun(sK1,bool)),X55),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X56))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,X57)) ) | $spl6),
  inference(superposition,[],[f1557,f3827])).
fof(f4766,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X51,fun(sK1,bool)),X52),s(X51,X53))),s(sK1,sK5)))),bool)) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X51,fun(sK1,bool)),X52),s('type.semanticPrimitives.error_result'(X51),'const.semanticPrimitives.Rraise_1'(s(X51,X53))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f1419,f3827])).
fof(f4765,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(X47,fun(sK1,bool)),X48),s(X47,X49))),s(sK1,sK5))) | s(bool,X50) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X47,fun(sK1,bool)),X48),s('type.semanticPrimitives.error_result'(X47),'const.semanticPrimitives.Rraise_1'(s(X47,X49))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X50)) ) | $spl6),
  inference(superposition,[],[f649,f3827])).
fof(f4764,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X43,fun(sK1,bool)),X44),s(X43,X45))),s(sK1,sK5)))) | s(bool,X46) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X43,fun(sK1,bool)),X44),s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rraise_1'(s(X43,X45))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X46)) ) | $spl6),
  inference(superposition,[],[f259,f3827])).
fof(f4759,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X23,fun(sK1,bool)),X24),s(X23,X25))),s(sK1,sK5)))) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X23,fun(sK1,bool)),X24),s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rraise_1'(s(X23,X25))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f171,f3827])).
fof(f4757,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X15,fun(sK1,bool)),X16),s(X15,X17))),s(sK1,sK5)))) | p(s(bool,X18)) | s(bool,X18) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X15,fun(sK1,bool)),X16),s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rraise_1'(s(X15,X17))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f168,f3827])).
fof(f4756,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X12,fun(sK1,bool)),X13),s(X12,X14))),s(sK1,sK5)))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X12,fun(sK1,bool)),X13),s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X14))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f133,f3827])).
fof(f4755,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X9,fun(sK1,bool)),X10),s(X9,X11))),s(sK1,sK5)))) | s(bool,'F_0') != s(bool,'const.evalProps.exc_rel_3'(s(fun(X9,fun(sK1,bool)),X10),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f132,f3827])).
fof(f4754,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X6,fun(sK1,bool)),X7),s(X6,X8))),s(sK1,sK5)))) | s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(sK1,bool)),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f129,f3827])).
fof(f4753,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(X3,fun(sK1,bool)),X4),s(X3,X5))),s(sK1,sK5))) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,X5))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) ) | $spl6),
  inference(superposition,[],[f128,f3827])).
fof(f4742,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,'const.evalProps.exc_rel_3'(s(fun(X102,fun(sK1,bool)),X103),s('type.semanticPrimitives.error_result'(X102),'const.semanticPrimitives.Rraise_1'(s(X102,X104))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) | s(bool,happ(s(fun(sK1,bool),happ(s(fun(X102,fun(sK1,bool)),X103),s(X102,X104))),s(sK1,sK5))) = s(bool,sK8(happ(s(fun(sK1,bool),happ(s(fun(X102,fun(sK1,bool)),X103),s(X102,X104))),s(sK1,sK5)),X105,bool))) ) | $spl6),
  inference(superposition,[],[f165,f3827])).
fof(f4741,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),X100),s(bool,'const.evalProps.exc_rel_3'(s(fun(X97,fun(sK1,bool)),X98),s('type.semanticPrimitives.error_result'(X97),'const.semanticPrimitives.Rraise_1'(s(X97,X99))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) | p(s(bool,happ(s(fun(bool,bool),X100),s(bool,X101)))) | s(bool,X101) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(X97,fun(sK1,bool)),X98),s(X97,X99))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f164,f3827])).
fof(f4739,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,X92) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X89,fun(sK1,bool)),X90),s('type.semanticPrimitives.error_result'(X89),'const.semanticPrimitives.Rraise_1'(s(X89,X91))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X92))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X89,fun(sK1,bool)),X90),s(X89,X91))),s(sK1,sK5)))))) ) | $spl6),
  inference(superposition,[],[f163,f3827])).
fof(f4737,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X84))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(X81,fun(sK1,bool)),X82),s('type.semanticPrimitives.error_result'(X81),'const.semanticPrimitives.Rraise_1'(s(X81,X83))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s(bool,X84) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X81,fun(sK1,bool)),X82),s(X81,X83))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f162,f3827])).
fof(f4733,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X68,fun(sK1,bool)),X69),s('type.semanticPrimitives.error_result'(X68),'const.semanticPrimitives.Rraise_1'(s(X68,X70))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(bool,X71) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X68,fun(sK1,bool)),X69),s(X68,X70))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f1557,f3827])).
fof(f4732,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(X65,fun(sK1,bool)),X66),s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,X67))),s('type.semanticPrimitives.error_result'(sK1),sK3)))),bool)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X65,fun(sK1,bool)),X66),s(X65,X67))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f1419,f3827])).
fof(f4731,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,'F_0') != s(bool,'const.evalProps.exc_rel_3'(s(fun(X61,fun(sK1,bool)),X62),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X63))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,X64) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X61,fun(sK1,bool)),X62),s(X61,X63))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X64)) ) | $spl6),
  inference(superposition,[],[f649,f3827])).
fof(f4730,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X57,fun(sK1,bool)),X58),s('type.semanticPrimitives.error_result'(X57),'const.semanticPrimitives.Rraise_1'(s(X57,X59))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(bool,X60) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X57,fun(sK1,bool)),X58),s(X57,X59))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f259,f3827])).
fof(f4725,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X37,fun(sK1,bool)),X38),s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,X39))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | ~p(s(bool,X40)) | s(bool,X40) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X37,fun(sK1,bool)),X38),s(X37,X39))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f171,f3827])).
fof(f4723,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X29,fun(sK1,bool)),X30),s('type.semanticPrimitives.error_result'(X29),'const.semanticPrimitives.Rraise_1'(s(X29,X31))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | p(s(bool,X32)) | s(bool,X32) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X29,fun(sK1,bool)),X30),s(X29,X31))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f168,f3827])).
fof(f4722,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X26,fun(sK1,bool)),X27),s('type.semanticPrimitives.error_result'(X26),'const.semanticPrimitives.Rraise_1'(s(X26,X28))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X26,fun(sK1,bool)),X27),s(X26,X28))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f133,f3827])).
fof(f4721,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X23,fun(sK1,bool)),X24),s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rraise_1'(s(X23,X25))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(X23,fun(sK1,bool)),X24),s(X23,X25))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f132,f3827])).
fof(f4720,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X20,fun(sK1,bool)),X21),s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,X22))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X20,fun(sK1,bool)),X21),s(X20,X22))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f129,f3827])).
fof(f4719,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,'const.evalProps.exc_rel_3'(s(fun(X17,fun(sK1,bool)),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X17,fun(sK1,bool)),X18),s(X17,X19))),s(sK1,sK5))))) ) | $spl6),
  inference(superposition,[],[f128,f3827])).
fof(f4717,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X11,fun(sK1,bool)),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(sK1,sK5) = s(sK1,sK8(sK5,happ(s(fun(X11,fun(sK1,bool)),X12),s(X11,X13)),sK1))) ) | $spl6),
  inference(superposition,[],[f165,f3827])).
fof(f4716,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X7,fun(sK1,bool)),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X7,fun(sK1,bool)),X8),s(X7,X9))),s(sK1,X10)))) | s(sK1,sK5) != s(sK1,X10)) ) | $spl6),
  inference(superposition,[],[f164,f3827])).
fof(f4713,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),happ(s(fun(X0,fun(sK1,bool)),X1),s(X0,X2)),sK1))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f3818,f3827])).
fof(f4706,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X3,fun(sK1,bool)),X4),s(X3,X5))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,X5))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f3827,f126])).
fof(f4705,plain,(
  ( ! [X6,X7,X5] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(sK1,bool)),X7),s('type.semanticPrimitives.error_result'(X5),X6),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X5,fun(sK1,bool)),X7),s(X5,sK6(X6,X5)))),s(sK1,sK5))) | s('type.semanticPrimitives.error_result'(X5),X6) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X6,X5))))) ) | $spl6),
  inference(superposition,[],[f3827,f153])).
fof(f4668,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X11,fun(sK1,bool)),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X11,fun(sK1,bool)),X12),s(X11,X13))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f3827,f202])).
fof(f4667,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X7,fun(sK1,bool)),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,X10) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X7,fun(sK1,bool)),X8),s(X7,X9))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f3827,f202])).
fof(f4666,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X4,fun(sK1,bool)),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(X4,fun(sK1,bool)),X5),s(X4,X6))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f3827,f126])).
fof(f4663,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK1,bool),happ(s(fun(X0,fun(sK1,bool)),X1),s(X0,X2))),s(sK1,sK5))))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(sK1,bool)),X1),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f3827,f163])).
fof(f4621,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X101,bool)),X102),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,X103))))) = s(bool,sK8(happ(s(fun(X101,bool),happ(s(fun(sK1,fun(X101,bool)),X102),s(sK1,sK5))),s(X101,X103)),X104,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X104),s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X101,bool)),X102),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,X103))))))))) ) | $spl6),
  inference(forward_demodulation,[],[f4585,f3826])).
fof(f4585,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,happ(s(fun(bool,bool),X104),s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X101,bool)),X102),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,X103)))))))) | s(bool,happ(s(fun(X101,bool),happ(s(fun(sK1,fun(X101,bool)),X102),s(sK1,sK5))),s(X101,X103))) = s(bool,sK8(happ(s(fun(X101,bool),happ(s(fun(sK1,fun(X101,bool)),X102),s(sK1,sK5))),s(X101,X103)),X104,bool))) ) | $spl6),
  inference(superposition,[],[f165,f3826])).
fof(f4574,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') != s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X60,bool)),X61),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,X62))))) | s(bool,X63) = s(bool,happ(s(fun(X60,bool),happ(s(fun(sK1,fun(X60,bool)),X61),s(sK1,sK5))),s(X60,X62))) | s(bool,'T_0') = s(bool,X63)) ) | $spl6),
  inference(superposition,[],[f649,f3826])).
fof(f4601,plain,(
  ( ! [X11] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(sK1,bool)),X11),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),happ(s(fun(sK1,fun(sK1,bool)),X11),s(sK1,sK5)),sK1))) ) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f4559,f186])).
fof(f4559,plain,(
  ( ! [X11] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(sK1,bool)),X11),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5)))))) | s(sK1,sK5) = s(sK1,sK8(sK6(sK3,sK1),happ(s(fun(sK1,fun(sK1,bool)),X11),s(sK1,sK5)),sK1))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f3818,f3826])).
fof(f4558,plain,(
  ( ! [X10] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun('type.semanticPrimitives.error_result'(sK1),bool)),X10),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)))))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),happ(s(fun(sK1,fun('type.semanticPrimitives.error_result'(sK1),bool)),X10),s(sK1,sK5)),'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(superposition,[],[f3850,f3826])).
fof(f4600,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X7,bool)),X8),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9)))))) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X7,bool)),X8),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,sK8(X9,happ(s(fun(sK1,fun(X7,bool)),X8),s(sK1,sK5)),X7)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f4557,f3826])).
fof(f4557,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X7,bool)),X8),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,sK8(X9,happ(s(fun(sK1,fun(X7,bool)),X8),s(sK1,sK5)),X7))))))) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(sK1,fun(X7,bool)),X8),s(sK1,sK5))),s(X7,X9))))) ) | $spl6),
  inference(superposition,[],[f166,f3826])).
fof(f4556,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X4,bool)),X5),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6)))))) | s(X4,X6) = s(X4,sK8(X6,happ(s(fun(sK1,fun(X4,bool)),X5),s(sK1,sK5)),X4))) ) | $spl6),
  inference(superposition,[],[f165,f3826])).
fof(f4599,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X0,bool)),X1),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3)))))) | ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X0,bool)),X1),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))) | s(X0,X2) != s(X0,X3)) ) | $spl6),
  inference(forward_demodulation,[],[f4555,f3826])).
fof(f4555,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X0,bool)),X1),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(sK1,fun(X0,bool)),X1),s(sK1,sK5))),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) ) | $spl6),
  inference(superposition,[],[f164,f3826])).
fof(f4551,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X10,bool)),X11),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X12))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),happ(s(fun(sK1,fun(X10,bool)),X11),s(sK1,sK5))),s(X10,X12)))) ) | $spl6),
  inference(superposition,[],[f3826,f202])).
fof(f4550,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))) | s(bool,X9) = s(bool,happ(s(fun(X6,bool),happ(s(fun(sK1,fun(X6,bool)),X7),s(sK1,sK5))),s(X6,X8))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f3826,f202])).
fof(f4549,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X3,bool)),X4),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,X5))))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),happ(s(fun(sK1,fun(X3,bool)),X4),s(sK1,sK5))),s(X3,X5)))) ) | $spl6),
  inference(superposition,[],[f3826,f126])).
fof(f4505,plain,(
  ( ! [X6] : (s(sK1,sK5) = s(sK1,sK8(sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1),X6,sK1)) | ~p(s(bool,happ(s(fun(sK1,bool),X6),s(sK1,sK5))))) ) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f4485,f4413])).
fof(f4485,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK1,bool),X6),s(sK1,sK5)))) | s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) = s(sK1,sK8(sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1),X6,sK1))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f165,f4413])).
fof(f4472,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X1,'type.semanticPrimitives.error_result'(sK1))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X1),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(subsumption_resolution,[],[f4467,f119])).
fof(f4467,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X1,'type.semanticPrimitives.error_result'(sK1))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X1),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f3850,f202])).
fof(f4465,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK6('const.semanticPrimitives.Rraise_1'(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0)),fun('type.semanticPrimitives.error_result'(sK1),bool)),'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(superposition,[],[f3850,f1323])).
fof(f4421,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1))) ) | $spl6),
  inference(superposition,[],[f162,f4410])).
fof(f4413,plain,(
  s(sK1,sK5) = s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) | ($spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f4410,f3812])).
fof(f4411,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.error_result'(sK1),X2) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),X2)),'type.semanticPrimitives.error_result'(sK1)),sK1)))) | s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X2)) ) | $spl6),
  inference(forward_demodulation,[],[f4407,f1323])).
fof(f4407,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X2) | s('type.semanticPrimitives.error_result'(sK1),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),X2)),'type.semanticPrimitives.error_result'(sK1))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),X2)),'type.semanticPrimitives.error_result'(sK1)),sK1))))) ) | $spl6),
  inference(superposition,[],[f3828,f1323])).
fof(f4410,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)))) | $spl6),
  inference(forward_demodulation,[],[f4402,f1323])).
fof(f4402,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)))) | $spl6),
  inference(unit_resulting_resolution,[],[f1323,f3828])).
fof(f4260,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK1,fun(X85,bool)),X86),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X87)))),X88,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0'))))) ) | $spl6),
  inference(forward_demodulation,[],[f4240,f3825])).
fof(f4240,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0')))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X85,bool)),X86),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X87))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK1,fun(X85,bool)),X86),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X87)))),X88,bool))) ) | $spl6),
  inference(superposition,[],[f165,f3825])).
fof(f4235,plain,(
  ( ! [X66,X64,X67,X65] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X67))) | s(bool,X67) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X64,bool)),X65),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X64),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X66)))))) ) | $spl6),
  inference(superposition,[],[f162,f3825])).
fof(f4252,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,X9) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f4213,f188])).
fof(f4213,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X9) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f3825,f202])).
fof(f4212,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X0,bool)),X2),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X0),X1))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) ) | $spl6),
  inference(superposition,[],[f3825,f153])).
fof(f4108,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X82,fun(sK1,bool)),X83),s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X84))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X85,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X85),s(bool,'F_0'))))) ) | $spl6),
  inference(forward_demodulation,[],[f4089,f3824])).
fof(f4089,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X85),s(bool,'F_0')))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X82,fun(sK1,bool)),X83),s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X84))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X82,fun(sK1,bool)),X83),s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X84))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X85,bool))) ) | $spl6),
  inference(superposition,[],[f165,f3824])).
fof(f4084,plain,(
  ( ! [X61,X64,X62,X63] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X64))) | s(bool,X64) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X61,fun(sK1,bool)),X62),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X63))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f162,f3824])).
fof(f4100,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,X9) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(sK1,bool)),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f4063,f188])).
fof(f4063,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X9) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(sK1,bool)),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f3824,f202])).
fof(f4062,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(sK1,bool)),X2),s('type.semanticPrimitives.error_result'(X0),X1),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK6(X1,X0))))) ) | $spl6),
  inference(superposition,[],[f3824,f153])).
fof(f3886,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X0) | s(sK1,sK5) = s(sK1,sK6(X0,sK1)) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X0,sK1))))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f3812,f153])).
fof(f3885,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,X0)),sK1))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f3812,f1323])).
fof(f3850,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X25),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X25,'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(forward_demodulation,[],[f3835,f186])).
fof(f3835,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X25),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X25,'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(superposition,[],[f165,f186])).
fof(f3828,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X16) | s('type.semanticPrimitives.error_result'(sK1),X16) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X16,sK1))))) ) | $spl6),
  inference(superposition,[],[f925,f186])).
fof(f3827,plain,(
  ( ! [X14,X12,X13] : (s(bool,happ(s(fun(sK1,bool),happ(s(fun(X12,fun(sK1,bool)),X13),s(X12,X14))),s(sK1,sK5))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X12,fun(sK1,bool)),X13),s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X14))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f175,f186])).
fof(f3826,plain,(
  ( ! [X10,X11,X9] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X9,bool)),X10),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))))) = s(bool,happ(s(fun(X9,bool),happ(s(fun(sK1,fun(X9,bool)),X10),s(sK1,sK5))),s(X9,X11)))) ) | $spl6),
  inference(superposition,[],[f175,f186])).
fof(f3825,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8)))))) ) | $spl6),
  inference(superposition,[],[f174,f186])).
fof(f3824,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X5))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f173,f186])).
fof(f3812,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,X0)) ) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f3811,f982])).
fof(f3811,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,sK6(sK3,sK1))) ) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f3792,f186])).
fof(f3792,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,sK6(sK3,sK1))) ) | $spl14),
  inference(superposition,[],[f162,f982])).
fof(f3744,plain,(
  ( ! [X25] : (s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X25,'type.semanticPrimitives.error_result'(sK1))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X25),s('type.semanticPrimitives.error_result'(sK1),sK3))))) ) | $spl6),
  inference(forward_demodulation,[],[f3716,f186])).
fof(f3716,plain,(
  ( ! [X25] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X25),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X25,'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(superposition,[],[f165,f186])).
fof(f3708,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X16) | s('type.semanticPrimitives.error_result'(sK1),X16) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X16,sK1))))) ) | $spl6),
  inference(superposition,[],[f925,f186])).
fof(f3707,plain,(
  ( ! [X14,X12,X13] : (s(bool,happ(s(fun(sK1,bool),happ(s(fun(X12,fun(sK1,bool)),X13),s(X12,X14))),s(sK1,sK5))) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X12,fun(sK1,bool)),X13),s('type.semanticPrimitives.error_result'(X12),'const.semanticPrimitives.Rraise_1'(s(X12,X14))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f175,f186])).
fof(f3706,plain,(
  ( ! [X10,X11,X9] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X9,bool)),X10),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))))) = s(bool,happ(s(fun(X9,bool),happ(s(fun(sK1,fun(X9,bool)),X10),s(sK1,sK5))),s(X9,X11)))) ) | $spl6),
  inference(superposition,[],[f175,f186])).
fof(f3705,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK1,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(sK1),sK3),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8)))))) ) | $spl6),
  inference(superposition,[],[f174,f186])).
fof(f3704,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X3,fun(sK1,bool)),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X5))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl6),
  inference(superposition,[],[f173,f186])).
fof(f3702,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,X0)) ) | $spl6),
  inference(superposition,[],[f162,f186])).
fof(f3625,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool)) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f3207,f126])).
fof(f3622,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool))) ) | $spl0),
  inference(superposition,[],[f3207,f1902])).
fof(f3668,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X41,fun(X42,bool)),X43),s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rraise_1'(s(X41,X44))),s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,X45)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3616,f1323])).
fof(f3616,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) != s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X41,fun(X42,bool)),X43),s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rraise_1'(s(X41,X44))),s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,X45)))))) ) | $spl0),
  inference(superposition,[],[f3207,f175])).
fof(f3660,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,X0) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl0),
  inference(condensation,[],[f3659])).
fof(f3659,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3606])).
fof(f3606,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool))) ) | $spl0),
  inference(superposition,[],[f3207,f1902])).
fof(f3592,plain,(
  s(bool,'T_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | $spl0),
  inference(unit_resulting_resolution,[],[f188,f3207])).
fof(f3590,plain,(
  ( ! [X80,X81,X82] : (s(bool,X81) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool)) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) ) | $spl0),
  inference(subsumption_resolution,[],[f3566,f187])).
fof(f3566,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'F_0')) | s(bool,X81) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool)) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) ) | $spl0),
  inference(superposition,[],[f3475,f1902])).
fof(f3589,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X79) != s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) | s(bool,'F_0') = s(bool,X79) | s(X76,X78) = s(X76,sK8(X78,X77,X76))) ) | $spl0),
  inference(subsumption_resolution,[],[f3565,f187])).
fof(f3565,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,'F_0')) | s(bool,X79) != s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) | s(bool,'F_0') = s(bool,X79) | s(X76,X78) = s(X76,sK8(X78,X77,X76))) ) | $spl0),
  inference(superposition,[],[f3475,f1741])).
fof(f3587,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X68) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X63,fun(X64,bool)),X65),s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rraise_1'(s(X63,X66))),s('type.semanticPrimitives.error_result'(X64),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X67))))) | s(bool,'F_0') = s(bool,X68)) ) | $spl0),
  inference(subsumption_resolution,[],[f3562,f187])).
fof(f3562,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,'F_0')) | s(bool,X68) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X63,fun(X64,bool)),X65),s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rraise_1'(s(X63,X66))),s('type.semanticPrimitives.error_result'(X64),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X67))))) | s(bool,'F_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f3475,f174])).
fof(f3586,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X57,fun(X58,bool)),X59),s('type.semanticPrimitives.error_result'(X57),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X60))),s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl0),
  inference(subsumption_resolution,[],[f3561,f187])).
fof(f3561,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (p(s(bool,'F_0')) | s(bool,X62) != s(bool,'const.evalProps.exc_rel_3'(s(fun(X57,fun(X58,bool)),X59),s('type.semanticPrimitives.error_result'(X57),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X60))),s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl0),
  inference(superposition,[],[f3475,f173])).
fof(f3583,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool))) ) | $spl0),
  inference(subsumption_resolution,[],[f3552,f187])).
fof(f3552,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'F_0')) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool))) ) | $spl0),
  inference(superposition,[],[f3475,f1902])).
fof(f3578,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) | $spl0),
  inference(subsumption_resolution,[],[f3546,f187])).
fof(f3546,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) | $spl0),
  inference(superposition,[],[f3475,f1620])).
fof(f3572,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f3540,f187])).
fof(f3540,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f3475,f126])).
fof(f3195,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f1902])).
fof(f3207,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3189])).
fof(f3189,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f1902])).
fof(f3521,plain,(
  ( ! [X116,X114,X115] : (s(bool,X114) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool),X115,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X115),s(bool,'F_0'))) | s(bool,X114) = s(bool,X116) | s(bool,'F_0') = s(bool,X116)) ) | $spl0),
  inference(forward_demodulation,[],[f3186,f1323])).
fof(f3186,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X115),s(bool,'F_0'))) | s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool)) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool),X115,bool)) | s(bool,X114) = s(bool,X116) | s(bool,'F_0') = s(bool,X116)) ) | $spl0),
  inference(superposition,[],[f1741,f1902])).
fof(f3520,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool),X81,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0')))) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) ) | $spl0),
  inference(forward_demodulation,[],[f3177,f1323])).
fof(f3177,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0')))) | s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool)) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool),X81,bool)) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) ) | $spl0),
  inference(superposition,[],[f165,f1902])).
fof(f3172,plain,(
  ( ! [X66,X64,X65] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X65))) | s(bool,X65) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool)) | s(bool,X64) = s(bool,X66) | s(bool,'F_0') = s(bool,X66)) ) | $spl0),
  inference(superposition,[],[f162,f1902])).
fof(f3516,plain,(
  ( ! [X62,X63] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) | s(bool,X62) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl0),
  inference(subsumption_resolution,[],[f3171,f167])).
fof(f3171,plain,(
  ( ! [X62,X63] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rraise_1'(s(bool,X62)),bool)))) | s(bool,X62) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl0),
  inference(superposition,[],[f153,f1902])).
fof(f3515,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X57) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool)) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,X56) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) ) | $spl0),
  inference(subsumption_resolution,[],[f3169,f187])).
fof(f3169,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'F_0')) | s(bool,X57) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool)) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,X56) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) ) | $spl0),
  inference(superposition,[],[f2437,f1902])).
fof(f3148,plain,(
  ( ! [X187,X185,X188,X186] : (s(X187,f(s(bool,'F_0'))) != s(X187,g(s(bool,'F_0'))) | s(fun(bool,X187),X185) = s(fun(bool,X187),X186) | s(bool,X188) = s(bool,sK9(X185,X186,X187,bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X188)),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f3145,plain,(
  ( ! [X178,X179] : (s('type.semanticPrimitives.error_result'(bool),X178) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X178) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X178,bool)))) | s(bool,X179) = s(bool,sK6(X178,bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X179)),bool))) ) | $spl0),
  inference(superposition,[],[f153,f1902])).
fof(f3144,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,'F_0') != s(bool,f(s(X174,sK9(X175,X176,bool,X174)))) | s(fun(X174,bool),X175) = s(fun(X174,bool),X176) | s(bool,X177) = s(bool,g(s(X174,sK9(X175,X176,bool,X174)))) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X177)),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f3143,plain,(
  ( ! [X169,X170,X168] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X169),s(bool,'F_0'))) | s(bool,X168) = s(bool,sK8(X168,X169,bool)) | s(bool,X168) = s(bool,X170) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X170)),bool))) ) | $spl0),
  inference(superposition,[],[f1741,f1902])).
fof(f3131,plain,(
  ( ! [X125,X126,X124] : (s(bool,'F_0') != s(bool,X125) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X125))) | s(bool,X124) = s(bool,X126) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X126)),bool))) ) | $spl0),
  inference(superposition,[],[f163,f1902])).
fof(f3129,plain,(
  ( ! [X118,X120,X119] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X119))) | s(bool,X118) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X120)),bool))) ) | $spl0),
  inference(superposition,[],[f162,f1902])).
fof(f3091,plain,(
  ( ! [X123,X124,X122] : (s(bool,X122) = s(bool,happ(s(fun(bool,bool),X123),s(bool,X122))) | s(bool,'F_0') = s(bool,sK8('F_0',X123,bool)) | s(bool,X122) = s(bool,X124) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X124)),bool))) ) | $spl0),
  inference(superposition,[],[f1741,f1902])).
fof(f3086,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X100) = s(bool,'const.evalProps.exc_rel_3'(s(fun(bool,fun(X101,bool)),X102),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))),s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X103))))) | s(bool,X100) = s(bool,X104) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X104)),bool))) ) | $spl0),
  inference(superposition,[],[f174,f1902])).
fof(f3085,plain,(
  ( ! [X99,X97,X95,X98,X96] : (s(bool,X95) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X96,fun(bool,bool)),X97),s('type.semanticPrimitives.error_result'(X96),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X98))),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95))))) | s(bool,X95) = s(bool,X99) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X99)),bool))) ) | $spl0),
  inference(superposition,[],[f173,f1902])).
fof(f3081,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) != s(bool,X81) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X81))) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X82)),bool))) ) | $spl0),
  inference(superposition,[],[f163,f1902])).
fof(f3057,plain,(
  ( ! [X12,X11] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11))) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X12)),bool))) ) | $spl0),
  inference(superposition,[],[f694,f1902])).
fof(f3056,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X7) = s(bool,X10) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool))) ) | $spl0),
  inference(superposition,[],[f649,f1902])).
fof(f3055,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X6)),bool))) ) | $spl0),
  inference(superposition,[],[f188,f1902])).
fof(f3052,plain,(
  ( ! [X210,X208,X211,X209] : (s(X210,f(s(bool,X211))) != s(X210,g(s(bool,X211))) | s(fun(bool,X210),X208) = s(fun(bool,X210),X209) | s(bool,'F_0') = s(bool,sK9(X208,X209,X210,bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X211)),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f3049,plain,(
  ( ! [X202,X201] : (s('type.semanticPrimitives.error_result'(bool),X201) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X202))) | s('type.semanticPrimitives.error_result'(bool),X201) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X201,bool)))) | s(bool,'F_0') = s(bool,sK6(X201,bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X202)),bool))) ) | $spl0),
  inference(superposition,[],[f153,f1902])).
fof(f3048,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,X200) != s(bool,f(s(X197,sK9(X198,X199,bool,X197)))) | s(fun(X197,bool),X198) = s(fun(X197,bool),X199) | s(bool,'F_0') = s(bool,g(s(X197,sK9(X198,X199,bool,X197)))) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X200)),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f3035,plain,(
  ( ! [X146,X147,X148] : (s(bool,X147) != s(bool,X148) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X148))) | s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X147)),bool))) ) | $spl0),
  inference(superposition,[],[f163,f1902])).
fof(f3033,plain,(
  ( ! [X142,X140,X141] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X141))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142))) | s(bool,X140) = s(bool,X142) | s(bool,'F_0') = s(bool,X140) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X141)),bool))) ) | $spl0),
  inference(superposition,[],[f162,f1902])).
fof(f3476,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) != s(bool,X75) | s(bool,X75) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X74)),bool))) ) | $spl0),
  inference(subsumption_resolution,[],[f3012,f188])).
fof(f3012,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) != s(bool,X75) | s(bool,X75) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X74)),bool))) ) | $spl0),
  inference(superposition,[],[f2088,f1902])).
fof(f3475,plain,(
  ( ! [X64,X65] : (p(s(bool,X65)) | s(bool,X64) != s(bool,X65) | s(bool,'F_0') = s(bool,X64)) ) | $spl0),
  inference(forward_demodulation,[],[f3474,f1323])).
fof(f3474,plain,(
  ( ! [X64,X65] : (s(bool,X64) != s(bool,X65) | p(s(bool,X65)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool))) ) | $spl0),
  inference(subsumption_resolution,[],[f3009,f188])).
fof(f3009,plain,(
  ( ! [X64,X65] : (s(bool,X64) != s(bool,X65) | p(s(bool,X65)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool))) ) | $spl0),
  inference(superposition,[],[f128,f1902])).
fof(f3470,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f754,f756,f759,f760,f762,f803,f805,f809,f811,f828,f873,f875,f887,f890,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f937,f938,f940,f922,f878,f808,f720,f694,f180,f990,f1026,f1049,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1605,f1607,f1619,f1620,f1737,f1739,f1741,f1745,f1656,f1748,f1749,f1750,f1751,f1752,f1686,f1767,f1769,f1770,f1797,f1820,f1862,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f1899,f1902,f1926,f1947,f2014,f2016,f2018,f2019,f2026,f2027,f2029,f2030,f2088,f2089,f2108,f2110,f2111,f2112,f2113,f2437,f2444,f2148,f2150,f2152,f2153,f2160,f2161,f2163,f2164,f2487,f2490,f2553,f2556,f2580,f2594,f2679,f2596,f2597,f2598,f2680,f2602,f2603,f2606,f2608,f2610,f2681,f2682,f2683,f2684,f2685,f2617,f2687,f2619,f2620,f2688,f2691,f2692,f2656,f2658,f2660,f2661,f2737,f2738,f2727,f2782,f2783,f2785,f3347,f3340,f3326,f3376,f2840,f2846,f2847,f2873,f3392,f2900,f3404,f2965,f2967,f2979,f2980,f2981,f2984,f3469])).
fof(f3469,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(condensation,[],[f3468])).
fof(f3468,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool))) ) | $spl0),
  inference(subsumption_resolution,[],[f2998,f649])).
fof(f2998,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool))) ) | $spl0),
  inference(superposition,[],[f649,f1902])).
fof(f2984,plain,(
  ( ! [X210,X208,X211,X209] : (s(X210,f(s(bool,X211))) != s(X210,g(s(bool,X211))) | s(fun(bool,X210),X208) = s(fun(bool,X210),X209) | s(bool,'F_0') = s(bool,X211) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,sK9(X208,X209,X210,bool))),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f2981,plain,(
  ( ! [X202,X201] : (s('type.semanticPrimitives.error_result'(bool),X201) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X202))) | s('type.semanticPrimitives.error_result'(bool),X201) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X201,bool)))) | s(bool,'F_0') = s(bool,X202) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,sK6(X201,bool))),bool))) ) | $spl0),
  inference(superposition,[],[f153,f1902])).
fof(f2980,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,X200) != s(bool,f(s(X197,sK9(X198,X199,bool,X197)))) | s(fun(X197,bool),X198) = s(fun(X197,bool),X199) | s(bool,'F_0') = s(bool,X200) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,g(s(X197,sK9(X198,X199,bool,X197))))),bool))) ) | $spl0),
  inference(superposition,[],[f172,f1902])).
fof(f2979,plain,(
  ( ! [X191,X192,X190] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X192),s(bool,X191))) | s(bool,X190) = s(bool,sK8(X190,X192,bool)) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X190)),bool))) ) | $spl0),
  inference(superposition,[],[f1741,f1902])).
fof(f2967,plain,(
  ( ! [X146,X147,X148] : (s(bool,X147) != s(bool,X148) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X148))) | s(bool,'F_0') = s(bool,X147) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X146)),bool))) ) | $spl0),
  inference(superposition,[],[f163,f1902])).
fof(f2965,plain,(
  ( ! [X142,X140,X141] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X141))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142))) | s(bool,X140) = s(bool,X142) | s(bool,'F_0') = s(bool,X141) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X140)),bool))) ) | $spl0),
  inference(superposition,[],[f162,f1902])).
fof(f3404,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f2907,f188])).
fof(f2907,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f1902,f202])).
fof(f2900,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X70,bool),happ(s(fun(X69,fun(X70,bool)),X71),s(X69,X72))),s(X70,X73)))),bool)) | s(bool,X74) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X69,fun(X70,bool)),X71),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X72))),s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,X73))))) | s(bool,'F_0') = s(bool,X74)) ) | $spl0),
  inference(superposition,[],[f1902,f175])).
fof(f3392,plain,(
  ( ! [X61,X60] : (s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X61)),bool)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(subsumption_resolution,[],[f2875,f188])).
fof(f2875,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool)) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X61)),bool)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(superposition,[],[f1902,f1419])).
fof(f2873,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(X50,bool),happ(s(fun(X49,fun(X50,bool)),X51),s(X49,X52))),s(X50,X53))) | s(bool,X54) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X49,fun(X50,bool)),X51),s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,X52))),s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rraise_1'(s(X50,X53))))) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X54)),bool))) ) | $spl0),
  inference(superposition,[],[f1902,f175])).
fof(f2847,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X6)),bool)) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f1902,f1620])).
fof(f2846,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X3)),bool))) ) | $spl0),
  inference(superposition,[],[f1902,f1902])).
fof(f2840,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (s(bool,X71) = s(bool,happ(s(fun(X67,bool),happ(s(fun(X66,fun(X67,bool)),X68),s(X66,X69))),s(X67,X70))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X66,fun(X67,bool)),X68),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X69))),s('type.semanticPrimitives.error_result'(X67),'const.semanticPrimitives.Rraise_1'(s(X67,X70))))) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X71)),bool))) ) | $spl0),
  inference(superposition,[],[f1902,f175])).
fof(f3376,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool)) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl0),
  inference(subsumption_resolution,[],[f2830,f188])).
fof(f2830,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool)) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl0),
  inference(superposition,[],[f1902,f202])).
fof(f3326,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))),bool)) | s(X0,X2) = s(X0,sK8(X2,X1,X0))) ) | $spl0),
  inference(condensation,[],[f3325])).
fof(f3325,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))),bool)) | s(X76,X78) = s(X76,sK8(X78,X77,X76))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2813])).
fof(f2813,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))),bool)) | s(X76,X78) = s(X76,sK8(X78,X77,X76))) ) | $spl0),
  inference(superposition,[],[f1902,f1741])).
fof(f3340,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(condensation,[],[f3339])).
fof(f3339,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool)) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2800])).
fof(f2800,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool)) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool))) ) | $spl0),
  inference(superposition,[],[f1902,f1902])).
fof(f3347,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(condensation,[],[f3346])).
fof(f3346,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2788])).
fof(f2788,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1902,f126])).
fof(f2785,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X88) = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))) | s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(X85,X87) = s(X85,sK8(X87,X86,X85))) ) | $spl0),
  inference(subsumption_resolution,[],[f2764,f187])).
fof(f2764,plain,(
  ( ! [X88,X87,X85,X89,X86] : (p(s(bool,'F_0')) | s(bool,X88) = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))) | s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(X85,X87) = s(X85,sK8(X87,X86,X85))) ) | $spl0),
  inference(superposition,[],[f2437,f1741])).
fof(f2783,plain,(
  ( ! [X70,X74,X72,X71,X69,X75,X73] : (s(bool,X74) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X69,fun(X70,bool)),X71),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X72))),s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X73))))) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75)) ) | $spl0),
  inference(subsumption_resolution,[],[f2761,f187])).
fof(f2761,plain,(
  ( ! [X70,X74,X72,X71,X69,X75,X73] : (p(s(bool,'F_0')) | s(bool,X74) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X69,fun(X70,bool)),X71),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X72))),s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X73))))) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75)) ) | $spl0),
  inference(superposition,[],[f2437,f174])).
fof(f2782,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,X67) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X62,fun(X63,bool)),X64),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X65))),s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rraise_1'(s(X63,X66))))) | s(bool,X67) = s(bool,X68) | s(bool,'F_0') = s(bool,X68)) ) | $spl0),
  inference(subsumption_resolution,[],[f2760,f187])).
fof(f2760,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (p(s(bool,'F_0')) | s(bool,X67) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X62,fun(X63,bool)),X64),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X65))),s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rraise_1'(s(X63,X66))))) | s(bool,X67) = s(bool,X68) | s(bool,'F_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f2437,f173])).
fof(f2661,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,happ(s(fun(bool,bool),X114),s(bool,'F_0')))) | s(bool,happ(s(fun(X111,bool),X112),s(X111,X113))) = s(bool,sK8(happ(s(fun(X111,bool),X112),s(X111,X113)),X114,bool)) | s(X111,X113) = s(X111,sK8(X113,X112,X111))) ) | $spl0),
  inference(superposition,[],[f165,f1741])).
fof(f2660,plain,(
  ( ! [X109,X107,X110,X108,X106] : (~p(s(bool,happ(s(fun(bool,bool),X109),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X109),s(bool,X110)))) | s(bool,X110) != s(bool,happ(s(fun(X106,bool),X107),s(X106,X108))) | s(X106,X108) = s(X106,sK8(X108,X107,X106))) ) | $spl0),
  inference(superposition,[],[f164,f1741])).
fof(f2658,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'F_0') != s(bool,X101) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100))))) | s(X98,X100) = s(X98,sK8(X100,X99,X98))) ) | $spl0),
  inference(superposition,[],[f163,f1741])).
fof(f2656,plain,(
  ( ! [X92,X90,X93,X91] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93))) | s(bool,X93) = s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))) | s(X90,X92) = s(X90,sK8(X92,X91,X90))) ) | $spl0),
  inference(superposition,[],[f162,f1741])).
fof(f2692,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) | s(X26,X28) = s(X26,sK8(X28,X27,X26))) ) | $spl0),
  inference(subsumption_resolution,[],[f2641,f187])).
fof(f2641,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) | s(X26,X28) = s(X26,sK8(X28,X27,X26))) ) | $spl0),
  inference(superposition,[],[f168,f1741])).
fof(f2691,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK8(X9,X8,X7)) = s(X7,sK8(sK8(X9,X8,X7),X8,X7))) ) | $spl0),
  inference(subsumption_resolution,[],[f2634,f187])).
fof(f2634,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK8(X9,X8,X7)) = s(X7,sK8(sK8(X9,X8,X7),X8,X7))) ) | $spl0),
  inference(superposition,[],[f166,f1741])).
fof(f2688,plain,(
  ( ! [X80,X81,X79] : (s('type.semanticPrimitives.abort',X79) = s('type.semanticPrimitives.abort',sK8(sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X79)),X80),X81,'type.semanticPrimitives.abort')) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',X79)))) ) | $spl0),
  inference(forward_demodulation,[],[f2621,f1261])).
fof(f2621,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',X79))) | s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X79)),X80)) = s('type.semanticPrimitives.abort',sK8(sK7('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X79)),X80),X81,'type.semanticPrimitives.abort'))) ) | $spl0),
  inference(superposition,[],[f1741,f1261])).
fof(f2620,plain,(
  ( ! [X78,X76,X77] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X76),bool),X78),s('type.semanticPrimitives.error_result'(X76),X77))) | s('type.semanticPrimitives.error_result'(X76),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X77,X76)))) = s('type.semanticPrimitives.error_result'(X76),sK8('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X77,X76))),X78,'type.semanticPrimitives.error_result'(X76))) | s('type.semanticPrimitives.error_result'(X76),X77) = s('type.semanticPrimitives.error_result'(X76),'const.semanticPrimitives.Rraise_1'(s(X76,sK6(X77,X76))))) ) | $spl0),
  inference(superposition,[],[f1741,f153])).
fof(f2619,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X73),bool),X75),s('type.semanticPrimitives.error_result'(X73),X74))) | s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rraise_1'(s(X73,sK6(X74,X73)))) = s('type.semanticPrimitives.error_result'(X73),sK8('const.semanticPrimitives.Rraise_1'(s(X73,sK6(X74,X73))),X75,'type.semanticPrimitives.error_result'(X73))) | s('type.semanticPrimitives.error_result'(X73),X74) = s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X74,X73))))) ) | $spl0),
  inference(superposition,[],[f1741,f153])).
fof(f2687,plain,(
  ( ! [X72,X71] : (s(bool,X71) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X71)),bool),X72,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X72),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X71)) ) | $spl0),
  inference(forward_demodulation,[],[f2618,f1323])).
fof(f2618,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X72),s(bool,'T_0'))) | s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X71)),bool)) = s(bool,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(bool,X71)),bool),X72,bool)) | s(bool,'F_0') = s(bool,X71)) ) | $spl0),
  inference(superposition,[],[f1741,f1419])).
fof(f2617,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X70),s(bool,'F_0'))) | s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) = s(bool,sK8(happ(s(fun(X67,bool),X68),s(X67,X69)),X70,bool)) | s(X67,X69) = s(X67,sK8(X69,X68,X67))) ) | $spl0),
  inference(superposition,[],[f1741,f1741])).
fof(f2685,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,happ(s(fun(X61,bool),happ(s(fun(X60,fun(X61,bool)),X62),s(X60,X63))),s(X61,X64))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X60,fun(X61,bool)),X62),s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,X63))),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X64)))),X65,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X65),s(bool,happ(s(fun(X61,bool),happ(s(fun(X60,fun(X61,bool)),X62),s(X60,X63))),s(X61,X64)))))) ) | $spl0),
  inference(forward_demodulation,[],[f2615,f175])).
fof(f2615,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X65),s(bool,happ(s(fun(X61,bool),happ(s(fun(X60,fun(X61,bool)),X62),s(X60,X63))),s(X61,X64))))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X60,fun(X61,bool)),X62),s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,X63))),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X64))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X60,fun(X61,bool)),X62),s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,X63))),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X64)))),X65,bool))) ) | $spl0),
  inference(superposition,[],[f1741,f175])).
fof(f2684,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X54,fun(X55,bool)),X56),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X57))),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X58)))),X59,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f2614,f174])).
fof(f2614,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'F_0'))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X54,fun(X55,bool)),X56),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X57))),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X58))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X54,fun(X55,bool)),X56),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X57))),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X58)))),X59,bool))) ) | $spl0),
  inference(superposition,[],[f1741,f174])).
fof(f2683,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X48,fun(X49,bool)),X50),s('type.semanticPrimitives.error_result'(X48),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X51))),s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,X52)))),X53,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X53),s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f2613,f173])).
fof(f2613,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X53),s(bool,'F_0'))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(X48,fun(X49,bool)),X50),s('type.semanticPrimitives.error_result'(X48),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X51))),s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,X52))))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(X48,fun(X49,bool)),X50),s('type.semanticPrimitives.error_result'(X48),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X51))),s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,X52)))),X53,bool))) ) | $spl0),
  inference(superposition,[],[f1741,f173])).
fof(f2682,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X47,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f2612,f990])).
fof(f2612,plain,(
  ( ! [X47] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0'))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X47,bool))) ) | $spl0),
  inference(superposition,[],[f1741,f990])).
fof(f2681,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X46,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,X45))) | s(bool,'F_0') = s(bool,X45)) ) | $spl0),
  inference(forward_demodulation,[],[f2611,f990])).
fof(f2611,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,X45))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X46,bool)) | s(bool,'F_0') = s(bool,X45)) ) | $spl0),
  inference(superposition,[],[f1741,f1605])).
fof(f2610,plain,(
  ( ! [X43,X44] : (s(bool,X43) = s(bool,happ(s(fun(bool,bool),X44),s(bool,X43))) | s(bool,'F_0') = s(bool,sK8('F_0',X44,bool)) | s(bool,'T_0') = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f1741,f126])).
fof(f2608,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))) | s(bool,'F_0') = s(bool,sK8('F_0',X38,bool)) | s(bool,X37) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) ) | $spl0),
  inference(superposition,[],[f1741,f1620])).
fof(f2606,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,X31))) | s(bool,'T_0') = s(bool,sK8('T_0',X32,bool)) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f1741,f202])).
fof(f2603,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,X23))) | s(bool,X22) = s(bool,sK8(X22,X24,bool)) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f1741,f1620])).
fof(f2602,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0'))) | s(bool,X19) = s(bool,sK8(X19,X20,bool)) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') = s(bool,X21)) ) | $spl0),
  inference(superposition,[],[f1741,f1620])).
fof(f2680,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,X17) = s(bool,sK8(X17,X18,bool)) | s(bool,'F_0') = s(bool,X17)) ) | $spl0),
  inference(forward_demodulation,[],[f2601,f990])).
fof(f2601,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s(bool,X17) = s(bool,sK8(X17,X18,bool)) | s(bool,'F_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f1741,f1605])).
fof(f2598,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,X9))) | s(bool,X8) = s(bool,sK8(X8,X10,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f1741,f202])).
fof(f2597,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | s(bool,X5) = s(bool,sK8(X5,X6,bool)) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f1741,f202])).
fof(f2596,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,X3) = s(bool,sK8(X3,X4,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1741,f126])).
fof(f2679,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0),X2,X0)) | s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) ) | $spl0),
  inference(forward_demodulation,[],[f2595,f1323])).
fof(f2595,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1))) | s(X0,sK6('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0)) = s(X0,sK8(sK6('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0),X2,X0))) ) | $spl0),
  inference(superposition,[],[f1741,f1323])).
fof(f2594,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK8(X2,sK6('const.semanticPrimitives.Rraise_1'(s(fun(X0,bool),X1)),fun(X0,bool)),X0))) ) | $spl0),
  inference(superposition,[],[f1741,f1323])).
fof(f2547,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(subsumption_resolution,[],[f2535,f119])).
fof(f2535,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(superposition,[],[f2524,f202])).
fof(f2532,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1)))) | $spl17),
  inference(unit_resulting_resolution,[],[f2524,f133])).
fof(f2533,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1)))) | $spl17),
  inference(unit_resulting_resolution,[],[f2524,f128])).
fof(f2490,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X66)),bool)) | s(bool,'F_0') = s(bool,X67)) ) | $spl0),
  inference(superposition,[],[f2088,f1323])).
fof(f2487,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,'T_0') != s(bool,happ(s(fun(X58,bool),happ(s(fun(X57,fun(X58,bool)),X59),s(X57,X60))),s(X58,X61))) | s(bool,X62) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X57,fun(X58,bool)),X59),s('type.semanticPrimitives.error_result'(X57),'const.semanticPrimitives.Rraise_1'(s(X57,X60))),s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl0),
  inference(superposition,[],[f2088,f175])).
fof(f2164,plain,(
  ( ! [X154,X152,X155,X153] : (s(X154,f(s(bool,'F_0'))) != s(X154,g(s(bool,'F_0'))) | s(fun(bool,X154),X152) = s(fun(bool,X154),X153) | s(bool,X155) = s(bool,sK9(X152,X153,X154,bool)) | s(bool,'F_0') = s(bool,X155)) ) | $spl0),
  inference(superposition,[],[f172,f1620])).
fof(f2163,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,happ(s(fun(bool,bool),X150),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X150),s(bool,X149)))) | s(bool,X151) = s(bool,sK8(X149,X150,bool)) | s(bool,'F_0') = s(bool,X151)) ) | $spl0),
  inference(superposition,[],[f166,f1620])).
fof(f2161,plain,(
  ( ! [X146,X145] : (s('type.semanticPrimitives.error_result'(bool),X145) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X145) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X145,bool)))) | s(bool,X146) = s(bool,sK6(X145,bool)) | s(bool,'F_0') = s(bool,X146)) ) | $spl0),
  inference(superposition,[],[f153,f1620])).
fof(f2160,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'F_0') != s(bool,f(s(X141,sK9(X142,X143,bool,X141)))) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(bool,X144) = s(bool,g(s(X141,sK9(X142,X143,bool,X141)))) | s(bool,'F_0') = s(bool,X144)) ) | $spl0),
  inference(superposition,[],[f172,f1620])).
fof(f2153,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),X111),s(bool,'F_0')))) | s(bool,X110) = s(bool,sK8(X110,X111,bool)) | s(bool,X110) = s(bool,X112) | s(bool,'F_0') = s(bool,X112)) ) | $spl0),
  inference(superposition,[],[f165,f1620])).
fof(f2152,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,happ(s(fun(bool,bool),X107),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X107),s(bool,X108)))) | s(bool,X106) != s(bool,X108) | s(bool,X106) = s(bool,X109) | s(bool,'F_0') = s(bool,X109)) ) | $spl0),
  inference(superposition,[],[f164,f1620])).
fof(f2150,plain,(
  ( ! [X101,X102,X100] : (s(bool,'F_0') != s(bool,X101) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101))) | s(bool,X100) = s(bool,X102) | s(bool,'F_0') = s(bool,X102)) ) | $spl0),
  inference(superposition,[],[f163,f1620])).
fof(f2148,plain,(
  ( ! [X94,X95,X96] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95))) | s(bool,X94) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'F_0') = s(bool,X96)) ) | $spl0),
  inference(superposition,[],[f162,f1620])).
fof(f2444,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(X85,bool),X86),s(X85,X87)))) | s(bool,X88) = s(bool,happ(s(fun(X85,bool),X86),s(X85,sK8(X87,X86,X85)))) | s(bool,'F_0') = s(bool,X88)) ) | $spl0),
  inference(subsumption_resolution,[],[f2144,f187])).
fof(f2144,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X85,bool),X86),s(X85,X87)))) | s(bool,X88) = s(bool,happ(s(fun(X85,bool),X86),s(X85,sK8(X87,X86,X85)))) | s(bool,'F_0') = s(bool,X88)) ) | $spl0),
  inference(superposition,[],[f166,f1620])).
fof(f2437,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f2121,f187])).
fof(f2121,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f168,f1620])).
fof(f2113,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X79) = s(bool,'const.evalProps.exc_rel_3'(s(fun(bool,fun(X80,bool)),X81),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X79))),s('type.semanticPrimitives.error_result'(X80),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X82))))) | s(bool,X79) = s(bool,X83) | s(bool,'F_0') = s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f174,f1620])).
fof(f2112,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X74) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X75,fun(bool,bool)),X76),s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X77))),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X74))))) | s(bool,X74) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl0),
  inference(superposition,[],[f173,f1620])).
fof(f2111,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,happ(s(fun(bool,bool),X70),s(bool,X69)))) | s(bool,'F_0') = s(bool,sK8('F_0',X70,bool)) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) ) | $spl0),
  inference(superposition,[],[f165,f1620])).
fof(f2110,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),X66),s(bool,X67)))) | s(bool,'F_0') != s(bool,X67) | s(bool,X65) = s(bool,X68) | s(bool,'F_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f164,f1620])).
fof(f2108,plain,(
  ( ! [X61,X59,X60] : (s(bool,X59) != s(bool,X60) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) | s(bool,X59) = s(bool,X61) | s(bool,'F_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f163,f1620])).
fof(f2089,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X7) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) ) | $spl0),
  inference(superposition,[],[f649,f1620])).
fof(f2088,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f188,f1620])).
fof(f2030,plain,(
  ( ! [X167,X169,X170,X168] : (s(X169,f(s(bool,X170))) != s(X169,g(s(bool,X170))) | s(fun(bool,X169),X167) = s(fun(bool,X169),X168) | s(bool,'F_0') = s(bool,X170) | s(bool,'F_0') = s(bool,sK9(X167,X168,X169,bool))) ) | $spl0),
  inference(superposition,[],[f172,f1620])).
fof(f2029,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,happ(s(fun(bool,bool),X165),s(bool,X166)))) | ~p(s(bool,happ(s(fun(bool,bool),X165),s(bool,X164)))) | s(bool,'F_0') = s(bool,X166) | s(bool,'F_0') = s(bool,sK8(X164,X165,bool))) ) | $spl0),
  inference(superposition,[],[f166,f1620])).
fof(f2027,plain,(
  ( ! [X161,X160] : (s('type.semanticPrimitives.error_result'(bool),X160) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X161))) | s('type.semanticPrimitives.error_result'(bool),X160) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X160,bool)))) | s(bool,'F_0') = s(bool,X161) | s(bool,'F_0') = s(bool,sK6(X160,bool))) ) | $spl0),
  inference(superposition,[],[f153,f1620])).
fof(f2026,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X159) != s(bool,f(s(X156,sK9(X157,X158,bool,X156)))) | s(fun(X156,bool),X157) = s(fun(X156,bool),X158) | s(bool,'F_0') = s(bool,X159) | s(bool,'F_0') = s(bool,g(s(X156,sK9(X157,X158,bool,X156))))) ) | $spl0),
  inference(superposition,[],[f172,f1620])).
fof(f2019,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,happ(s(fun(bool,bool),X126),s(bool,X125)))) | s(bool,X124) = s(bool,sK8(X124,X126,bool)) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') = s(bool,X124)) ) | $spl0),
  inference(superposition,[],[f165,f1620])).
fof(f2018,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,happ(s(fun(bool,bool),X122),s(bool,X121)))) | p(s(bool,happ(s(fun(bool,bool),X122),s(bool,X123)))) | s(bool,X120) != s(bool,X123) | s(bool,'F_0') = s(bool,X121) | s(bool,'F_0') = s(bool,X120)) ) | $spl0),
  inference(superposition,[],[f164,f1620])).
fof(f2016,plain,(
  ( ! [X116,X114,X115] : (s(bool,X115) != s(bool,X116) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X116))) | s(bool,'F_0') = s(bool,X115) | s(bool,'F_0') = s(bool,X114)) ) | $spl0),
  inference(superposition,[],[f163,f1620])).
fof(f2014,plain,(
  ( ! [X109,X110,X108] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X109))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X110))) | s(bool,X108) = s(bool,X110) | s(bool,'F_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X108)) ) | $spl0),
  inference(superposition,[],[f162,f1620])).
fof(f1947,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,'F_0') = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,X46) = s(bool,'const.evalProps.exc_rel_3'(s(fun(X41,fun(X42,bool)),X43),s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rraise_1'(s(X41,X44))),s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,X45))))) | s(bool,'F_0') = s(bool,X46)) ) | $spl0),
  inference(superposition,[],[f1620,f175])).
fof(f1926,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1620,f1620])).
fof(f1902,plain,(
  ( ! [X64,X65] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool))) ) | $spl0),
  inference(superposition,[],[f1620,f1323])).
fof(f1899,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,X60) = s(bool,happ(s(fun(X56,bool),happ(s(fun(X55,fun(X56,bool)),X57),s(X55,X58))),s(X56,X59))) | s(bool,'F_0') = s(bool,X60) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X55,fun(X56,bool)),X57),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,X58))),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X59)))))) ) | $spl0),
  inference(superposition,[],[f1620,f175])).
fof(f1862,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),happ(s(fun(X5,fun(X6,bool)),X7),s(X5,X8))),s(X6,X9))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X9)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1821,f990])).
fof(f1821,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X5,fun(X6,bool)),X7),s(X5,X8))),s(X6,X9))) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X5,fun(X6,bool)),X7),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X8))),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X9)))))) ) | $spl0),
  inference(superposition,[],[f175,f1605])).
fof(f1770,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X39,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'F_0') = s(bool,X38)) ) | $spl0),
  inference(forward_demodulation,[],[f1691,f990])).
fof(f1691,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X39,bool)) | s(bool,'F_0') = s(bool,X38)) ) | $spl0),
  inference(superposition,[],[f165,f1605])).
fof(f1769,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X37)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl0),
  inference(forward_demodulation,[],[f1690,f990])).
fof(f1690,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X37)))) | s(bool,X37) != s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,X35)) ) | $spl0),
  inference(superposition,[],[f164,f1605])).
fof(f1767,plain,(
  ( ! [X31,X32] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X32))) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X31)) ) | $spl0),
  inference(forward_demodulation,[],[f1688,f990])).
fof(f1688,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X32) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X32))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s(bool,'F_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f163,f1605])).
fof(f1686,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,X28) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,X27)) ) | $spl0),
  inference(superposition,[],[f162,f1605])).
fof(f1752,plain,(
  ( ! [X88,X87,X89] : (s(X89,f(s(bool,'T_0'))) != s(X89,g(s(bool,'T_0'))) | s(fun(bool,X89),X87) = s(fun(bool,X89),X88) | s(bool,'F_0') = s(bool,sK9(X87,X88,X89,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f1668,f990])).
fof(f1668,plain,(
  ( ! [X88,X87,X89] : (s(X89,f(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) != s(X89,g(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s(fun(bool,X89),X87) = s(fun(bool,X89),X88) | s(bool,'F_0') = s(bool,sK9(X87,X88,X89,bool))) ) | $spl0),
  inference(superposition,[],[f172,f1605])).
fof(f1751,plain,(
  ( ! [X85,X86] : (p(s(bool,happ(s(fun(bool,bool),X86),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X85)))) | s(bool,'F_0') = s(bool,sK8(X85,X86,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f1667,f990])).
fof(f1667,plain,(
  ( ! [X85,X86] : (p(s(bool,happ(s(fun(bool,bool),X86),s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) | ~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X85)))) | s(bool,'F_0') = s(bool,sK8(X85,X86,bool))) ) | $spl0),
  inference(superposition,[],[f166,f1605])).
fof(f1750,plain,(
  ( ! [X83] : (s('type.semanticPrimitives.error_result'(bool),X83) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'(bool),X83) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X83,bool)))) | s(bool,'F_0') = s(bool,sK6(X83,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f1665,f990])).
fof(f1665,plain,(
  ( ! [X83] : (s('type.semanticPrimitives.error_result'(bool),X83) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s('type.semanticPrimitives.error_result'(bool),X83) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X83,bool)))) | s(bool,'F_0') = s(bool,sK6(X83,bool))) ) | $spl0),
  inference(superposition,[],[f153,f1605])).
fof(f1749,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,f(s(X80,sK9(X81,X82,bool,X80)))) | s(fun(X80,bool),X81) = s(fun(X80,bool),X82) | s(bool,'F_0') = s(bool,g(s(X80,sK9(X81,X82,bool,X80))))) ) | $spl0),
  inference(forward_demodulation,[],[f1664,f990])).
fof(f1664,plain,(
  ( ! [X80,X81,X82] : (s(bool,f(s(X80,sK9(X81,X82,bool,X80)))) != s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(fun(X80,bool),X81) = s(fun(X80,bool),X82) | s(bool,'F_0') = s(bool,g(s(X80,sK9(X81,X82,bool,X80))))) ) | $spl0),
  inference(superposition,[],[f172,f1605])).
fof(f1748,plain,(
  ( ! [X64,X65] : (~p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'T_0')))) | s(bool,X64) = s(bool,sK8(X64,X65,bool)) | s(bool,'F_0') = s(bool,X64)) ) | $spl0),
  inference(forward_demodulation,[],[f1659,f990])).
fof(f1659,plain,(
  ( ! [X64,X65] : (~p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))))) | s(bool,X64) = s(bool,sK8(X64,X65,bool)) | s(bool,'F_0') = s(bool,X64)) ) | $spl0),
  inference(superposition,[],[f165,f1605])).
fof(f1656,plain,(
  ( ! [X57,X58] : (s(bool,X58) != s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X58))) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(superposition,[],[f163,f1605])).
fof(f1745,plain,(
  ( ! [X54,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X54))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl0),
  inference(forward_demodulation,[],[f1654,f990])).
fof(f1654,plain,(
  ( ! [X54,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X54))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl0),
  inference(superposition,[],[f162,f1605])).
fof(f1741,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,X47))) | s(X45,X47) = s(X45,sK8(X47,X46,X45))) ) | $spl0),
  inference(subsumption_resolution,[],[f1740,f119])).
fof(f1740,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | s(X45,X47) = s(X45,sK8(X47,X46,X45)) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,X47)))) ) | $spl0),
  inference(forward_demodulation,[],[f1649,f990])).
fof(f1649,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s(X45,X47) = s(X45,sK8(X47,X46,X45)) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,X47)))) ) | $spl0),
  inference(superposition,[],[f165,f1605])).
fof(f1739,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,happ(s(fun(X41,bool),X42),s(X41,X44)))) | s(X41,X43) != s(X41,X44) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1738,f119])).
fof(f1738,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X41,bool),X42),s(X41,X44)))) | s(X41,X43) != s(X41,X44) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))) ) | $spl0),
  inference(forward_demodulation,[],[f1648,f990])).
fof(f1648,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | p(s(bool,happ(s(fun(X41,bool),X42),s(X41,X44)))) | s(X41,X43) != s(X41,X44) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))) ) | $spl0),
  inference(superposition,[],[f164,f1605])).
fof(f1737,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.semanticPrimitives.abort',X34) = s('type.semanticPrimitives.abort',X35) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X31,fun(X32,bool)),X33),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X34))),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X35)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1736,f119])).
fof(f1736,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.abort',X34) = s('type.semanticPrimitives.abort',X35) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X31,fun(X32,bool)),X33),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X34))),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X35)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1646,f990])).
fof(f1646,plain,(
  ( ! [X35,X33,X31,X34,X32] : (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.abort',X34) = s('type.semanticPrimitives.abort',X35) | s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X31,fun(X32,bool)),X33),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X34))),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X35)))))) ) | $spl0),
  inference(superposition,[],[f176,f1605])).
fof(f1620,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1605,f1605])).
fof(f1619,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,X25)),bool))) ) | $spl0),
  inference(superposition,[],[f1605,f1323])).
fof(f1605,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(subsumption_resolution,[],[f1591,f119])).
fof(f1591,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | s(bool,X22) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f1557,f990])).
fof(f1049,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X11,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f1030,f990])).
fof(f1030,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0')))) | s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s(bool,sK8('const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)),X11,bool))) ) | $spl0),
  inference(superposition,[],[f165,f990])).
fof(f1026,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl0),
  inference(superposition,[],[f162,f990])).
fof(f990,plain,(
  s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | $spl0),
  inference(unit_resulting_resolution,[],[f180,f129])).
fof(f978,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X0) | s(sK1,sK5) = s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1))),sK1)) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X0,sK1))))) ) | $spl6),
  inference(subsumption_resolution,[],[f974,f167])).
fof(f974,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X0) | s(sK1,sK5) = s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1))),sK1)) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1)))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rraise_1'(s(sK1,sK6(X0,sK1))),sK1)))) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X0,sK1))))) ) | $spl6),
  inference(superposition,[],[f936,f153])).
fof(f972,plain,(
  s(sK1,sK5) = s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1)) | $spl6),
  inference(unit_resulting_resolution,[],[f186,f167,f936])).
fof(f936,plain,(
  ( ! [X11] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X11) | s(sK1,sK5) = s(sK1,sK6(X11,sK1)) | s('type.semanticPrimitives.error_result'(sK1),X11) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X11,sK1))))) ) | $spl6),
  inference(superposition,[],[f713,f153])).
fof(f913,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X0,'type.semanticPrimitives.error_result'(sK1))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f907,f119])).
fof(f907,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X0,'type.semanticPrimitives.error_result'(sK1))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f822,f202])).
fof(f822,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X0,'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(forward_demodulation,[],[f802,f186])).
fof(f802,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK1),bool),X0),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) = s('type.semanticPrimitives.error_result'(sK1),sK8('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),X0,'type.semanticPrimitives.error_result'(sK1)))) ) | $spl6),
  inference(superposition,[],[f165,f186])).
fof(f713,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,X0)) ) | $spl6),
  inference(superposition,[],[f162,f186])).
fof(f4799,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl14 | $spl17)),
  inference(global_subsumption,[],[f182,f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f713,f754,f756,f759,f760,f762,f822,f803,f805,f809,f811,f828,f873,f875,f887,f890,f913,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f936,f937,f938,f940,f972,f978,f976,f922,f878,f808,f720,f694,f990,f1026,f1049,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1605,f1607,f1619,f1620,f1737,f1739,f1741,f1745,f1656,f1748,f1749,f1750,f1751,f1752,f1686,f1767,f1769,f1770,f1797,f1820,f1862,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f1899,f1902,f1926,f1947,f2014,f2016,f2018,f2019,f2026,f2027,f2029,f2030,f2088,f2089,f2108,f2110,f2111,f2112,f2113,f2437,f2444,f2148,f2150,f2152,f2153,f2160,f2161,f2163,f2164,f2487,f2490,f2524,f2533,f2532,f2547,f2553,f2556,f2580,f2594,f2679,f2596,f2597,f2598,f2680,f2602,f2603,f2606,f2608,f2610,f2681,f2682,f2683,f2684,f2685,f2617,f2687,f2619,f2620,f2688,f2691,f2692,f2656,f2658,f2660,f2661,f2737,f2738,f2727,f2782,f2783,f2785,f3347,f3340,f3326,f3376,f2840,f2846,f2847,f2873,f3392,f2900,f3404,f2965,f2967,f2979,f2980,f2981,f2984,f3469,f3470,f3475,f3476,f3033,f3035,f3048,f3049,f3052,f3055,f3056,f3057,f3081,f3085,f3086,f3091,f3129,f3131,f3143,f3144,f3145,f3148,f3515,f3516,f3172,f3520,f3521,f3207,f3195,f3531,f3532,f3572,f3578,f3583,f3586,f3587,f3589,f3590,f3592,f3660,f3668,f3622,f3625,f3672,f3675,f3702,f3704,f3705,f3706,f3707,f3708,f3744,f180,f186,f982,f3812,f3818,f3824,f3825,f3826,f3827,f3828,f3850,f3885,f3886,f4062,f4100,f4084,f4108,f4212,f4252,f4235,f4260,f4410,f4411,f4413,f4421,f4465,f4472,f4505,f4506,f4515,f4549,f4550,f4551,f4599,f4556,f4600,f4558,f4601,f4574,f4621,f4663,f4666,f4667,f4668,f4705,f4706,f4713,f4716,f4717,f4719,f4720,f4721,f4722,f4723,f4725,f4730,f4731,f4732,f4733,f4737,f4739,f4741,f4742,f4753,f4754,f4755,f4756,f4757,f4759,f4764,f4765,f4766,f4767,f4771,f4773,f4775,f4776,f4797])).
fof(f4798,plain,(
  $false | ($spl0 | $spl6 | $spl14 | $spl17)),
  inference(global_subsumption,[],[f118,f117,f116,f119,f187,f126,f129,f128,f133,f132,f153,f161,f160,f163,f162,f166,f165,f164,f167,f171,f168,f172,f173,f174,f175,f177,f176,f188,f202,f259,f375,f441,f649,f650,f660,f699,f701,f703,f709,f713,f754,f756,f759,f760,f762,f822,f803,f805,f809,f811,f828,f873,f875,f887,f890,f913,f916,f915,f917,f919,f923,f925,f926,f927,f929,f932,f934,f936,f937,f938,f940,f972,f978,f976,f922,f878,f808,f720,f694,f990,f1026,f1049,f1053,f1066,f1076,f1081,f1111,f1099,f1118,f1140,f1146,f1147,f1178,f1165,f1185,f1213,f1214,f1215,f1217,f1219,f1220,f1222,f1242,f1261,f1265,f1283,f1285,f1287,f1288,f1290,f1292,f1298,f1305,f1319,f1323,f1343,f1345,f1347,f1350,f1378,f1380,f1388,f1415,f1417,f1419,f1476,f1442,f1486,f1446,f1450,f1555,f1557,f1558,f1527,f1559,f1561,f1562,f1572,f1598,f1605,f1607,f1619,f1620,f1737,f1739,f1741,f1745,f1656,f1748,f1749,f1750,f1751,f1752,f1686,f1767,f1769,f1770,f1797,f1820,f1862,f1822,f1824,f1826,f1840,f1841,f1844,f1876,f1899,f1902,f1926,f1947,f2014,f2016,f2018,f2019,f2026,f2027,f2029,f2030,f2088,f2089,f2108,f2110,f2111,f2112,f2113,f2437,f2444,f2148,f2150,f2152,f2153,f2160,f2161,f2163,f2164,f2487,f2490,f2524,f2533,f2532,f2547,f2553,f2556,f2580,f2594,f2679,f2596,f2597,f2598,f2680,f2602,f2603,f2606,f2608,f2610,f2681,f2682,f2683,f2684,f2685,f2617,f2687,f2619,f2620,f2688,f2691,f2692,f2656,f2658,f2660,f2661,f2737,f2738,f2727,f2782,f2783,f2785,f3347,f3340,f3326,f3376,f2840,f2846,f2847,f2873,f3392,f2900,f3404,f2965,f2967,f2979,f2980,f2981,f2984,f3469,f3470,f3475,f3476,f3033,f3035,f3048,f3049,f3052,f3055,f3056,f3057,f3081,f3085,f3086,f3091,f3129,f3131,f3143,f3144,f3145,f3148,f3515,f3516,f3172,f3520,f3521,f3207,f3195,f3531,f3532,f3572,f3578,f3583,f3586,f3587,f3589,f3590,f3592,f3660,f3668,f3622,f3625,f3672,f3675,f3702,f3704,f3705,f3706,f3707,f3708,f3744,f180,f186,f982,f3812,f3818,f3824,f3825,f3826,f3827,f3828,f3850,f3885,f3886,f4062,f4100,f4084,f4108,f4212,f4252,f4235,f4260,f4410,f4411,f4413,f4421,f4465,f4472,f4505,f4506,f4515,f4549,f4550,f4551,f4599,f4556,f4600,f4558,f4601,f4574,f4621,f4663,f4666,f4667,f4668,f4705,f4706,f4713,f4716,f4717,f4719,f4720,f4721,f4722,f4723,f4725,f4730,f4731,f4732,f4733,f4737,f4739,f4741,f4742,f4753,f4754,f4755,f4756,f4757,f4759,f4764,f4765,f4766,f4767,f4771,f4773,f4775,f4776,f4797])).
fof(f4793,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4712,f3867])).
fof(f3867,plain,(
  s(bool,'T_0') != s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | $spl1),
  inference(unit_resulting_resolution,[],[f181,f128])).
fof(f4712,plain,(
  s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | ($spl4 | $spl6)),
  inference(superposition,[],[f192,f3827])).
fof(f192,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5))) | $spl4),
  inference(unit_resulting_resolution,[],[f184,f129])).
fof(f4790,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4664,f3867])).
fof(f4664,plain,(
  s(bool,'T_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))) | ($spl4 | $spl6)),
  inference(superposition,[],[f3827,f192])).
fof(f4052,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl22),
  inference(cnf_transformation,[],[f4052_D])).
fof(f4052_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4054,plain,(
  s(bool,'T_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | $spl25),
  inference(cnf_transformation,[],[f4054_D])).
fof(f4054_D,plain,(
  s(bool,'T_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f3994,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl18),
  inference(equality_factoring,[],[f3673])).
fof(f3783,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f186,f3750])).
fof(f3750,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X2)))) ) | $spl12),
  inference(superposition,[],[f167,f980])).
fof(f3776,plain,(
  $false | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f3775,f187])).
fof(f3775,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl12)),
  inference(backward_demodulation,[],[f3752,f180])).
fof(f3752,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(sK1,bool)),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl12),
  inference(superposition,[],[f174,f980])).
fof(f3738,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3737,f119])).
fof(f3737,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f3736,f192])).
fof(f3736,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f3735,f1323])).
fof(f3735,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1))))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f3725,f167])).
fof(f3725,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1))))) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1)))) | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f3709])).
fof(f3709,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1))))) | s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1)))) | ($spl2 | $spl6)),
  inference(superposition,[],[f1007,f186])).
fof(f1007,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),X0) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(X0,sK1))))) | s('type.semanticPrimitives.error_result'(sK1),X0) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(X0,sK1))))) ) | $spl2),
  inference(superposition,[],[f182,f153])).
fof(f3734,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3733,f119])).
fof(f3733,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f3726,f192])).
fof(f3726,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f3701])).
fof(f3701,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl2 | $spl6)),
  inference(superposition,[],[f182,f186])).
fof(f3732,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3731,f119])).
fof(f3731,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f3696,f192])).
fof(f3696,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f186,f182])).
fof(f3730,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3729,f119])).
fof(f3729,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f3728,f192])).
fof(f3728,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f3697,f1323])).
fof(f3697,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6('const.semanticPrimitives.Rraise_1'(s(sK1,sK5)),sK1))))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f167,f186,f1007])).
fof(f3727,plain,(
  $false | ($spl4 | $spl6 | $spl17)),
  inference(subsumption_resolution,[],[f3698,f3694])).
fof(f3694,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | ($spl4 | $spl6 | $spl17)),
  inference(backward_demodulation,[],[f186,f2549])).
fof(f2549,plain,(
  s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | ($spl4 | $spl17)),
  inference(unit_resulting_resolution,[],[f2525,f162])).
fof(f2525,plain,(
  s(sK1,sK5) != s(sK1,sK6(sK3,sK1)) | ($spl4 | $spl17)),
  inference(unit_resulting_resolution,[],[f2524,f1241])).
fof(f1241,plain,(
  ( ! [X0] : (s(sK1,sK5) != s(sK1,X0) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1229,f119])).
fof(f1229,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X0)))) | s(sK1,sK5) != s(sK1,X0)) ) | $spl4),
  inference(superposition,[],[f164,f192])).
fof(f3698,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK6(sK3,sK1)))) | $spl6),
  inference(unit_resulting_resolution,[],[f186,f925])).
fof(f3673,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl18),
  inference(cnf_transformation,[],[f3673_D])).
fof(f3673_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3674,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4)) ) | $spl20),
  inference(cnf_transformation,[],[f3674_D])).
fof(f3674_D,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3675,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f3644,f126])).
fof(f3644,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f3207,f202])).
fof(f3672,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f3641,f126])).
fof(f3641,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f3207,f202])).
fof(f980,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) | $spl12),
  inference(cnf_transformation,[],[f980_D])).
fof(f980_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2524,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1))))) | $spl17),
  inference(cnf_transformation,[],[f2524_D])).
fof(f2524_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f2519,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK6(sK3,sK1))))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) | $spl2),
  inference(equality_resolution,[],[f1007])).
fof(f982,plain,(
  s(sK1,sK5) = s(sK1,sK6(sK3,sK1)) | $spl14),
  inference(cnf_transformation,[],[f982_D])).
fof(f982_D,plain,(
  s(sK1,sK5) = s(sK1,sK6(sK3,sK1)) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f976,plain,(
  s(sK1,sK5) = s(sK1,sK6(sK3,sK1)) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7(sK3,sK1)))) | $spl6),
  inference(equality_resolution,[],[f936])).
fof(f888,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8))))) ) | $spl8),
  inference(cnf_transformation,[],[f888_D])).
fof(f888_D,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f889,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl10),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f890,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK8(X14,X13,X12))))) )),
  inference(subsumption_resolution,[],[f883,f128])).
fof(f883,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK8(X14,X13,X12))))) )),
  inference(superposition,[],[f166,f202])).
fof(f887,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8))))) )),
  inference(subsumption_resolution,[],[f882,f128])).
fof(f882,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f166,f202])).
fof(f180,plain,(
  p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f186,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | $spl6),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f116,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | (s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,sK5))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))))) & (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | ! [X5] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X5))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X6))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X6)))))) & (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) | ! [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) != s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5)))))))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) | ? [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5)))))) & (~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) | ! [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) != s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) <~> ? [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) <=> ? [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X0,fun(X1,bool)),X4),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))),s('type.semanticPrimitives.error_result'(X1),X3)))) <=> ? [X5] : (s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X5))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X5))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))),s('type.semanticPrimitives.error_result'(X7),X9)))) <=> ? [X11] : (s('type.semanticPrimitives.error_result'(X7),X9) = s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X11))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X8))),s(X7,X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.evalProps.exc_rel_3'(s(fun(X6,fun(X7,bool)),X10),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))),s('type.semanticPrimitives.error_result'(X7),X9)))) <=> ? [X11] : (s('type.semanticPrimitives.error_result'(X7),X9) = s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X11))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X8))),s(X7,X11))))))),
  file('thm.evalProps.exc_rel_raise|49|','thm.evalProps.exc_rel_raise|49|')).
fof(f184,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | $spl4),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f117,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,sK5)))) | p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))),
  inference(cnf_transformation,[],[f86])).
fof(f181,plain,(
  ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) | $spl1),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f182,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X5))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X5))))) ) | $spl2),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X5))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X5))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f118,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,sK2))),s(sK1,X5)))) | s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X5))) | ~p(s(bool,'const.evalProps.exc_rel_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK2))),s('type.semanticPrimitives.error_result'(sK1),sK3))))) )),
  inference(cnf_transformation,[],[f86])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 105
% Passive clauses: 521
% Generated clauses: 4724
% Final active clauses: 47
% Final passive clauses: 136
% Input formulas: 32
% Initial clauses: 62
% 
% Duplicate literals: 270
% Trivial inequalities: 73
% Fw subsumption resolutions: 237
% Fw demodulations: 475
% Bw demodulations: 15
% Condensations: 427
% Global subsumptions: 12
% 
% Simple tautologies: 92
% Equational tautologies: 93
% Forward subsumptions: 2403
% Fw demodulations to eq. taut.: 93
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 330
% Factoring: 1
% Forward superposition: 1122
% Backward superposition: 1469
% Self superposition: 83
% Equality factoring: 70
% Equality resolution: 17
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 575
% SAT solver unit clauses: 18
% SAT solver binary clauses: 125
% 
% Sat splits: 10
% Sat splitting refutations: 20
% 
% Memory used [KB]: 3965
% Time elapsed: 0.318 s
% ------------------------------
% ------------------------------
% Success in time 1.68 s
