Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% (128255)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 151
% Passive clauses: 1364
% Generated clauses: 12566
% Final active clauses: 121
% Final passive clauses: 910
% Input formulas: 55
% Initial clauses: 99
% 
% Duplicate literals: 370
% Trivial inequalities: 384
% Fw subsumption resolutions: 1337
% Fw demodulations: 1100
% Bw demodulations: 24
% 
% Simple tautologies: 309
% Equational tautologies: 22
% Forward subsumptions: 7452
% 
% Binary resolution: 134
% Forward superposition: 4600
% Backward superposition: 4178
% Self superposition: 156
% Equality factoring: 50
% Equality resolution: 48
% 
% Split clauses: 159
% Split components: 322
% SAT solver clauses: 399
% SAT solver unit clauses: 15
% SAT solver binary clauses: 99
% 
% Sat splits: 159
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 397
% TWLsolver calls for satisfiability: 77
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_numberProof
% (128292)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 259
% Generated clauses: 12548
% Final active clauses: 88
% Final passive clauses: 138
% Input formulas: 55
% Initial clauses: 99
% 
% Duplicate literals: 388
% Trivial inequalities: 180
% Fw subsumption resolutions: 1192
% Fw demodulations: 1095
% Bw demodulations: 26
% Condensations: 1868
% Global subsumptions: 8
% 
% Simple tautologies: 154
% Equational tautologies: 22
% Forward subsumptions: 6890
% 
% Binary resolution: 19
% Forward superposition: 4127
% Backward superposition: 3016
% Self superposition: 185
% Equality factoring: 51
% Equality resolution: 24
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 566
% SAT solver unit clauses: 134
% SAT solver binary clauses: 38
% SAT solver learnt clauses: 16
% 
% Sat splits: 43
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 534
% TWLsolver calls for satisfiability: 359
% 
% Memory used [KB]: 9978
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_numberProof
% (128331)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 592
% Passive clauses: 5616
% Generated clauses: 161405
% Final active clauses: 100
% Final passive clauses: 495
% Input formulas: 55
% Initial clauses: 99
% 
% Duplicate literals: 3108
% Trivial inequalities: 1077
% Fw subsumption resolutions: 12766
% Fw demodulations: 23455
% Bw demodulations: 4657
% Condensations: 16347
% Global subsumptions: 384
% 
% Simple tautologies: 1074
% Equational tautologies: 971
% Forward subsumptions: 95409
% Fw demodulations to eq. taut.: 22
% Bw demodulations to eq. taut.: 88
% 
% Unit resulting resolution: 9147
% Forward superposition: 51371
% Backward superposition: 32859
% Self superposition: 757
% Equality factoring: 130
% Equality resolution: 214
% 
% Split clauses: 223
% Split components: 463
% SAT solver clauses: 7158
% SAT solver unit clauses: 477
% SAT solver binary clauses: 1147
% 
% Sat splits: 223
% Sat splitting refutations: 710
% 
% Memory used [KB]: 57568
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f12511,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f270,f271,f272,f273,f274,f276,f277,f287,f294,f342,f352,f350,f348,f346,f344,f305,f306,f307,f308,f309,f354,f310,f312,f323,f353,f388,f378,f408,f441,f443,f437,f436,f479,f480,f476,f475,f518,f515,f514,f513,f537,f538,f539,f540,f667,f670,f674,f679,f686,f690,f297,f722,f723,f876,f915,f917,f918,f920,f922,f933,f952,f954,f955,f957,f959,f970,f988,f990,f991,f992,f995,f1005,f1020,f1022,f1023,f1025,f1027,f1038,f1056,f1058,f1059,f1061,f1063,f1074,f1090,f1196,f1202,f1205,f1210,f1215,f1226,f1249,f1251,f1252,f1253,f1256,f1266,f1386,f1392,f1395,f1400,f1405,f1416,f1439,f1546,f1552,f1555,f1560,f1565,f1576,f1049,f1880,f1886,f1889,f1894,f1899,f1910,f1084,f2190,f2197,f2199,f2201,f2202,f2192,f2251,f2252,f2462,f2465,f2485,f235,f231,f2556,f213,f2555,f2557,f2559,f2560,f2562,f2564,f2570,f2577,f2579,f2581,f2582,f2584,f2586,f2592,f2634,f2642,f2644,f2645,f2646,f2649,f2657,f2856,f2903,f2905,f2906,f2907,f2910,f2918,f2920,f2922,f2923,f2924,f2927,f2935,f2964,f2969,f2971,f2972,f2974,f2976,f2983,f3061,f3063,f3064,f3065,f3068,f3076,f3099,f3101,f3102,f3103,f3106,f3114,f3130,f3132,f3134,f3135,f3137,f3139,f3146,f3248,f3250,f3251,f3252,f3255,f3263,f3309,f3264,f3266,f3267,f3269,f3271,f3278,f3312,f3320,f3322,f3323,f3324,f3327,f3335,f3377,f3336,f3338,f3339,f3341,f3343,f3350,f3380,f4330,f4332,f4370,f4371,f4372,f4377,f4380,f4383,f4387,f4391,f4403,f4409,f4414,f4417,f4423,f4427,f4441,f4446,f4449,f4452,f4456,f4460,f4472,f4478,f4483,f4486,f4492,f4496,f4510,f4512,f4513,f4514,f4517,f4525,f4579,f4526,f4528,f4529,f4531,f4533,f4540,f4582,f4586,f4588,f4673,f4674,f3095,f4682,f4684,f4683,f4690,f4692,f4693,f4695,f4697,f4704,f4743,f4796,f4798,f4799,f4800,f4803,f4811,f4983,f4814,f4816,f4817,f4819,f4821,f4828,f4990,f4835,f4837,f4838,f4840,f4842,f4849,f4991,f4856,f4858,f4859,f4861,f4863,f4870,f4992,f5004,f5006,f5198,f5200,f5201,f5203,f5205,f5212,f5252,f5262,f5264,f5265,f5266,f5269,f5277,f5310,f5312,f5440,f5443,f5446,f5448,f5451,f5455,f3029,f5463,f5464,f5469,f5471,f5472,f5473,f5476,f5484,f5516,f5518,f5520,f5519,f5526,f5528,f5529,f5530,f5533,f5541,f5573,f5575,f5576,f5311,f5602,f5604,f5605,f5607,f5609,f5616,f5797,f5800,f5803,f5805,f5808,f5812,f3186,f5820,f5821,f5851,f5853,f5854,f5855,f5858,f5866,f5900,f5902,f5904,f5899,f5908,f5907,f5953,f5955,f5956,f5958,f5960,f5967,f945,f5998,f6000,f6001,f6003,f6005,f6012,f981,f6081,f6083,f6143,f6145,f6186,f6187,f7727,f7735,f7724,f7748,f7723,f7720,f7719,f7718,f7717,f7716,f7715,f7714,f7713,f7710,f7757,f7709,f7706,f7763,f7705,f7770,f7702,f7775,f7777,f7699,f7696,f7751,f7784,f7785,f7695,f7790,f7694,f7693,f7795,f7844,f7845,f7852,f7853,f7689,f7864,f7866,f7867,f7869,f7872,f7874,f7875,f7877,f7879,f7881,f7887,f7889,f7891,f7932,f7934,f5308,f7935,f7857,f7938,f7940,f7678,f7676,f7673,f7953,f7955,f7954,f7956,f7957,f7958,f7959,f7960,f7962,f7961,f7963,f7964,f7965,f7966,f7967,f7968,f7969,f7970,f7971,f7972,f7667,f7973,f7974,f6082,f7666,f7976,f7977,f6144,f7978,f7979,f7980,f7981,f7982,f7983,f7985,f7988,f7993,f7995,f7996,f8000,f8002,f8004,f8009,f8010,f8011,f8013,f8014,f7665,f8016,f7664,f7663,f8017,f8019,f8022,f8023,f8025,f8026,f8027,f8028,f8030,f8032,f8034,f8035,f6811,f8037,f8041,f8043,f8044,f8045,f6822,f8046,f8048,f8052,f8053,f8054,f8055,f8056,f7650,f8074,f7649,f8075,f8076,f8077,f8078,f8079,f8080,f8081,f8082,f8083,f8089,f8093,f8096,f8098,f8103,f8105,f8108,f8111,f8115,f8117,f8118,f8120,f7648,f8122,f7647,f7646,f8126,f8129,f8130,f7645,f7644,f8139,f8140,f8143,f8144,f8145,f8147,f7641,f8154,f8155,f7640,f8042,f8158,f8024,f8159,f5903,f8171,f8172,f8173,f8177,f8178,f8179,f8180,f8181,f7031,f7032,f7033,f8185,f8186,f8187,f8188,f8189,f8190,f8191,f7041,f8193,f7043,f8195,f7045,f7752,f8197,f7050,f7051,f7052,f7053,f7054,f7055,f7056,f8198,f8201,f8203,f8204,f8206,f8208,f8209,f8210,f8212,f8213,f8216,f8217,f8221,f8222,f7635,f8223,f8224,f8225,f8226,f8227,f7633,f8229,f8231,f8234,f8237,f8239,f8240,f8243,f8244,f8245,f8248,f8250,f8253,f8254,f8199,f8255,f8258,f8263,f8049,f8271,f8272,f7623,f7621,f8274,f8275,f8276,f8277,f8278,f8287,f8288,f8290,f8291,f8295,f8298,f8301,f8304,f8306,f8313,f8318,f8320,f8330,f8331,f8337,f8342,f8347,f8350,f8355,f8362,f5660,f8371,f8372,f8384,f7734,f8393,f7577,f8406,f8408,f8409,f8410,f8413,f8421,f1015,f8444,f8447,f8449,f8450,f8451,f8498,f8499,f8566,f9364,f9363,f9362,f9359,f9358,f9371,f9356,f9353,f9352,f9374,f9375,f9349,f9346,f9345,f9379,f9344,f9382,f9384,f9387,f9388,f9389,f9391,f9393,f9394,f9397,f9398,f9399,f9400,f9401,f9343,f9402,f9342,f9341,f9403,f9404,f9406,f9408,f9410,f9411,f9412,f9413,f9414,f9416,f9417,f9418,f9419,f9421,f9422,f9423,f9424,f9338,f9427,f9430,f9431,f9432,f9433,f9434,f9435,f9436,f9437,f9438,f9439,f9440,f9451,f9452,f9455,f9456,f9457,f9460,f9461,f9458,f9462,f9465,f9468,f9470,f9471,f9474,f9475,f9476,f9479,f9481,f9482,f9483,f9484,f9336,f9485,f9335,f9334,f9486,f9487,f9488,f9333,f9332,f9490,f9492,f9493,f9494,f9495,f8146,f9497,f9329,f296,f9328,f9500,f9501,f9510,f9511,f9512,f9515,f9516,f2200,f9517,f2198,f9518,f9519,f9002,f8182,f9003,f8183,f9004,f8184,f9005,f5005,f9520,f9521,f9522,f7778,f9523,f9524,f9525,f7776,f9012,f8192,f9013,f4331,f9014,f8446,f9015,f8194,f9016,f4587,f9017,f8196,f9018,f4333,f9022,f9023,f9024,f9025,f9026,f9027,f9028,f9526,f9528,f9531,f9533,f9534,f9536,f9538,f9539,f9540,f9542,f9543,f9546,f9547,f9550,f9552,f9553,f9554,f9322,f9555,f9556,f9557,f9558,f9320,f9559,f9560,f9562,f9563,f9564,f9567,f9570,f9572,f9573,f9576,f9577,f9578,f9581,f9583,f9586,f9587,f9529,f9588,f9591,f9592,f9595,f9527,f9309,f9365,f9599,f9367,f9600,f9601,f9308,f9306,f9603,f9604,f9605,f9606,f7975,f9607,f8230,f9616,f9617,f9619,f9620,f9296,f9625,f9627,f9628,f9630,f9632,f9639,f9702,f7736,f10903,f10910,f10917,f10918,f10919,f10920,f10921,f10922,f10923,f10926,f10927,f10928,f10933,f10934,f10935,f10971,f7933,f10978,f7939,f10985,f10986,f7865,f10987,f7868,f10988,f7871,f10989,f10990,f7873,f10991,f442,f7876,f10994,f7878,f10995,f7880,f10996,f7882,f11001,f7888,f11002,f7890,f11003,f7892,f11004,f11009,f11012,f11017,f11021,f11028,f11036,f11037,f11049,f11055,f11060,f11065,f11068,f11073,f11080,f11090,f11104,f11111,f11116,f11119,f11124,f11128,f11135,f11143,f11144,f8015,f444,f11156,f11162,f11167,f11172,f11175,f11180,f11187,f7989,f11197,f8029,f11211,f7762,f11219,f11221,f11222,f11223,f11226,f11234,f1277,f11263,f11266,f11268,f11269,f11270,f8160,f11265,f11346,f8033,f2191,f11331,f8319,f11362,f11364,f11365,f11367,f11369,f11376,f11413,f11421,f11422,f11423,f11425,f11426,f11424,f11428,f11430,f11431,f11433,f11435,f11442,f11479,f11487,f11488,f11489,f11491,f11492,f11490,f11839,f11832,f11825,f11824,f11927,f11966,f347,f11816,f7733,f8394,f12036,f12038,f12039,f295,f12041,f12043,f298,f12050,f12089,f300,f12097,f12098,f12099,f7737,f12101,f12102,f299,f12100,f12443,f12444,f12510])).
fof(f12510,plain,(
  $false | $spl21),
  inference(subsumption_resolution,[],[f350,f266])).
fof(f266,plain,(
  s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3)))),
  inference(resolution,[],[f172,f254])).
fof(f254,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))))) & (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))))) & ((s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))))) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) <=> (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) <=> (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X16] : ! [X8,X9] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X16),X8),s('type.closSem.state'(X16),X9)))) <=> (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X8))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X9))) & s('type.ffi.ffi_state'(X16),'const.closSem.state_ffi_1'(s('type.closSem.state'(X16),X8))) = s('type.ffi.ffi_state'(X16),'const.closSem.state_ffi_1'(s('type.closSem.state'(X16),X9))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X16),X8))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X16),X9)))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X16),X8))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X16),X9)))))) & p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X16),X8))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X16),X9))))))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.clos_numberProof.state_rel_def')).
fof(f172,plain,(
  p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))))),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) & ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f129])).
fof(f129,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3)))) & ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X3)))))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3)))) => p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X3)))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3)))) => p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X6),X8),s('type.closSem.state'(X6),X9)))) => p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.closSem.state'(X6),X8))),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.closSem.state'(X6),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X6),X8),s('type.closSem.state'(X6),X9)))) => p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.closSem.state'(X6),X8))),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.closSem.state'(X6),X9)))))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb')).
fof(f12444,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f12119,f250])).
fof(f250,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb',aHL_FALSITY)).
fof(f12119,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(backward_demodulation,[],[f7733,f268])).
fof(f268,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(resolution,[],[f172,f256])).
fof(f256,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f12443,plain,(
  $false | $spl132),
  inference(subsumption_resolution,[],[f12118,f250])).
fof(f12118,plain,(
  p(s(bool,'F_0')) | $spl132),
  inference(backward_demodulation,[],[f7733,f269])).
fof(f269,plain,(
  p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(resolution,[],[f172,f257])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f12100,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X68))) = s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X69))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X66),X68))),s('type.closSem.state'(X66),X69)))) ) | $spl322),
  inference(cnf_transformation,[],[f12100_D])).
fof(f12100_D,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X68))) = s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X69))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X66),X68))),s('type.closSem.state'(X66),X69)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f299,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl10),
  inference(cnf_transformation,[],[f299_D])).
fof(f299_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f12102,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,X103)) | s('type.ffi.ffi_state'(X99),'const.closSem.state_ffi_1'(s('type.closSem.state'(X99),X101))) = s('type.ffi.ffi_state'(X99),'const.closSem.state_ffi_1'(s('type.closSem.state'(X99),X102))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X99),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X100),s('type.closSem.state'(X99),X101))),s('type.closSem.state'(X99),X102)))) )),
  inference(subsumption_resolution,[],[f12083,f242])).
fof(f242,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f12083,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,X103)) | s('type.ffi.ffi_state'(X99),'const.closSem.state_ffi_1'(s('type.closSem.state'(X99),X101))) = s('type.ffi.ffi_state'(X99),'const.closSem.state_ffi_1'(s('type.closSem.state'(X99),X102))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X99),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X100),s('type.closSem.state'(X99),X101))),s('type.closSem.state'(X99),X102))) | s(bool,'F_0') = s(bool,X103)) )),
  inference(superposition,[],[f3155,f7622])).
fof(f7622,plain,(
  ( ! [X526,X527] : (s(bool,X526) = s(bool,X527) | s(bool,'F_0') = s(bool,X527) | s(bool,'F_0') = s(bool,X526)) )),
  inference(trivial_inequality_removal,[],[f7268])).
fof(f7268,plain,(
  ( ! [X526,X527] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X526) = s(bool,X527) | s(bool,'F_0') = s(bool,X527) | s(bool,'F_0') = s(bool,X526)) )),
  inference(superposition,[],[f560,f3476])).
fof(f3476,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f3074,f261])).
fof(f261,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb',aHL_BOOL_CASES)).
fof(f3074,plain,(
  ( ! [X50,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X51)) )),
  inference(resolution,[],[f2660,f247])).
fof(f247,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f2660,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f243,f261])).
fof(f243,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f165])).
fof(f560,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f551,f283])).
fof(f283,plain,(
  s(bool,'F_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))),
  inference(resolution,[],[f172,f242])).
fof(f551,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))),
  inference(superposition,[],[f283,f261])).
fof(f3155,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X8),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.closSem.state'(X8),X10))),s('type.closSem.state'(X8),X11)))) | s('type.ffi.ffi_state'(X8),'const.closSem.state_ffi_1'(s('type.closSem.state'(X8),X10))) = s('type.ffi.ffi_state'(X8),'const.closSem.state_ffi_1'(s('type.closSem.state'(X8),X11)))) )),
  inference(superposition,[],[f254,f178])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X16] : ! [X8,X17] : s('type.ffi.ffi_state'(X16),'const.closSem.state_ffi_1'(s('type.closSem.state'(X16),X8))) = s('type.ffi.ffi_state'(X16),'const.closSem.state_ffi_1'(s('type.closSem.state'(X16),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_73_prove_ltb')).
fof(f12101,plain,(
  ( ! [X94,X97,X95,X98,X96] : (~p(s(bool,X98)) | s('type.ffi.ffi_state'(X94),'const.closSem.state_ffi_1'(s('type.closSem.state'(X94),X96))) = s('type.ffi.ffi_state'(X94),'const.closSem.state_ffi_1'(s('type.closSem.state'(X94),X97))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X94),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X95),s('type.closSem.state'(X94),X96))),s('type.closSem.state'(X94),X97)))) )),
  inference(subsumption_resolution,[],[f12082,f242])).
fof(f12082,plain,(
  ( ! [X94,X97,X95,X98,X96] : (~p(s(bool,X98)) | s('type.ffi.ffi_state'(X94),'const.closSem.state_ffi_1'(s('type.closSem.state'(X94),X96))) = s('type.ffi.ffi_state'(X94),'const.closSem.state_ffi_1'(s('type.closSem.state'(X94),X97))) | s(bool,'F_0') = s(bool,X98) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X94),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X95),s('type.closSem.state'(X94),X96))),s('type.closSem.state'(X94),X97)))) )),
  inference(superposition,[],[f3155,f7622])).
fof(f7737,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl138),
  inference(cnf_transformation,[],[f7737_D])).
fof(f7737_D,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f12099,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X68))) = s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X69))) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X66),X68))),s('type.closSem.state'(X66),X69)))) )),
  inference(subsumption_resolution,[],[f12077,f2660])).
fof(f12077,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,X70)) | s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X68))) = s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X69))) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X67),s('type.closSem.state'(X66),X68))),s('type.closSem.state'(X66),X69)))) )),
  inference(superposition,[],[f3155,f3476])).
fof(f12098,plain,(
  ( ! [X61,X62,X60,X65,X63] : (s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X62))) = s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X61),s('type.closSem.state'(X60),X62))),s('type.closSem.state'(X60),X63)))) ) | $spl320),
  inference(cnf_transformation,[],[f12098_D])).
fof(f12098_D,plain,(
  ( ! [X61,X62,X60,X65,X63] : (s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X62))) = s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X61),s('type.closSem.state'(X60),X62))),s('type.closSem.state'(X60),X63)))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f12097,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~p(s(bool,X64)) | s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X62))) = s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X63))) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X61),s('type.closSem.state'(X60),X62))),s('type.closSem.state'(X60),X63))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f12076,f242])).
fof(f12076,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~p(s(bool,X64)) | s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X62))) = s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X63))) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X61),s('type.closSem.state'(X60),X62))),s('type.closSem.state'(X60),X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f3155,f3476])).
fof(f300,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl12),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12089,plain,(
  ( ! [X6,X4,X10,X8,X5] : (s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X5))) = s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X6))) | s(bool,X8) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X4),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.closSem.state'(X4),X5))),s('type.closSem.state'(X4),X6))) | ~p(s(bool,X8))) ) | $spl318),
  inference(cnf_transformation,[],[f12089_D])).
fof(f12089_D,plain,(
  ( ! [X6,X4,X10,X8,X5] : (s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X5))) = s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X6))) | s(bool,X8) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X4),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.closSem.state'(X4),X5))),s('type.closSem.state'(X4),X6))) | ~p(s(bool,X8))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f12050,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s('type.ffi.ffi_state'(X80),'const.closSem.state_ffi_1'(s('type.closSem.state'(X80),X81))) = s('type.ffi.ffi_state'(X80),'const.closSem.state_ffi_1'(s('type.closSem.state'(X80),X82))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X85),s('type.closSem.state'(X80),X81))),s('type.closSem.state'(X80),X82)))) )),
  inference(resolution,[],[f3155,f262])).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f200])).
fof(f200,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X14,X10,X15] : ((s(bool,X7) = s(bool,X14) & (p(s(bool,X14)) => s(bool,X10) = s(bool,X15))) => ((p(s(bool,X7)) => p(s(bool,X10))) <=> (p(s(bool,X14)) => p(s(bool,X15)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CONG')).
fof(f298,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl8),
  inference(cnf_transformation,[],[f298_D])).
fof(f298_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f12043,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (s('type.ffi.ffi_state'(X52),'const.closSem.state_ffi_1'(s('type.closSem.state'(X52),X53))) = s('type.ffi.ffi_state'(X52),'const.closSem.state_ffi_1'(s('type.closSem.state'(X52),X54))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X52),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X57),s('type.closSem.state'(X52),X53))),s('type.closSem.state'(X52),X54))) | s(bool,X55) != s(bool,X58)) )),
  inference(resolution,[],[f3155,f203])).
fof(f203,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f12041,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s('type.ffi.ffi_state'(X39),'const.closSem.state_ffi_1'(s('type.closSem.state'(X39),X40))) = s('type.ffi.ffi_state'(X39),'const.closSem.state_ffi_1'(s('type.closSem.state'(X39),X41))) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X39),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X45),s('type.closSem.state'(X39),X40))),s('type.closSem.state'(X39),X41)))) )),
  inference(resolution,[],[f3155,f202])).
fof(f202,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f295,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f295_D])).
fof(f295_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f12039,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s('type.ffi.ffi_state'(X25),'const.closSem.state_ffi_1'(s('type.closSem.state'(X25),X26))) = s('type.ffi.ffi_state'(X25),'const.closSem.state_ffi_1'(s('type.closSem.state'(X25),X27))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X25),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X31),s('type.closSem.state'(X25),X26))),s('type.closSem.state'(X25),X27)))) )),
  inference(resolution,[],[f3155,f201])).
fof(f201,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f12038,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s('type.ffi.ffi_state'(X18),'const.closSem.state_ffi_1'(s('type.closSem.state'(X18),X19))) = s('type.ffi.ffi_state'(X18),'const.closSem.state_ffi_1'(s('type.closSem.state'(X18),X20))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X23),s('type.closSem.state'(X18),X19))),s('type.closSem.state'(X18),X20))) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f3155,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f12036,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X5))) = s('type.ffi.ffi_state'(X4),'const.closSem.state_ffi_1'(s('type.closSem.state'(X4),X6))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X4),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.closSem.state'(X4),X5))),s('type.closSem.state'(X4),X6)))) )),
  inference(resolution,[],[f3155,f197])).
fof(f197,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f8394,plain,(
  ( ! [X504,X503] : (s(bool,X503) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X504) | s(bool,X503) = s(bool,X504)) ) | $spl272),
  inference(cnf_transformation,[],[f8394_D])).
fof(f8394_D,plain,(
  ( ! [X504,X503] : (s(bool,X503) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X504) | s(bool,X503) = s(bool,X504)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f7733,plain,(
  ( ! [X46] : (s(bool,'F_0') = s(bool,X46)) ) | $spl132),
  inference(cnf_transformation,[],[f7733_D])).
fof(f7733_D,plain,(
  ( ! [X46] : (s(bool,'F_0') = s(bool,X46)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11816,plain,(
  ( ! [X263,X265,X264] : (s(bool,'F_0') = s(bool,X263) | s(bool,X264) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X264) = s(bool,X265) | s(bool,'T_0') = s(bool,X265)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f11791])).
fof(f11791,plain,(
  ( ! [X263,X265,X264] : (s(bool,'F_0') = s(bool,X263) | s(bool,X264) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X264) = s(bool,X265) | s(bool,'T_0') = s(bool,X265) | s(bool,'F_0') = s(bool,X263)) ) | $spl144),
  inference(superposition,[],[f3476,f7762])).
fof(f347,plain,(
  p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | $spl18),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f11966,plain,(
  ( ! [X8] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl144),
  inference(subsumption_resolution,[],[f11690,f243])).
fof(f11690,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl144),
  inference(superposition,[],[f255,f7762])).
fof(f255,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f11927,plain,(
  ( ! [X275,X274] : (s(bool,X274) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X275) | s(bool,X274) != s(bool,X275)) ) | $spl316),
  inference(cnf_transformation,[],[f11927_D])).
fof(f11927_D,plain,(
  ( ! [X275,X274] : (s(bool,X274) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X275) | s(bool,X274) != s(bool,X275)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f11824,plain,(
  ( ! [X273,X275,X274] : (s(bool,X274) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X273) | s(bool,X274) != s(bool,X275) | s(bool,'T_0') = s(bool,X275)) ) | ($spl60 | $spl144)),
  inference(duplicate_literal_removal,[],[f11640])).
fof(f11640,plain,(
  ( ! [X273,X275,X274] : (s(bool,X274) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X273) | s(bool,X274) != s(bool,X275) | s(bool,'T_0') = s(bool,X275) | s(bool,'F_0') = s(bool,X273)) ) | ($spl60 | $spl144)),
  inference(superposition,[],[f4623,f7762])).
fof(f4623,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl60),
  inference(subsumption_resolution,[],[f4603,f560])).
fof(f4603,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X21) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X19)) ) | $spl60),
  inference(superposition,[],[f3111,f3074])).
fof(f3111,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'F_0') = s(bool,X48) | s(bool,X48) != s(bool,X49)) ) | $spl60),
  inference(resolution,[],[f2856,f242])).
fof(f11825,plain,(
  ( ! [X251,X249,X250] : (s(bool,X250) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,'F_0') = s(bool,X249)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f11632])).
fof(f11632,plain,(
  ( ! [X251,X249,X250] : (s(bool,X250) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,'F_0') = s(bool,X249) | s(bool,'F_0') = s(bool,X249)) ) | $spl144),
  inference(superposition,[],[f3476,f7762])).
fof(f11832,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f11529])).
fof(f11529,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl144),
  inference(superposition,[],[f7762,f3476])).
fof(f11839,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f11505])).
fof(f11505,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X22)) ) | $spl144),
  inference(superposition,[],[f7762,f3476])).
fof(f11490,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X81))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),X81),s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X82),'type.ffi.ffi_state'(X80)),X83),s('type.closSem.state'(X82),X84)))))) ) | $spl314),
  inference(cnf_transformation,[],[f11490_D])).
fof(f11490_D,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X81))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),X81),s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X82),'type.ffi.ffi_state'(X80)),X83),s('type.closSem.state'(X82),X84)))))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f11492,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,X124)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X119),X120))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X121),X123))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X119),X120),s('type.closSem.state'(X119),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X121),'type.ffi.ffi_state'(X119)),X122),s('type.closSem.state'(X121),X123)))))) )),
  inference(subsumption_resolution,[],[f11475,f242])).
fof(f11475,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,X124)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X119),X120))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X121),X123))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X119),X120),s('type.closSem.state'(X119),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X121),'type.ffi.ffi_state'(X119)),X122),s('type.closSem.state'(X121),X123))))) | s(bool,'F_0') = s(bool,X124)) )),
  inference(superposition,[],[f3001,f7622])).
fof(f3001,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X16),X20),s('type.closSem.state'(X16),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X17),'type.ffi.ffi_state'(X16)),X18),s('type.closSem.state'(X17),X19)))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X20))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X17),X19)))) )),
  inference(superposition,[],[f253,f182])).
fof(f182,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X1),X2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X1),X2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X1),X2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X18,X16] : ! [X8,X17] : s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X8))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X18),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X16),'type.ffi.ffi_state'(X18)),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_77_prove_ltb')).
fof(f253,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f11491,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,X118)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X113),X114))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X115),X117))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X113),X114),s('type.closSem.state'(X113),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X115),'type.ffi.ffi_state'(X113)),X116),s('type.closSem.state'(X115),X117)))))) )),
  inference(subsumption_resolution,[],[f11474,f242])).
fof(f11474,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,X118)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X113),X114))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X115),X117))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X113),X114),s('type.closSem.state'(X113),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X115),'type.ffi.ffi_state'(X113)),X116),s('type.closSem.state'(X115),X117)))))) )),
  inference(superposition,[],[f3001,f7622])).
fof(f11489,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X81))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84))) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),X81),s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X82),'type.ffi.ffi_state'(X80)),X83),s('type.closSem.state'(X82),X84)))))) )),
  inference(subsumption_resolution,[],[f11469,f2660])).
fof(f11469,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X85)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X81))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84))) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),X81),s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X82),'type.ffi.ffi_state'(X80)),X83),s('type.closSem.state'(X82),X84)))))) )),
  inference(superposition,[],[f3001,f3476])).
fof(f11488,plain,(
  ( ! [X76,X74,X79,X77,X75,X73] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X74))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X75),X77))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),X74),s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X75),'type.ffi.ffi_state'(X73)),X76),s('type.closSem.state'(X75),X77)))))) ) | $spl312),
  inference(cnf_transformation,[],[f11488_D])).
fof(f11488_D,plain,(
  ( ! [X76,X74,X79,X77,X75,X73] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X74))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X75),X77))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),X74),s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X75),'type.ffi.ffi_state'(X73)),X76),s('type.closSem.state'(X75),X77)))))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f11487,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,X78)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X74))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X75),X77))) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),X74),s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X75),'type.ffi.ffi_state'(X73)),X76),s('type.closSem.state'(X75),X77))))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f11468,f242])).
fof(f11468,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,X78)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X74))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X75),X77))) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),X74),s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X75),'type.ffi.ffi_state'(X73)),X76),s('type.closSem.state'(X75),X77))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78)) )),
  inference(superposition,[],[f3001,f3476])).
fof(f11479,plain,(
  ( ! [X6,X12,X10,X8,X7,X5] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),X6),s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8))))) | ~p(s(bool,X10))) ) | $spl310),
  inference(cnf_transformation,[],[f11479_D])).
fof(f11479_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X5] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),X6),s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8))))) | ~p(s(bool,X10))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f11442,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X95),X96))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X95),X96),s('type.closSem.state'(X95),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X97),'type.ffi.ffi_state'(X95)),X101),s('type.closSem.state'(X97),X98)))))) )),
  inference(resolution,[],[f3001,f262])).
fof(f11435,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X60),X61))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X62),X63))) | ~p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),X61),s('type.closSem.state'(X60),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X62),'type.ffi.ffi_state'(X60)),X66),s('type.closSem.state'(X62),X63))))) | s(bool,X64) != s(bool,X67)) )),
  inference(resolution,[],[f3001,f203])).
fof(f11433,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X46))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X45),X46),s('type.closSem.state'(X45),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X47),'type.ffi.ffi_state'(X45)),X52),s('type.closSem.state'(X47),X48)))))) )),
  inference(resolution,[],[f3001,f202])).
fof(f11431,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X31),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X31),'type.ffi.ffi_state'(X29)),X36),s('type.closSem.state'(X31),X32)))))) )),
  inference(resolution,[],[f3001,f201])).
fof(f11430,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X22))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X24))) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X21),X22),s('type.closSem.state'(X21),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X23),'type.ffi.ffi_state'(X21)),X27),s('type.closSem.state'(X23),X24))))) | s(bool,X25) != s(bool,X28)) )),
  inference(resolution,[],[f3001,f199])).
fof(f11428,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),X6),s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8)))))) )),
  inference(resolution,[],[f3001,f197])).
fof(f11424,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X84))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X81),X83))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X81),'type.ffi.ffi_state'(X80)),X82),s('type.closSem.state'(X81),X83))),s('type.closSem.state'(X80),X84)))) ) | $spl308),
  inference(cnf_transformation,[],[f11424_D])).
fof(f11424_D,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X84))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X81),X83))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X81),'type.ffi.ffi_state'(X80)),X82),s('type.closSem.state'(X81),X83))),s('type.closSem.state'(X80),X84)))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f11426,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,X124)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X119),X123))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X120),X122))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X119),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X120),'type.ffi.ffi_state'(X119)),X121),s('type.closSem.state'(X120),X122))),s('type.closSem.state'(X119),X123)))) )),
  inference(subsumption_resolution,[],[f11409,f242])).
fof(f11409,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (~p(s(bool,X124)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X119),X123))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X120),X122))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X119),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X120),'type.ffi.ffi_state'(X119)),X121),s('type.closSem.state'(X120),X122))),s('type.closSem.state'(X119),X123))) | s(bool,'F_0') = s(bool,X124)) )),
  inference(superposition,[],[f2992,f7622])).
fof(f2992,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X16),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X17),'type.ffi.ffi_state'(X16)),X18),s('type.closSem.state'(X17),X19))),s('type.closSem.state'(X16),X20)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X20))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X17),X19)))) )),
  inference(superposition,[],[f253,f182])).
fof(f11425,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,X118)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X113),X117))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X114),X116))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X113),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X114),'type.ffi.ffi_state'(X113)),X115),s('type.closSem.state'(X114),X116))),s('type.closSem.state'(X113),X117)))) )),
  inference(subsumption_resolution,[],[f11408,f242])).
fof(f11408,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,X118)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X113),X117))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X114),X116))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X113),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X114),'type.ffi.ffi_state'(X113)),X115),s('type.closSem.state'(X114),X116))),s('type.closSem.state'(X113),X117)))) )),
  inference(superposition,[],[f2992,f7622])).
fof(f11423,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X84))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X81),X83))) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X81),'type.ffi.ffi_state'(X80)),X82),s('type.closSem.state'(X81),X83))),s('type.closSem.state'(X80),X84)))) )),
  inference(subsumption_resolution,[],[f11403,f2660])).
fof(f11403,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X85)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X80),X84))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X81),X83))) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X80),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X81),'type.ffi.ffi_state'(X80)),X82),s('type.closSem.state'(X81),X83))),s('type.closSem.state'(X80),X84)))) )),
  inference(superposition,[],[f2992,f3476])).
fof(f11422,plain,(
  ( ! [X76,X74,X79,X77,X75,X73] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X77))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X74),X76))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X74),'type.ffi.ffi_state'(X73)),X75),s('type.closSem.state'(X74),X76))),s('type.closSem.state'(X73),X77)))) ) | $spl306),
  inference(cnf_transformation,[],[f11422_D])).
fof(f11422_D,plain,(
  ( ! [X76,X74,X79,X77,X75,X73] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X77))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X74),X76))) | s(bool,'T_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X74),'type.ffi.ffi_state'(X73)),X75),s('type.closSem.state'(X74),X76))),s('type.closSem.state'(X73),X77)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f11421,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,X78)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X77))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X74),X76))) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X74),'type.ffi.ffi_state'(X73)),X75),s('type.closSem.state'(X74),X76))),s('type.closSem.state'(X73),X77))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f11402,f242])).
fof(f11402,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,X78)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X73),X77))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X74),X76))) | s(bool,X79) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X73),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X74),'type.ffi.ffi_state'(X73)),X75),s('type.closSem.state'(X74),X76))),s('type.closSem.state'(X73),X77))) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78)) )),
  inference(superposition,[],[f2992,f3476])).
fof(f11413,plain,(
  ( ! [X6,X12,X10,X8,X7,X5] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8))),s('type.closSem.state'(X5),X6))) | ~p(s(bool,X10))) ) | $spl304),
  inference(cnf_transformation,[],[f11413_D])).
fof(f11413_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X5] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8))),s('type.closSem.state'(X5),X6))) | ~p(s(bool,X10))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f11376,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X95),X96))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X95),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X97),'type.ffi.ffi_state'(X95)),X101),s('type.closSem.state'(X97),X98))),s('type.closSem.state'(X95),X96)))) )),
  inference(resolution,[],[f2992,f262])).
fof(f11369,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X60),X61))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X62),X63))) | ~p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X62),'type.ffi.ffi_state'(X60)),X66),s('type.closSem.state'(X62),X63))),s('type.closSem.state'(X60),X61))) | s(bool,X64) != s(bool,X67)) )),
  inference(resolution,[],[f2992,f203])).
fof(f11367,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X46))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X45),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X47),'type.ffi.ffi_state'(X45)),X52),s('type.closSem.state'(X47),X48))),s('type.closSem.state'(X45),X46)))) )),
  inference(resolution,[],[f2992,f202])).
fof(f11365,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X31),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X31),'type.ffi.ffi_state'(X29)),X36),s('type.closSem.state'(X31),X32))),s('type.closSem.state'(X29),X30)))) )),
  inference(resolution,[],[f2992,f201])).
fof(f11364,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X22))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X24))) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X21),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X23),'type.ffi.ffi_state'(X21)),X27),s('type.closSem.state'(X23),X24))),s('type.closSem.state'(X21),X22))) | s(bool,X25) != s(bool,X28)) )),
  inference(resolution,[],[f2992,f199])).
fof(f11362,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X5),X6))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X7),X8))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X5),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X5)),X12),s('type.closSem.state'(X7),X8))),s('type.closSem.state'(X5),X6)))) )),
  inference(resolution,[],[f2992,f197])).
fof(f8319,plain,(
  ( ! [X379,X376,X378] : (s(bool,'T_0') != s(bool,X376) | s(bool,'T_0') = s(bool,X379) | s(bool,X378) = s(bool,X379) | s(bool,X376) = s(bool,X378)) ) | $spl266),
  inference(cnf_transformation,[],[f8319_D])).
fof(f8319_D,plain,(
  ( ! [X379,X376,X378] : (s(bool,'T_0') != s(bool,X376) | s(bool,'T_0') = s(bool,X379) | s(bool,X378) = s(bool,X379) | s(bool,X376) = s(bool,X378)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11331,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(duplicate_literal_removal,[],[f11320])).
fof(f11320,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f8164,f3476])).
fof(f8164,plain,(
  ( ! [X546,X547] : (s(bool,'T_0') != s(bool,X546) | s(bool,X546) = s(bool,X547) | s(bool,'F_0') = s(bool,X547)) )),
  inference(subsumption_resolution,[],[f7013,f560])).
fof(f7013,plain,(
  ( ! [X546,X547] : (s(bool,'T_0') != s(bool,X546) | s(bool,X546) = s(bool,X547) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X547)) )),
  inference(superposition,[],[f560,f3476])).
fof(f2191,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) | $spl44),
  inference(cnf_transformation,[],[f2191_D])).
fof(f2191_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8033,plain,(
  ( ! [X335,X337,X338] : (s(bool,X335) != s(bool,X337) | s(bool,'F_0') = s(bool,X338) | s(bool,X335) = s(bool,X338) | s(bool,'F_0') = s(bool,X337)) ) | $spl214),
  inference(cnf_transformation,[],[f8033_D])).
fof(f8033_D,plain,(
  ( ! [X335,X337,X338] : (s(bool,X335) != s(bool,X337) | s(bool,'F_0') = s(bool,X338) | s(bool,X335) = s(bool,X338) | s(bool,'F_0') = s(bool,X337)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f11346,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,X45) = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X44)) )),
  inference(duplicate_literal_removal,[],[f11287])).
fof(f11287,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,X45) = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f8164,f3476])).
fof(f11265,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | $spl300),
  inference(cnf_transformation,[],[f11265_D])).
fof(f11265_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f8160,plain,(
  ( ! [X451,X454] : (p(s(bool,X451)) | s(bool,X451) = s(bool,X454)) ) | $spl244),
  inference(cnf_transformation,[],[f8160_D])).
fof(f8160_D,plain,(
  ( ! [X451,X454] : (p(s(bool,X451)) | s(bool,X451) = s(bool,X454)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f11270,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f11253,f2915])).
fof(f2915,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,X48) = s(bool,X49) | p(s(bool,X48))) )),
  inference(resolution,[],[f2659,f242])).
fof(f2659,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(superposition,[],[f243,f243])).
fof(f11253,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f268,f3476])).
fof(f11269,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl302),
  inference(cnf_transformation,[],[f11269_D])).
fof(f11269_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f11268,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f11252,f246])).
fof(f246,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f11252,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f268,f3476])).
fof(f11266,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f11246,f246])).
fof(f11246,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(superposition,[],[f268,f3074])).
fof(f11263,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f11245,f246])).
fof(f11245,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f268,f3074])).
fof(f1277,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) ) | $spl42),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f11234,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f262])).
fof(f11226,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f268,f203])).
fof(f11223,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f202])).
fof(f11222,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f201])).
fof(f11221,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f268,f199])).
fof(f11219,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f197])).
fof(f7762,plain,(
  ( ! [X95] : (s(bool,X95) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X95)) ) | $spl144),
  inference(cnf_transformation,[],[f7762_D])).
fof(f7762_D,plain,(
  ( ! [X95] : (s(bool,X95) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X95)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11211,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') = s(bool,X419) | s(bool,X420) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(subsumption_resolution,[],[f10679,f261])).
fof(f10679,plain,(
  ( ! [X420,X419] : (s(bool,'F_0') != s(bool,X419) | s(bool,'T_0') = s(bool,X419) | s(bool,X420) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(superposition,[],[f3407,f7622])).
fof(f3407,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(trivial_inequality_removal,[],[f3404])).
fof(f3404,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X20) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f2882,f285])).
fof(f285,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))),
  inference(resolution,[],[f172,f247])).
fof(f2882,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2590,f261])).
fof(f2590,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | s(bool,'T_0') != s(bool,X35)) )),
  inference(resolution,[],[f242,f246])).
fof(f8029,plain,(
  ( ! [X325,X323,X326] : (p(s(bool,X325)) | s(bool,'F_0') = s(bool,X326) | s(bool,X323) = s(bool,X326) | s(bool,X323) = s(bool,X325)) ) | $spl212),
  inference(cnf_transformation,[],[f8029_D])).
fof(f8029_D,plain,(
  ( ! [X325,X323,X326] : (p(s(bool,X325)) | s(bool,'F_0') = s(bool,X326) | s(bool,X323) = s(bool,X326) | s(bool,X323) = s(bool,X325)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f11197,plain,(
  ( ! [X277,X275,X276,X274] : (p(s(bool,X276)) | p(s(bool,X275)) | s(bool,X274) = s(bool,X276) | s(bool,X274) = s(bool,X277) | s(bool,'F_0') = s(bool,X277)) ) | $spl62),
  inference(subsumption_resolution,[],[f10634,f243])).
fof(f10634,plain,(
  ( ! [X277,X275,X276,X274] : (s(bool,'F_0') != s(bool,X275) | p(s(bool,X276)) | p(s(bool,X275)) | s(bool,X274) = s(bool,X276) | s(bool,X274) = s(bool,X277) | s(bool,'F_0') = s(bool,X277)) ) | $spl62),
  inference(superposition,[],[f2964,f7622])).
fof(f7989,plain,(
  ( ! [X24,X29,X27] : (~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X29) | s(bool,X24) = s(bool,X29) | s(bool,X24) != s(bool,X27)) ) | $spl206),
  inference(cnf_transformation,[],[f7989_D])).
fof(f7989_D,plain,(
  ( ! [X24,X29,X27] : (~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X29) | s(bool,X24) = s(bool,X29) | s(bool,X24) != s(bool,X27)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f11187,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(subsumption_resolution,[],[f10616,f250])).
fof(f10616,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,'F_0')) | p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(superposition,[],[f262,f7622])).
fof(f11180,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f10592,f250])).
fof(f10592,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,'F_0')) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f203,f7622])).
fof(f11175,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f10584,f250])).
fof(f10584,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f202,f7622])).
fof(f11172,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f10579,f250])).
fof(f10579,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'F_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f201,f7622])).
fof(f11167,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f10572,f250])).
fof(f10572,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f199,f7622])).
fof(f11162,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f10565,f250])).
fof(f10565,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f197,f7622])).
fof(f11156,plain,(
  ( ! [X399,X401,X400,X402] : (s(bool,X399) != s(bool,X400) | s(bool,'F_0') = s(bool,X400) | s(bool,'T_0') = s(bool,X401) | s(bool,X399) = s(bool,X402) | s(bool,'F_0') = s(bool,X402)) ) | $spl60),
  inference(subsumption_resolution,[],[f10554,f261])).
fof(f10554,plain,(
  ( ! [X399,X401,X400,X402] : (s(bool,X399) != s(bool,X400) | s(bool,'F_0') = s(bool,X400) | s(bool,'F_0') != s(bool,X401) | s(bool,'T_0') = s(bool,X401) | s(bool,X399) = s(bool,X402) | s(bool,'F_0') = s(bool,X402)) ) | $spl60),
  inference(superposition,[],[f4623,f7622])).
fof(f444,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | $spl30),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8015,plain,(
  ( ! [X242,X240] : (s(bool,'T_0') != s(bool,X240) | s(bool,'F_0') = s(bool,X242) | s(bool,X240) = s(bool,X242)) ) | $spl208),
  inference(cnf_transformation,[],[f8015_D])).
fof(f8015_D,plain,(
  ( ! [X242,X240] : (s(bool,'T_0') != s(bool,X240) | s(bool,'F_0') = s(bool,X242) | s(bool,X240) = s(bool,X242)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f11144,plain,(
  ( ! [X308,X306,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,'F_0') != s(bool,X307) | s(bool,X306) = s(bool,X308) | s(bool,'F_0') = s(bool,X308)) )),
  inference(subsumption_resolution,[],[f10531,f2590])).
fof(f10531,plain,(
  ( ! [X308,X306,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,'F_0') != s(bool,X307) | s(bool,'T_0') = s(bool,X307) | s(bool,X306) = s(bool,X308) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f2882,f7622])).
fof(f11143,plain,(
  ( ! [X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | p(s(bool,X304)) | s(bool,X303) = s(bool,X305) | s(bool,'F_0') = s(bool,X305)) )),
  inference(subsumption_resolution,[],[f10530,f243])).
fof(f10530,plain,(
  ( ! [X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | s(bool,'F_0') != s(bool,X304) | p(s(bool,X304)) | s(bool,X303) = s(bool,X305) | s(bool,'F_0') = s(bool,X305)) )),
  inference(superposition,[],[f2881,f7622])).
fof(f2881,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) )),
  inference(superposition,[],[f2590,f243])).
fof(f11135,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f10514,f250])).
fof(f10514,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | p(s(bool,'F_0')) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(superposition,[],[f262,f7622])).
fof(f11128,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(subsumption_resolution,[],[f10489,f250])).
fof(f10489,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,'F_0')) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f203,f7622])).
fof(f11124,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f10482,f250])).
fof(f10482,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | p(s(bool,'F_0')) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f202,f7622])).
fof(f11119,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f10475,f250])).
fof(f10475,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | p(s(bool,'F_0')) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(superposition,[],[f201,f7622])).
fof(f11116,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f10470,f250])).
fof(f10470,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,'F_0')) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f199,f7622])).
fof(f11111,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f10460,f250])).
fof(f10460,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,'F_0')) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f197,f7622])).
fof(f11104,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') = s(bool,X419) | s(bool,X420) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(subsumption_resolution,[],[f10423,f261])).
fof(f10423,plain,(
  ( ! [X420,X419] : (s(bool,'F_0') != s(bool,X419) | s(bool,'T_0') = s(bool,X419) | s(bool,X420) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(superposition,[],[f3407,f7622])).
fof(f11090,plain,(
  ( ! [X277,X275,X276,X274] : (p(s(bool,X276)) | p(s(bool,X275)) | s(bool,X274) = s(bool,X276) | s(bool,X274) = s(bool,X277) | s(bool,'F_0') = s(bool,X277)) ) | $spl62),
  inference(subsumption_resolution,[],[f10378,f243])).
fof(f10378,plain,(
  ( ! [X277,X275,X276,X274] : (s(bool,'F_0') != s(bool,X275) | p(s(bool,X276)) | p(s(bool,X275)) | s(bool,X274) = s(bool,X276) | s(bool,X274) = s(bool,X277) | s(bool,'F_0') = s(bool,X277)) ) | $spl62),
  inference(superposition,[],[f2964,f7622])).
fof(f11080,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(subsumption_resolution,[],[f10360,f250])).
fof(f10360,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,'F_0')) | p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(superposition,[],[f262,f7622])).
fof(f11073,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f10336,f250])).
fof(f10336,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,'F_0')) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f203,f7622])).
fof(f11068,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f10328,f250])).
fof(f10328,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f202,f7622])).
fof(f11065,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f10323,f250])).
fof(f10323,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'F_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f201,f7622])).
fof(f11060,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f10316,f250])).
fof(f10316,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f199,f7622])).
fof(f11055,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f10309,f250])).
fof(f10309,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f197,f7622])).
fof(f11049,plain,(
  ( ! [X399,X401,X400,X402] : (s(bool,X399) != s(bool,X400) | s(bool,'F_0') = s(bool,X400) | s(bool,'T_0') = s(bool,X401) | s(bool,X399) = s(bool,X402) | s(bool,'F_0') = s(bool,X402)) ) | $spl60),
  inference(subsumption_resolution,[],[f10298,f261])).
fof(f10298,plain,(
  ( ! [X399,X401,X400,X402] : (s(bool,X399) != s(bool,X400) | s(bool,'F_0') = s(bool,X400) | s(bool,'F_0') != s(bool,X401) | s(bool,'T_0') = s(bool,X401) | s(bool,X399) = s(bool,X402) | s(bool,'F_0') = s(bool,X402)) ) | $spl60),
  inference(superposition,[],[f4623,f7622])).
fof(f11037,plain,(
  ( ! [X308,X306,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,'F_0') != s(bool,X307) | s(bool,X306) = s(bool,X308) | s(bool,'F_0') = s(bool,X308)) )),
  inference(subsumption_resolution,[],[f10275,f2590])).
fof(f10275,plain,(
  ( ! [X308,X306,X307] : (s(bool,'T_0') != s(bool,X306) | s(bool,'F_0') != s(bool,X307) | s(bool,'T_0') = s(bool,X307) | s(bool,X306) = s(bool,X308) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f2882,f7622])).
fof(f11036,plain,(
  ( ! [X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | p(s(bool,X304)) | s(bool,X303) = s(bool,X305) | s(bool,'F_0') = s(bool,X305)) )),
  inference(subsumption_resolution,[],[f10274,f243])).
fof(f10274,plain,(
  ( ! [X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | s(bool,'F_0') != s(bool,X304) | p(s(bool,X304)) | s(bool,X303) = s(bool,X305) | s(bool,'F_0') = s(bool,X305)) )),
  inference(superposition,[],[f2881,f7622])).
fof(f11028,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f10258,f250])).
fof(f10258,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | p(s(bool,'F_0')) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(superposition,[],[f262,f7622])).
fof(f11021,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(subsumption_resolution,[],[f10233,f250])).
fof(f10233,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,'F_0')) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f203,f7622])).
fof(f11017,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f10226,f250])).
fof(f10226,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | p(s(bool,'F_0')) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f202,f7622])).
fof(f11012,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f10219,f250])).
fof(f10219,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | p(s(bool,'F_0')) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(superposition,[],[f201,f7622])).
fof(f11009,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f10214,f250])).
fof(f10214,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,'F_0')) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f199,f7622])).
fof(f11004,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f10204,f250])).
fof(f10204,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,'F_0')) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f197,f7622])).
fof(f7892,plain,(
  ( ! [X679,X680,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) | $spl176),
  inference(cnf_transformation,[],[f7892_D])).
fof(f7892_D,plain,(
  ( ! [X679,X680,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f11003,plain,(
  ( ! [X561,X563,X564,X562] : (~p(s(bool,X564)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X561),X562),s('type.closSem.state'(X561),X563)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X563)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X563)))))) | s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X562))) != s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X563))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X562))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X563))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X562))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X563)))))) )),
  inference(subsumption_resolution,[],[f10187,f242])).
fof(f10187,plain,(
  ( ! [X561,X563,X564,X562] : (~p(s(bool,X564)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X561),X562),s('type.closSem.state'(X561),X563)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X563)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X563)))))) | s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X562))) != s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X563))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X562))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X563))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X562))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X563))))) | s(bool,'F_0') = s(bool,X564)) )),
  inference(superposition,[],[f258,f7622])).
fof(f258,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f7890,plain,(
  ( ! [X668,X670,X669] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) | $spl174),
  inference(cnf_transformation,[],[f7890_D])).
fof(f7890_D,plain,(
  ( ! [X668,X670,X669] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f11002,plain,(
  ( ! [X553,X555,X556,X554] : (~p(s(bool,X556)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X553),X554),s('type.closSem.state'(X553),X555)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X555)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X554))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X555)))))) | s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X554))) != s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X555))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X554))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X555))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X555)))))) )),
  inference(subsumption_resolution,[],[f10185,f242])).
fof(f10185,plain,(
  ( ! [X553,X555,X556,X554] : (~p(s(bool,X556)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X553),X554),s('type.closSem.state'(X553),X555)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X555)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X554))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X555)))))) | s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X554))) != s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X555))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X554))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X555))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X555))))) | s(bool,'F_0') = s(bool,X556)) )),
  inference(superposition,[],[f258,f7622])).
fof(f7888,plain,(
  ( ! [X660,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) | $spl172),
  inference(cnf_transformation,[],[f7888_D])).
fof(f7888_D,plain,(
  ( ! [X660,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f11001,plain,(
  ( ! [X546,X545,X547,X544] : (~p(s(bool,X547)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X544),X545),s('type.closSem.state'(X544),X546)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X545))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X546)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X546)))))) | s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X545))) != s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X546))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X545))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X546))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X546)))))) )),
  inference(subsumption_resolution,[],[f10182,f242])).
fof(f10182,plain,(
  ( ! [X546,X545,X547,X544] : (~p(s(bool,X547)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X544),X545),s('type.closSem.state'(X544),X546)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X545))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X546)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X546)))))) | s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X545))) != s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X546))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X545))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X546))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X546))))) | s(bool,'F_0') = s(bool,X547)) )),
  inference(superposition,[],[f258,f7622])).
fof(f7882,plain,(
  ( ! [X617,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) | $spl170),
  inference(cnf_transformation,[],[f7882_D])).
fof(f7882_D,plain,(
  ( ! [X617,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f10996,plain,(
  ( ! [X511,X512,X510,X513] : (~p(s(bool,X513)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X511))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X512)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X510),X511),s('type.closSem.state'(X510),X512)))) )),
  inference(subsumption_resolution,[],[f10174,f242])).
fof(f10174,plain,(
  ( ! [X511,X512,X510,X513] : (~p(s(bool,X513)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X511))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X512)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X510),X511),s('type.closSem.state'(X510),X512))) | s(bool,'F_0') = s(bool,X513)) )),
  inference(superposition,[],[f257,f7622])).
fof(f7880,plain,(
  ( ! [X613,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) | $spl168),
  inference(cnf_transformation,[],[f7880_D])).
fof(f7880_D,plain,(
  ( ! [X613,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f10995,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,X509)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X507))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X508)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X506),X507),s('type.closSem.state'(X506),X508)))) )),
  inference(subsumption_resolution,[],[f10173,f242])).
fof(f10173,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,X509)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X507))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X508)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X506),X507),s('type.closSem.state'(X506),X508))) | s(bool,'F_0') = s(bool,X509)) )),
  inference(superposition,[],[f256,f7622])).
fof(f7878,plain,(
  ( ! [X606,X607,X608] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) | $spl166),
  inference(cnf_transformation,[],[f7878_D])).
fof(f7878_D,plain,(
  ( ! [X606,X607,X608] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10994,plain,(
  ( ! [X505,X502,X504,X503] : (~p(s(bool,X505)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X503))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X504)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X502),X503),s('type.closSem.state'(X502),X504)))) )),
  inference(subsumption_resolution,[],[f10172,f242])).
fof(f10172,plain,(
  ( ! [X505,X502,X504,X503] : (~p(s(bool,X505)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X503))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X504)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X502),X503),s('type.closSem.state'(X502),X504))) | s(bool,'F_0') = s(bool,X505)) )),
  inference(superposition,[],[f255,f7622])).
fof(f7876,plain,(
  ( ! [X593] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) ) | $spl164),
  inference(cnf_transformation,[],[f7876_D])).
fof(f7876_D,plain,(
  ( ! [X593] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f442,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl28),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f10991,plain,(
  ( ! [X492,X493] : (s(bool,'T_0') != s(bool,X493) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X492))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X492)))) )),
  inference(subsumption_resolution,[],[f10169,f2590])).
fof(f10169,plain,(
  ( ! [X492,X493] : (s(bool,'T_0') != s(bool,X493) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X492))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X492))) | s(bool,'F_0') = s(bool,X493)) )),
  inference(superposition,[],[f1072,f7622])).
fof(f1072,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X41))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X41)))) )),
  inference(resolution,[],[f783,f246])).
fof(f783,plain,(
  ( ! [X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1)))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1)))) )),
  inference(superposition,[],[f266,f254])).
fof(f7873,plain,(
  ( ! [X587] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) ) | $spl162),
  inference(cnf_transformation,[],[f7873_D])).
fof(f7873_D,plain,(
  ( ! [X587] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f10990,plain,(
  ( ! [X490,X491] : (s(bool,'T_0') != s(bool,X491) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X490))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X490)))) )),
  inference(subsumption_resolution,[],[f10168,f2590])).
fof(f10168,plain,(
  ( ! [X490,X491] : (s(bool,'T_0') != s(bool,X491) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X490))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X490))) | s(bool,'F_0') = s(bool,X491)) )),
  inference(superposition,[],[f968,f7622])).
fof(f968,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X41))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X41)))) )),
  inference(resolution,[],[f544,f246])).
fof(f544,plain,(
  ( ! [X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1)))) )),
  inference(superposition,[],[f265,f253])).
fof(f265,plain,(
  s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3)))),
  inference(resolution,[],[f172,f253])).
fof(f10989,plain,(
  ( ! [X488,X489] : (~p(s(bool,X489)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X488))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X488)))) )),
  inference(subsumption_resolution,[],[f10167,f242])).
fof(f10167,plain,(
  ( ! [X488,X489] : (~p(s(bool,X489)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X488))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X488))) | s(bool,'F_0') = s(bool,X489)) )),
  inference(superposition,[],[f544,f7622])).
fof(f7871,plain,(
  s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | $spl160),
  inference(cnf_transformation,[],[f7871_D])).
fof(f7871_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f10988,plain,(
  ( ! [X484] : (s(bool,'T_0') != s(bool,X484) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f10164,f2590])).
fof(f10164,plain,(
  ( ! [X484] : (s(bool,'T_0') != s(bool,X484) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X484)) )),
  inference(superposition,[],[f321,f7622])).
fof(f321,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(resolution,[],[f173,f246])).
fof(f173,plain,(
  ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))))),
  inference(cnf_transformation,[],[f139])).
fof(f7868,plain,(
  ( ! [X577] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) ) | $spl158),
  inference(cnf_transformation,[],[f7868_D])).
fof(f7868_D,plain,(
  ( ! [X577] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f10987,plain,(
  ( ! [X482,X483] : (s(bool,'T_0') != s(bool,X483) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X482))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X482),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f10163,f2590])).
fof(f10163,plain,(
  ( ! [X482,X483] : (s(bool,'T_0') != s(bool,X483) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X482))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X482),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X483)) )),
  inference(superposition,[],[f1036,f7622])).
fof(f1036,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X41),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X41)))) )),
  inference(resolution,[],[f782,f246])).
fof(f782,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X0),s('type.closSem.state'(sK0),sK3)))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X0)))) )),
  inference(superposition,[],[f266,f254])).
fof(f7865,plain,(
  ( ! [X571] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) ) | $spl156),
  inference(cnf_transformation,[],[f7865_D])).
fof(f7865_D,plain,(
  ( ! [X571] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f10986,plain,(
  ( ! [X480,X481] : (s(bool,'T_0') != s(bool,X481) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X480))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X480),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f10162,f2590])).
fof(f10162,plain,(
  ( ! [X480,X481] : (s(bool,'T_0') != s(bool,X481) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X480))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X480),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X481)) )),
  inference(superposition,[],[f931,f7622])).
fof(f931,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X41),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X41)))) )),
  inference(resolution,[],[f543,f246])).
fof(f543,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X0),s('type.closSem.state'(sK0),sK3)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X0)))) )),
  inference(superposition,[],[f265,f253])).
fof(f10985,plain,(
  ( ! [X478,X479] : (~p(s(bool,X479)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X478))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X478),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f10161,f242])).
fof(f10161,plain,(
  ( ! [X478,X479] : (~p(s(bool,X479)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X478))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X478),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X479)) )),
  inference(superposition,[],[f543,f7622])).
fof(f7939,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'F_0') = s(bool,X412) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) | $spl180),
  inference(cnf_transformation,[],[f7939_D])).
fof(f7939_D,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'F_0') = s(bool,X412) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f10978,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s(bool,X324) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X325),X326),s('type.closSem.state'(X325),X327))) | s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'F_0') = s(bool,X324)) ) | $spl70),
  inference(subsumption_resolution,[],[f10112,f242])).
fof(f10112,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s(bool,X324) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X325),X326),s('type.closSem.state'(X325),X327))) | s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'F_0') = s(bool,X323) | s(bool,'F_0') = s(bool,X324)) ) | $spl70),
  inference(superposition,[],[f3186,f7622])).
fof(f7933,plain,(
  ( ! [X364,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'F_0') = s(bool,X362) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) | $spl178),
  inference(cnf_transformation,[],[f7933_D])).
fof(f7933_D,plain,(
  ( ! [X364,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'F_0') = s(bool,X362) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f10971,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X285) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X286),X287),s('type.closSem.state'(X286),X288))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'F_0') = s(bool,X285)) ) | $spl64),
  inference(subsumption_resolution,[],[f10101,f242])).
fof(f10101,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X285) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X286),X287),s('type.closSem.state'(X286),X288))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'F_0') = s(bool,X284) | s(bool,'F_0') = s(bool,X285)) ) | $spl64),
  inference(superposition,[],[f3029,f7622])).
fof(f10935,plain,(
  ( ! [X561,X563,X564,X562] : (~p(s(bool,X564)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X561),X562),s('type.closSem.state'(X561),X563)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X563)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X563)))))) | s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X562))) != s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X563))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X562))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X563))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X562))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X563)))))) )),
  inference(subsumption_resolution,[],[f10019,f242])).
fof(f10019,plain,(
  ( ! [X561,X563,X564,X562] : (~p(s(bool,X564)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X561),X562),s('type.closSem.state'(X561),X563)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X561),X563)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X562))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X561),X563)))))) | s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X562))) != s('type.ffi.ffi_state'(X561),'const.closSem.state_ffi_1'(s('type.closSem.state'(X561),X563))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X562))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X561),X563))) | s(bool,'F_0') = s(bool,X564) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X562))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X561),X563)))))) )),
  inference(superposition,[],[f258,f7622])).
fof(f10934,plain,(
  ( ! [X553,X555,X556,X554] : (~p(s(bool,X556)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X553),X554),s('type.closSem.state'(X553),X555)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X555)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X554))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X555)))))) | s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X554))) != s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X555))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X554))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X555))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X555)))))) )),
  inference(subsumption_resolution,[],[f10017,f242])).
fof(f10017,plain,(
  ( ! [X553,X555,X556,X554] : (~p(s(bool,X556)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X553),X554),s('type.closSem.state'(X553),X555)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X553),X555)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X554))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X553),X555)))))) | s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X554))) != s('type.ffi.ffi_state'(X553),'const.closSem.state_ffi_1'(s('type.closSem.state'(X553),X555))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X554))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X553),X555))) | s(bool,'F_0') = s(bool,X556) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X554))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X553),X555)))))) )),
  inference(superposition,[],[f258,f7622])).
fof(f10933,plain,(
  ( ! [X546,X545,X547,X544] : (~p(s(bool,X547)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X544),X545),s('type.closSem.state'(X544),X546)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X545))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X546)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X546)))))) | s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X545))) != s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X546))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X545))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X546))) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X546)))))) )),
  inference(subsumption_resolution,[],[f10014,f242])).
fof(f10014,plain,(
  ( ! [X546,X545,X547,X544] : (~p(s(bool,X547)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X544),X545),s('type.closSem.state'(X544),X546)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X545))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X544),X546)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X544),X546)))))) | s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X545))) != s('type.ffi.ffi_state'(X544),'const.closSem.state_ffi_1'(s('type.closSem.state'(X544),X546))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X545))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X544),X546))) | s(bool,'F_0') = s(bool,X547) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X545))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X544),X546)))))) )),
  inference(superposition,[],[f258,f7622])).
fof(f10928,plain,(
  ( ! [X511,X512,X510,X513] : (~p(s(bool,X513)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X511))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X512)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X510),X511),s('type.closSem.state'(X510),X512)))) )),
  inference(subsumption_resolution,[],[f10006,f242])).
fof(f10006,plain,(
  ( ! [X511,X512,X510,X513] : (~p(s(bool,X513)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X511))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X510),X512)))))) | s(bool,'F_0') = s(bool,X513) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X510),X511),s('type.closSem.state'(X510),X512)))) )),
  inference(superposition,[],[f257,f7622])).
fof(f10927,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,X509)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X507))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X508)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X506),X507),s('type.closSem.state'(X506),X508)))) )),
  inference(subsumption_resolution,[],[f10005,f242])).
fof(f10005,plain,(
  ( ! [X509,X507,X508,X506] : (~p(s(bool,X509)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X507))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X506),X508)))))) | s(bool,'F_0') = s(bool,X509) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X506),X507),s('type.closSem.state'(X506),X508)))) )),
  inference(superposition,[],[f256,f7622])).
fof(f10926,plain,(
  ( ! [X505,X502,X504,X503] : (~p(s(bool,X505)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X503))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X504)))))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X502),X503),s('type.closSem.state'(X502),X504)))) )),
  inference(subsumption_resolution,[],[f10004,f242])).
fof(f10004,plain,(
  ( ! [X505,X502,X504,X503] : (~p(s(bool,X505)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X503))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X502),X504)))))) | s(bool,'F_0') = s(bool,X505) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X502),X503),s('type.closSem.state'(X502),X504)))) )),
  inference(superposition,[],[f255,f7622])).
fof(f10923,plain,(
  ( ! [X492,X493] : (s(bool,'T_0') != s(bool,X493) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X492))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X492)))) )),
  inference(subsumption_resolution,[],[f10001,f2590])).
fof(f10001,plain,(
  ( ! [X492,X493] : (s(bool,'T_0') != s(bool,X493) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X492))) | s(bool,'F_0') = s(bool,X493) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X492)))) )),
  inference(superposition,[],[f1072,f7622])).
fof(f10922,plain,(
  ( ! [X490,X491] : (s(bool,'T_0') != s(bool,X491) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X490))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X490)))) )),
  inference(subsumption_resolution,[],[f10000,f2590])).
fof(f10000,plain,(
  ( ! [X490,X491] : (s(bool,'T_0') != s(bool,X491) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X490))) | s(bool,'F_0') = s(bool,X491) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X490)))) )),
  inference(superposition,[],[f968,f7622])).
fof(f10921,plain,(
  ( ! [X488,X489] : (~p(s(bool,X489)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X488))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X488)))) )),
  inference(subsumption_resolution,[],[f9999,f242])).
fof(f9999,plain,(
  ( ! [X488,X489] : (~p(s(bool,X489)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X488))) | s(bool,'F_0') = s(bool,X489) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X488)))) )),
  inference(superposition,[],[f544,f7622])).
fof(f10920,plain,(
  ( ! [X484] : (s(bool,'T_0') != s(bool,X484) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f9996,f2590])).
fof(f9996,plain,(
  ( ! [X484] : (s(bool,'T_0') != s(bool,X484) | s(bool,'F_0') = s(bool,X484) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(superposition,[],[f321,f7622])).
fof(f10919,plain,(
  ( ! [X482,X483] : (s(bool,'T_0') != s(bool,X483) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X482))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X482),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f9995,f2590])).
fof(f9995,plain,(
  ( ! [X482,X483] : (s(bool,'T_0') != s(bool,X483) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X482))) | s(bool,'F_0') = s(bool,X483) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X482),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f1036,f7622])).
fof(f10918,plain,(
  ( ! [X480,X481] : (s(bool,'T_0') != s(bool,X481) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X480))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X480),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f9994,f2590])).
fof(f9994,plain,(
  ( ! [X480,X481] : (s(bool,'T_0') != s(bool,X481) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X480))) | s(bool,'F_0') = s(bool,X481) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X480),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f931,f7622])).
fof(f10917,plain,(
  ( ! [X478,X479] : (~p(s(bool,X479)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X478))) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X478),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f9993,f242])).
fof(f9993,plain,(
  ( ! [X478,X479] : (~p(s(bool,X479)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X478))) | s(bool,'F_0') = s(bool,X479) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X478),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f543,f7622])).
fof(f10910,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s(bool,X324) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X325),X326),s('type.closSem.state'(X325),X327))) | s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'F_0') = s(bool,X324)) ) | $spl70),
  inference(subsumption_resolution,[],[f9944,f242])).
fof(f9944,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s(bool,X324) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X325),X326),s('type.closSem.state'(X325),X327))) | s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'F_0') = s(bool,X324) | s(bool,'F_0') = s(bool,X323)) ) | $spl70),
  inference(superposition,[],[f3186,f7622])).
fof(f10903,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X285) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X286),X287),s('type.closSem.state'(X286),X288))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'F_0') = s(bool,X285)) ) | $spl64),
  inference(subsumption_resolution,[],[f9933,f242])).
fof(f9933,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X285) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X286),X287),s('type.closSem.state'(X286),X288))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'F_0') = s(bool,X285) | s(bool,'F_0') = s(bool,X284)) ) | $spl64),
  inference(superposition,[],[f3029,f7622])).
fof(f7736,plain,(
  ( ! [X55,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) | $spl136),
  inference(cnf_transformation,[],[f7736_D])).
fof(f7736_D,plain,(
  ( ! [X55,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f9702,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X34)) ) | $spl142),
  inference(subsumption_resolution,[],[f9676,f2660])).
fof(f9676,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X34)) ) | $spl142),
  inference(superposition,[],[f7752,f3476])).
fof(f9639,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X52) != s(bool,X55)) ) | $spl142),
  inference(resolution,[],[f7752,f262])).
fof(f9632,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl142),
  inference(resolution,[],[f7752,f203])).
fof(f9630,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl142),
  inference(resolution,[],[f7752,f202])).
fof(f9628,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl142),
  inference(resolution,[],[f7752,f201])).
fof(f9627,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl142),
  inference(resolution,[],[f7752,f199])).
fof(f9625,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl142),
  inference(resolution,[],[f7752,f197])).
fof(f9296,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl138),
  inference(equality_factoring,[],[f7737])).
fof(f9620,plain,(
  ( ! [X518,X521,X517,X519] : (s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X518))) = s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X519))) | s(bool,X521) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X517),X518),s('type.closSem.state'(X517),X519)))) ) | $spl298),
  inference(cnf_transformation,[],[f9620_D])).
fof(f9620_D,plain,(
  ( ! [X518,X521,X517,X519] : (s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X518))) = s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X519))) | s(bool,X521) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X517),X518),s('type.closSem.state'(X517),X519)))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f9619,plain,(
  ( ! [X520,X518,X521,X517,X519] : (s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X518))) = s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X519))) | ~p(s(bool,X520)) | s(bool,X521) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X517),X518),s('type.closSem.state'(X517),X519)))) ) | ($spl70 | $spl138)),
  inference(subsumption_resolution,[],[f9278,f247])).
fof(f9278,plain,(
  ( ! [X520,X518,X521,X517,X519] : (s(bool,'T_0') != s(bool,X520) | s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X518))) = s('type.ffi.ffi_state'(X517),'const.closSem.state_ffi_1'(s('type.closSem.state'(X517),X519))) | ~p(s(bool,X520)) | s(bool,X521) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X517),X518),s('type.closSem.state'(X517),X519)))) ) | ($spl70 | $spl138)),
  inference(superposition,[],[f3186,f7737])).
fof(f9617,plain,(
  ( ! [X505,X508,X504,X506] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X505))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X506))) | s(bool,X508) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X504),X505),s('type.closSem.state'(X504),X506)))) ) | $spl296),
  inference(cnf_transformation,[],[f9617_D])).
fof(f9617_D,plain,(
  ( ! [X505,X508,X504,X506] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X505))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X506))) | s(bool,X508) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X504),X505),s('type.closSem.state'(X504),X506)))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f9616,plain,(
  ( ! [X505,X507,X508,X504,X506] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X505))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X506))) | ~p(s(bool,X507)) | s(bool,X508) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X504),X505),s('type.closSem.state'(X504),X506)))) ) | ($spl64 | $spl138)),
  inference(subsumption_resolution,[],[f9275,f247])).
fof(f9275,plain,(
  ( ! [X505,X507,X508,X504,X506] : (s(bool,'T_0') != s(bool,X507) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X505))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X504),X506))) | ~p(s(bool,X507)) | s(bool,X508) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X504),X505),s('type.closSem.state'(X504),X506)))) ) | ($spl64 | $spl138)),
  inference(superposition,[],[f3029,f7737])).
fof(f8230,plain,(
  ( ! [X494,X495] : (s(bool,X494) != s(bool,X495) | s(bool,'T_0') != s(bool,X495)) ) | $spl260),
  inference(cnf_transformation,[],[f8230_D])).
fof(f8230_D,plain,(
  ( ! [X494,X495] : (s(bool,X494) != s(bool,X495) | s(bool,'T_0') != s(bool,X495)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f9607,plain,(
  ( ! [X455,X454,X456] : (s(bool,X454) != s(bool,X455) | s(bool,'T_0') != s(bool,X455) | s(bool,'F_0') = s(bool,X456)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9258,f2882])).
fof(f9258,plain,(
  ( ! [X455,X454,X456] : (s(bool,'T_0') != s(bool,X454) | s(bool,X454) != s(bool,X455) | s(bool,'T_0') != s(bool,X455) | s(bool,'F_0') = s(bool,X456)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f5870,f7737])).
fof(f5870,plain,(
  ( ! [X90,X89] : (s(bool,'F_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | s(bool,'T_0') != s(bool,X90)) ) | $spl58),
  inference(resolution,[],[f4702,f250])).
fof(f4702,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,X62)) | s(bool,X60) != s(bool,X62) | s(bool,X60) != s(bool,X61) | s(bool,'T_0') != s(bool,X61)) ) | $spl58),
  inference(resolution,[],[f3036,f246])).
fof(f3036,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) ) | $spl58),
  inference(superposition,[],[f2634,f243])).
fof(f7975,plain,(
  ( ! [X639,X640] : (s(bool,X639) != s(bool,X640) | s(bool,'F_0') = s(bool,X639)) ) | $spl198),
  inference(cnf_transformation,[],[f7975_D])).
fof(f7975_D,plain,(
  ( ! [X639,X640] : (s(bool,X639) != s(bool,X640) | s(bool,'F_0') = s(bool,X639)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f9606,plain,(
  ( ! [X447,X449,X448] : (s(bool,'F_0') = s(bool,X448) | s(bool,X447) != s(bool,X448) | s(bool,'F_0') = s(bool,X449)) ) | ($spl6 | $spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9256,f828])).
fof(f828,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f812,f263])).
fof(f263,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f210])).
fof(f210,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f812,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f766,f261])).
fof(f766,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(subsumption_resolution,[],[f740,f213])).
fof(f740,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,'F_0')) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X18)) ) | $spl6),
  inference(superposition,[],[f297,f261])).
fof(f9256,plain,(
  ( ! [X447,X449,X448] : (s(bool,'T_0') != s(bool,X447) | s(bool,'F_0') = s(bool,X448) | s(bool,X447) != s(bool,X448) | s(bool,'F_0') = s(bool,X449)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f3111,f7737])).
fof(f9605,plain,(
  ( ! [X445,X444,X446] : (s(bool,'T_0') != s(bool,X444) | s(bool,'T_0') = s(bool,X445) | s(bool,'F_0') = s(bool,X446)) ) | $spl138),
  inference(subsumption_resolution,[],[f9255,f2882])).
fof(f9255,plain,(
  ( ! [X445,X444,X446] : (s(bool,'T_0') != s(bool,X444) | s(bool,'T_0') = s(bool,X445) | s(bool,X444) = s(bool,X445) | s(bool,'F_0') = s(bool,X446)) ) | $spl138),
  inference(superposition,[],[f3073,f7737])).
fof(f3073,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,X48) = s(bool,X49)) )),
  inference(resolution,[],[f2660,f242])).
fof(f9604,plain,(
  ( ! [X441,X443,X442] : (s(bool,'T_0') != s(bool,X441) | p(s(bool,X442)) | s(bool,'F_0') = s(bool,X443)) ) | $spl138),
  inference(subsumption_resolution,[],[f9254,f2881])).
fof(f9254,plain,(
  ( ! [X441,X443,X442] : (s(bool,'T_0') != s(bool,X441) | s(bool,X441) = s(bool,X442) | p(s(bool,X442)) | s(bool,'F_0') = s(bool,X443)) ) | $spl138),
  inference(superposition,[],[f2915,f7737])).
fof(f9603,plain,(
  ( ! [X438,X440,X439] : (s(bool,X438) != s(bool,X439) | ~p(s(bool,X439)) | s(bool,'F_0') = s(bool,X440)) ) | ($spl6 | $spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9253,f766])).
fof(f9253,plain,(
  ( ! [X438,X440,X439] : (s(bool,'T_0') != s(bool,X438) | s(bool,X438) != s(bool,X439) | ~p(s(bool,X439)) | s(bool,'F_0') = s(bool,X440)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f2634,f7737])).
fof(f9306,plain,(
  ( ! [X436,X437] : (s(bool,'T_0') != s(bool,X436) | s(bool,'F_0') = s(bool,X437)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9252])).
fof(f9252,plain,(
  ( ! [X436,X437] : (s(bool,'T_0') != s(bool,X436) | s(bool,'T_0') != s(bool,X436) | s(bool,'F_0') = s(bool,X437)) ) | $spl138),
  inference(superposition,[],[f2590,f7737])).
fof(f9308,plain,(
  ( ! [X434,X433] : (s(bool,'T_0') = s(bool,X433) | s(bool,'F_0') = s(bool,X434)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9250])).
fof(f9250,plain,(
  ( ! [X434,X433] : (s(bool,'T_0') = s(bool,X433) | s(bool,'T_0') = s(bool,X433) | s(bool,'F_0') = s(bool,X434)) ) | $spl138),
  inference(superposition,[],[f261,f7737])).
fof(f9601,plain,(
  ( ! [X430,X431] : (p(s(bool,X430)) | s(bool,'F_0') = s(bool,X431)) ) | $spl138),
  inference(subsumption_resolution,[],[f9248,f246])).
fof(f9248,plain,(
  ( ! [X430,X431] : (s(bool,'T_0') = s(bool,X430) | p(s(bool,X430)) | s(bool,'F_0') = s(bool,X431)) ) | $spl138),
  inference(superposition,[],[f243,f7737])).
fof(f9600,plain,(
  ( ! [X428,X429] : (~p(s(bool,X428)) | s(bool,'F_0') = s(bool,X429)) ) | $spl138),
  inference(subsumption_resolution,[],[f9247,f247])).
fof(f9247,plain,(
  ( ! [X428,X429] : (s(bool,'T_0') != s(bool,X428) | ~p(s(bool,X428)) | s(bool,'F_0') = s(bool,X429)) ) | $spl138),
  inference(superposition,[],[f242,f7737])).
fof(f9367,plain,(
  ( ! [X48,X46] : (s(bool,X46) = s(bool,X48)) ) | $spl280),
  inference(cnf_transformation,[],[f9367_D])).
fof(f9367_D,plain,(
  ( ! [X48,X46] : (s(bool,X46) = s(bool,X48)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f9599,plain,(
  ( ! [X360,X362,X361,X363] : (s(bool,'T_0') != s(bool,X361) | s(bool,'T_0') = s(bool,X362) | s(bool,X360) = s(bool,X363)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9598,f7737])).
fof(f9598,plain,(
  ( ! [X360,X362,X361,X363] : (s(bool,'T_0') != s(bool,X361) | s(bool,X360) != s(bool,X362) | s(bool,'T_0') = s(bool,X362) | s(bool,X360) = s(bool,X363)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9223,f2590])).
fof(f9223,plain,(
  ( ! [X360,X362,X361,X363] : (s(bool,'T_0') != s(bool,X361) | s(bool,'F_0') = s(bool,X361) | s(bool,X360) != s(bool,X362) | s(bool,'T_0') = s(bool,X362) | s(bool,X360) = s(bool,X363)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f4623,f7737])).
fof(f9365,plain,(
  ( ! [X24,X25,X22] : (s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,X22) = s(bool,X25)) ) | $spl278),
  inference(cnf_transformation,[],[f9365_D])).
fof(f9365_D,plain,(
  ( ! [X24,X25,X22] : (s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,X22) = s(bool,X25)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f9309,plain,(
  ( ! [X337,X339,X336,X338] : (s(bool,'T_0') = s(bool,X337) | s(bool,X336) = s(bool,X338) | s(bool,'F_0') = s(bool,X338) | s(bool,X336) = s(bool,X339)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9217])).
fof(f9217,plain,(
  ( ! [X337,X339,X336,X338] : (s(bool,'T_0') = s(bool,X337) | s(bool,X336) = s(bool,X338) | s(bool,'T_0') = s(bool,X337) | s(bool,'F_0') = s(bool,X338) | s(bool,X336) = s(bool,X339)) ) | $spl138),
  inference(superposition,[],[f3476,f7737])).
fof(f9527,plain,(
  ( ! [X59,X57,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) = s(bool,X59)) ) | $spl292),
  inference(cnf_transformation,[],[f9527_D])).
fof(f9527_D,plain,(
  ( ! [X59,X57,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) = s(bool,X59)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f9595,plain,(
  ( ! [X292,X290,X293,X291] : (s(bool,X290) != s(bool,X291) | s(bool,'T_0') = s(bool,X292) | s(bool,X290) = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9594,f263])).
fof(f9594,plain,(
  ( ! [X292,X290,X293,X291] : (~p(s(bool,'T_0')) | s(bool,X290) != s(bool,X291) | s(bool,'T_0') = s(bool,X292) | s(bool,X290) = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9205,f7737])).
fof(f9205,plain,(
  ( ! [X292,X290,X293,X291] : (~p(s(bool,'T_0')) | s(bool,X290) != s(bool,X291) | s(bool,X291) != s(bool,X292) | s(bool,'T_0') = s(bool,X292) | s(bool,X290) = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f3037,f7737])).
fof(f3037,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl58),
  inference(superposition,[],[f2634,f261])).
fof(f9592,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,'T_0') != s(bool,X278) | p(s(bool,X279)) | p(s(bool,X277)) | s(bool,X277) = s(bool,X280)) ) | ($spl62 | $spl138)),
  inference(subsumption_resolution,[],[f9202,f2881])).
fof(f9202,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,'T_0') != s(bool,X278) | p(s(bool,X279)) | p(s(bool,X277)) | s(bool,X278) = s(bool,X279) | s(bool,X277) = s(bool,X280)) ) | ($spl62 | $spl138)),
  inference(superposition,[],[f2964,f7737])).
fof(f9591,plain,(
  ( ! [X269,X267,X268] : (s(bool,'T_0') = s(bool,X268) | s(bool,X267) = s(bool,X269)) ) | $spl138),
  inference(subsumption_resolution,[],[f9312,f7737])).
fof(f9312,plain,(
  ( ! [X269,X267,X268] : (s(bool,X267) != s(bool,X268) | s(bool,'T_0') = s(bool,X268) | s(bool,X267) = s(bool,X269)) ) | $spl138),
  inference(trivial_inequality_removal,[],[f9199])).
fof(f9199,plain,(
  ( ! [X269,X267,X268] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X267) != s(bool,X268) | s(bool,'T_0') = s(bool,X268) | s(bool,X267) = s(bool,X269)) ) | $spl138),
  inference(superposition,[],[f2882,f7737])).
fof(f9588,plain,(
  ( ! [X231,X229,X230] : (s(bool,'T_0') != s(bool,X230) | p(s(bool,X229)) | s(bool,X229) = s(bool,X231)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9186,f246])).
fof(f9186,plain,(
  ( ! [X231,X229,X230] : (s(bool,'T_0') != s(bool,X230) | p(s(bool,X229)) | ~p(s(bool,X230)) | s(bool,X229) = s(bool,X231)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f297,f7737])).
fof(f9529,plain,(
  ( ! [X90,X87,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X86) = s(bool,X90) | p(s(bool,X87))) ) | $spl294),
  inference(cnf_transformation,[],[f9529_D])).
fof(f9529_D,plain,(
  ( ! [X90,X87,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X86) = s(bool,X90) | p(s(bool,X87))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f9587,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228)) ) | $spl138),
  inference(subsumption_resolution,[],[f9185,f263])).
fof(f9185,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9586,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'T_0') != s(bool,X214) | p(s(bool,X215)) | p(s(bool,X213)) | s(bool,X213) = s(bool,X216)) ) | $spl138),
  inference(subsumption_resolution,[],[f9182,f246])).
fof(f9182,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'T_0') != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X214)) | p(s(bool,X213)) | s(bool,X213) = s(bool,X216)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9583,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169)) ) | $spl138),
  inference(subsumption_resolution,[],[f9162,f263])).
fof(f9162,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9581,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X145)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,X148) | s(bool,X145) = s(bool,X149)) ) | $spl138),
  inference(subsumption_resolution,[],[f9158,f246])).
fof(f9158,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X145)) | ~p(s(bool,X147)) | ~p(s(bool,X146)) | s(bool,X147) != s(bool,X148) | s(bool,X145) = s(bool,X149)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9578,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129)) ) | $spl138),
  inference(subsumption_resolution,[],[f9154,f263])).
fof(f9154,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9577,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X113)) | p(s(bool,X110)) | s(bool,X110) = s(bool,X114)) ) | $spl138),
  inference(subsumption_resolution,[],[f9151,f246])).
fof(f9151,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X111)) | ~p(s(bool,X113)) | p(s(bool,X110)) | s(bool,X110) = s(bool,X114)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9576,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104)) ) | $spl138),
  inference(subsumption_resolution,[],[f9149,f263])).
fof(f9149,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9573,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X77) != s(bool,X78) | s(bool,X75) = s(bool,X79)) ) | $spl138),
  inference(subsumption_resolution,[],[f9144,f246])).
fof(f9144,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | ~p(s(bool,X76)) | p(s(bool,X75)) | s(bool,X77) != s(bool,X78) | s(bool,X75) = s(bool,X79)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9572,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69)) ) | $spl138),
  inference(subsumption_resolution,[],[f9142,f263])).
fof(f9142,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9570,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,X53) | s(bool,X50) = s(bool,X54)) ) | $spl138),
  inference(subsumption_resolution,[],[f9139,f246])).
fof(f9139,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X52)) | ~p(s(bool,X51)) | s(bool,X52) != s(bool,X53) | s(bool,X50) = s(bool,X54)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9567,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34)) ) | $spl138),
  inference(subsumption_resolution,[],[f9135,f263])).
fof(f9135,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9564,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) = s(bool,X4)) ) | $spl138),
  inference(subsumption_resolution,[],[f9129,f246])).
fof(f9129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) = s(bool,X4)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9563,plain,(
  ( ! [X424,X426,X423,X425] : (p(s(bool,X423)) | s(bool,X424) != s(bool,X425) | s(bool,'T_0') != s(bool,X425) | s(bool,X423) = s(bool,X426)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9126,f2882])).
fof(f9126,plain,(
  ( ! [X424,X426,X423,X425] : (p(s(bool,X423)) | s(bool,'T_0') != s(bool,X424) | s(bool,X424) != s(bool,X425) | s(bool,'T_0') != s(bool,X425) | s(bool,X423) = s(bool,X426)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f4702,f7737])).
fof(f9562,plain,(
  ( ! [X420,X422,X421,X419] : (s(bool,X419) != s(bool,X420) | s(bool,'T_0') = s(bool,X421) | s(bool,X419) = s(bool,X422)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9561,f560])).
fof(f9561,plain,(
  ( ! [X420,X422,X421,X419] : (s(bool,X419) != s(bool,X420) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X421) | s(bool,X419) = s(bool,X422)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9125,f7737])).
fof(f9125,plain,(
  ( ! [X420,X422,X421,X419] : (s(bool,X419) != s(bool,X420) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X420) != s(bool,X421) | s(bool,'T_0') = s(bool,X421) | s(bool,X419) = s(bool,X422)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f4623,f7737])).
fof(f9560,plain,(
  ( ! [X412,X414,X413,X411] : (s(bool,'T_0') = s(bool,X413) | s(bool,'T_0') = s(bool,X412) | s(bool,X411) = s(bool,X414)) ) | $spl138),
  inference(subsumption_resolution,[],[f9319,f7737])).
fof(f9319,plain,(
  ( ! [X412,X414,X413,X411] : (s(bool,X411) != s(bool,X412) | s(bool,'T_0') = s(bool,X413) | s(bool,'T_0') = s(bool,X412) | s(bool,X411) = s(bool,X414)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9123])).
fof(f9123,plain,(
  ( ! [X412,X414,X413,X411] : (s(bool,X411) != s(bool,X412) | s(bool,'T_0') = s(bool,X413) | s(bool,'T_0') = s(bool,X413) | s(bool,'T_0') = s(bool,X412) | s(bool,X411) = s(bool,X414)) ) | $spl138),
  inference(superposition,[],[f4329,f7737])).
fof(f4329,plain,(
  ( ! [X339,X340,X338] : (s(bool,X338) != s(bool,X339) | s(bool,'T_0') = s(bool,X340) | s(bool,X339) = s(bool,X340) | s(bool,'T_0') = s(bool,X338)) )),
  inference(subsumption_resolution,[],[f3613,f560])).
fof(f3613,plain,(
  ( ! [X339,X340,X338] : (s(bool,X338) != s(bool,X339) | s(bool,'T_0') = s(bool,X340) | s(bool,X339) = s(bool,X340) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X338)) )),
  inference(superposition,[],[f3073,f3074])).
fof(f9559,plain,(
  ( ! [X405,X407,X404,X406] : (p(s(bool,X404)) | s(bool,'T_0') = s(bool,X405) | s(bool,'T_0') = s(bool,X406) | s(bool,X404) = s(bool,X407)) ) | $spl138),
  inference(subsumption_resolution,[],[f9121,f7737])).
fof(f9121,plain,(
  ( ! [X405,X407,X404,X406] : (p(s(bool,X404)) | s(bool,'T_0') = s(bool,X405) | s(bool,X405) != s(bool,X406) | s(bool,'T_0') = s(bool,X406) | s(bool,X404) = s(bool,X407)) ) | $spl138),
  inference(superposition,[],[f4328,f7737])).
fof(f4328,plain,(
  ( ! [X335,X337,X336] : (p(s(bool,X337)) | s(bool,X336) = s(bool,X337) | s(bool,X335) != s(bool,X336) | s(bool,'T_0') = s(bool,X335)) )),
  inference(subsumption_resolution,[],[f3612,f560])).
fof(f3612,plain,(
  ( ! [X335,X337,X336] : (s(bool,X335) != s(bool,X336) | s(bool,X336) = s(bool,X337) | p(s(bool,X337)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X335)) )),
  inference(superposition,[],[f2915,f3074])).
fof(f9320,plain,(
  ( ! [X385,X387,X388,X386] : (s(bool,X385) = s(bool,X386) | s(bool,'T_0') = s(bool,X387) | s(bool,'F_0') = s(bool,X386) | s(bool,X385) = s(bool,X388)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9116])).
fof(f9116,plain,(
  ( ! [X385,X387,X388,X386] : (s(bool,X385) = s(bool,X386) | s(bool,'T_0') = s(bool,X387) | s(bool,'T_0') = s(bool,X387) | s(bool,'F_0') = s(bool,X386) | s(bool,X385) = s(bool,X388)) ) | $spl138),
  inference(superposition,[],[f3476,f7737])).
fof(f9558,plain,(
  ( ! [X377,X374,X376,X375] : (p(s(bool,X374)) | s(bool,X375) != s(bool,X376) | s(bool,'F_0') = s(bool,X376) | s(bool,X374) = s(bool,X377)) ) | ($spl6 | $spl68 | $spl138)),
  inference(subsumption_resolution,[],[f9115,f828])).
fof(f9115,plain,(
  ( ! [X377,X374,X376,X375] : (p(s(bool,X374)) | s(bool,'T_0') != s(bool,X375) | s(bool,X375) != s(bool,X376) | s(bool,'F_0') = s(bool,X376) | s(bool,X374) = s(bool,X377)) ) | ($spl68 | $spl138)),
  inference(superposition,[],[f3130,f7737])).
fof(f9557,plain,(
  ( ! [X372,X370,X369,X371] : (p(s(bool,X369)) | s(bool,'T_0') != s(bool,X370) | s(bool,'T_0') = s(bool,X371) | s(bool,X369) = s(bool,X372)) ) | ($spl66 | $spl138)),
  inference(subsumption_resolution,[],[f9114,f2882])).
fof(f9114,plain,(
  ( ! [X372,X370,X369,X371] : (p(s(bool,X369)) | s(bool,'T_0') != s(bool,X370) | s(bool,X370) = s(bool,X371) | s(bool,'T_0') = s(bool,X371) | s(bool,X369) = s(bool,X372)) ) | ($spl66 | $spl138)),
  inference(superposition,[],[f3095,f7737])).
fof(f9556,plain,(
  ( ! [X344,X343,X345,X342] : (s(bool,X342) != s(bool,X343) | p(s(bool,X344)) | p(s(bool,X343)) | s(bool,X342) = s(bool,X345)) ) | ($spl62 | $spl138)),
  inference(subsumption_resolution,[],[f9108,f246])).
fof(f9108,plain,(
  ( ! [X344,X343,X345,X342] : (s(bool,X342) != s(bool,X343) | p(s(bool,X344)) | p(s(bool,X343)) | s(bool,'T_0') = s(bool,X344) | s(bool,X342) = s(bool,X345)) ) | ($spl62 | $spl138)),
  inference(superposition,[],[f2964,f7737])).
fof(f9555,plain,(
  ( ! [X333,X332,X334] : (p(s(bool,X332)) | s(bool,'T_0') != s(bool,X333) | s(bool,X332) = s(bool,X334)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9105,f2590])).
fof(f9105,plain,(
  ( ! [X333,X332,X334] : (p(s(bool,X332)) | s(bool,'T_0') != s(bool,X333) | s(bool,'F_0') = s(bool,X333) | s(bool,X332) = s(bool,X334)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f2856,f7737])).
fof(f9322,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,X329)) | s(bool,'T_0') = s(bool,X330) | s(bool,X329) = s(bool,X331)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f9104])).
fof(f9104,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,X329)) | s(bool,'T_0') = s(bool,X330) | s(bool,'T_0') = s(bool,X330) | s(bool,X329) = s(bool,X331)) ) | $spl138),
  inference(superposition,[],[f2660,f7737])).
fof(f9554,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,X326)) | p(s(bool,X327)) | s(bool,X326) = s(bool,X328)) ) | $spl138),
  inference(subsumption_resolution,[],[f9103,f246])).
fof(f9103,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,X326)) | p(s(bool,X327)) | s(bool,'T_0') = s(bool,X327) | s(bool,X326) = s(bool,X328)) ) | $spl138),
  inference(superposition,[],[f2659,f7737])).
fof(f9553,plain,(
  ( ! [X325,X323,X324] : (p(s(bool,X323)) | p(s(bool,X324)) | s(bool,X323) = s(bool,X325)) ) | $spl138),
  inference(subsumption_resolution,[],[f9102,f246])).
fof(f9102,plain,(
  ( ! [X325,X323,X324] : (p(s(bool,X323)) | p(s(bool,X324)) | s(bool,'T_0') = s(bool,X324) | s(bool,X323) = s(bool,X325)) ) | $spl138),
  inference(superposition,[],[f2659,f7737])).
fof(f9552,plain,(
  ( ! [X316,X318,X317] : (s(bool,'T_0') = s(bool,X317) | s(bool,X316) = s(bool,X318)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9551,f560])).
fof(f9551,plain,(
  ( ! [X316,X318,X317] : (s(bool,'T_0') = s(bool,X317) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X316) = s(bool,X318)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9100,f7737])).
fof(f9100,plain,(
  ( ! [X316,X318,X317] : (s(bool,X316) != s(bool,X317) | s(bool,'T_0') = s(bool,X317) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X316) = s(bool,X318)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f828,f7737])).
fof(f9550,plain,(
  ( ! [X313,X315,X314] : (s(bool,'T_0') = s(bool,X314) | s(bool,X313) = s(bool,X315)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9549,f263])).
fof(f9549,plain,(
  ( ! [X313,X315,X314] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X314) | s(bool,X313) = s(bool,X315)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9099,f7737])).
fof(f9099,plain,(
  ( ! [X313,X315,X314] : (s(bool,X313) != s(bool,X314) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X314) | s(bool,X313) = s(bool,X315)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f766,f7737])).
fof(f9547,plain,(
  ( ! [X294,X296,X293,X295] : (p(s(bool,X293)) | p(s(bool,X294)) | s(bool,'T_0') != s(bool,X295) | s(bool,X293) = s(bool,X296)) ) | $spl138),
  inference(subsumption_resolution,[],[f9093,f246])).
fof(f9093,plain,(
  ( ! [X294,X296,X293,X295] : (p(s(bool,X293)) | p(s(bool,X294)) | ~p(s(bool,X295)) | s(bool,'T_0') != s(bool,X295) | s(bool,X293) = s(bool,X296)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9546,plain,(
  ( ! [X288,X285,X287,X286] : (s(bool,X285) != s(bool,X286) | p(s(bool,X287)) | p(s(bool,X286)) | s(bool,X285) = s(bool,X288)) ) | $spl138),
  inference(subsumption_resolution,[],[f9091,f263])).
fof(f9091,plain,(
  ( ! [X288,X285,X287,X286] : (s(bool,X285) != s(bool,X286) | p(s(bool,X287)) | ~p(s(bool,'T_0')) | p(s(bool,X286)) | s(bool,X285) = s(bool,X288)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9543,plain,(
  ( ! [X239,X237,X240,X238,X236] : (p(s(bool,X236)) | ~p(s(bool,X237)) | s(bool,'T_0') != s(bool,X238) | s(bool,X237) != s(bool,X239) | s(bool,X236) = s(bool,X240)) ) | $spl138),
  inference(subsumption_resolution,[],[f9073,f246])).
fof(f9073,plain,(
  ( ! [X239,X237,X240,X238,X236] : (p(s(bool,X236)) | ~p(s(bool,X237)) | ~p(s(bool,X238)) | s(bool,'T_0') != s(bool,X238) | s(bool,X237) != s(bool,X239) | s(bool,X236) = s(bool,X240)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9542,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,X231) != s(bool,X232) | p(s(bool,X232)) | ~p(s(bool,X233)) | s(bool,X233) != s(bool,X234) | s(bool,X231) = s(bool,X235)) ) | $spl138),
  inference(subsumption_resolution,[],[f9072,f263])).
fof(f9072,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,X231) != s(bool,X232) | p(s(bool,X232)) | ~p(s(bool,X233)) | ~p(s(bool,'T_0')) | s(bool,X233) != s(bool,X234) | s(bool,X231) = s(bool,X235)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9540,plain,(
  ( ! [X198,X196,X200,X199,X197] : (p(s(bool,X196)) | p(s(bool,X197)) | ~p(s(bool,X199)) | s(bool,'T_0') != s(bool,X198) | s(bool,X196) = s(bool,X200)) ) | $spl138),
  inference(subsumption_resolution,[],[f9065,f246])).
fof(f9065,plain,(
  ( ! [X198,X196,X200,X199,X197] : (p(s(bool,X196)) | p(s(bool,X197)) | ~p(s(bool,X198)) | ~p(s(bool,X199)) | s(bool,'T_0') != s(bool,X198) | s(bool,X196) = s(bool,X200)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9539,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,X188)) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,X186) = s(bool,X190)) ) | $spl138),
  inference(subsumption_resolution,[],[f9063,f263])).
fof(f9063,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,X188)) | ~p(s(bool,'T_0')) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,X186) = s(bool,X190)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9538,plain,(
  ( ! [X175,X173,X171,X174,X172] : (p(s(bool,X171)) | p(s(bool,X172)) | s(bool,X172) != s(bool,X174) | s(bool,'T_0') != s(bool,X173) | s(bool,X171) = s(bool,X175)) ) | $spl138),
  inference(subsumption_resolution,[],[f9060,f246])).
fof(f9060,plain,(
  ( ! [X175,X173,X171,X174,X172] : (p(s(bool,X171)) | p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X172) != s(bool,X174) | s(bool,'T_0') != s(bool,X173) | s(bool,X171) = s(bool,X175)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9536,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | p(s(bool,X152)) | s(bool,X153) != s(bool,X154) | s(bool,X151) = s(bool,X155)) ) | $spl138),
  inference(subsumption_resolution,[],[f9056,f263])).
fof(f9056,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,'T_0')) | p(s(bool,X152)) | s(bool,X153) != s(bool,X154) | s(bool,X151) = s(bool,X155)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9534,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,'T_0') != s(bool,X138) | s(bool,X137) != s(bool,X139) | s(bool,X136) = s(bool,X140)) ) | $spl138),
  inference(subsumption_resolution,[],[f9053,f246])).
fof(f9053,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X136)) | ~p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,'T_0') != s(bool,X138) | s(bool,X137) != s(bool,X139) | s(bool,X136) = s(bool,X140)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9533,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,X126) != s(bool,X127) | p(s(bool,X127)) | ~p(s(bool,X128)) | s(bool,X128) != s(bool,X129) | s(bool,X126) = s(bool,X130)) ) | $spl138),
  inference(subsumption_resolution,[],[f9051,f263])).
fof(f9051,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,X126) != s(bool,X127) | p(s(bool,X127)) | ~p(s(bool,X128)) | ~p(s(bool,'T_0')) | s(bool,X128) != s(bool,X129) | s(bool,X126) = s(bool,X130)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9531,plain,(
  ( ! [X103,X101,X105,X102,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X102) != s(bool,X104) | s(bool,'T_0') != s(bool,X103) | s(bool,X101) = s(bool,X105)) ) | $spl138),
  inference(subsumption_resolution,[],[f9046,f246])).
fof(f9046,plain,(
  ( ! [X103,X101,X105,X102,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,'T_0') != s(bool,X103) | s(bool,X101) = s(bool,X105)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9528,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | p(s(bool,X88)) | p(s(bool,X87)) | s(bool,X88) != s(bool,X89) | s(bool,X86) = s(bool,X90)) ) | $spl138),
  inference(subsumption_resolution,[],[f9043,f263])).
fof(f9043,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | p(s(bool,X88)) | ~p(s(bool,'T_0')) | p(s(bool,X87)) | s(bool,X88) != s(bool,X89) | s(bool,X86) = s(bool,X90)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9526,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,X58) | s(bool,X56) = s(bool,X59)) ) | $spl138),
  inference(subsumption_resolution,[],[f9038,f7737])).
fof(f9038,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,X58) | s(bool,X56) = s(bool,X59)) ) | $spl138),
  inference(superposition,[],[f2882,f7737])).
fof(f9028,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) = s(bool,X20)) ) | $spl138),
  inference(superposition,[],[f240,f7737])).
fof(f240,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9] : ((p(s(bool,'T_0')) & p(s(bool,X9))) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.AND_CLAUSES|split|0')).
fof(f9027,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) = s(bool,X17)) ) | $spl138),
  inference(superposition,[],[f238,f7737])).
fof(f238,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ((p(s(bool,X9)) & p(s(bool,'T_0'))) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.AND_CLAUSES|split|1')).
fof(f9026,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) = s(bool,X14)) ) | $spl138),
  inference(superposition,[],[f225,f7737])).
fof(f225,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CLAUSES|split|3')).
fof(f9025,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) = s(bool,X11)) ) | $spl138),
  inference(superposition,[],[f224,f7737])).
fof(f224,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f9024,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) = s(bool,X8)) ) | $spl138),
  inference(superposition,[],[f221,f7737])).
fof(f221,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CLAUSES|split|1')).
fof(f9023,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5)) ) | $spl138),
  inference(superposition,[],[f218,f7737])).
fof(f218,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ((p(s(bool,'T_0')) => p(s(bool,X9))) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CLAUSES|split|0')).
fof(f9022,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X2)) ) | $spl138),
  inference(superposition,[],[f216,f7737])).
fof(f216,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9] : ((p(s(bool,'F_0')) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f4333,plain,(
  ( ! [X372,X373,X371] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373)))))) ) | $spl82),
  inference(cnf_transformation,[],[f4333_D])).
fof(f4333_D,plain,(
  ( ! [X372,X373,X371] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9018,plain,(
  ( ! [X565,X563,X564,X562] : (p(s(bool,X565)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X562),X563),s('type.closSem.state'(X562),X564)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X562),X563))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X562),X564)))))) ) | $spl138),
  inference(superposition,[],[f256,f7737])).
fof(f8196,plain,(
  ( ! [X679,X680,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) | $spl256),
  inference(cnf_transformation,[],[f8196_D])).
fof(f8196_D,plain,(
  ( ! [X679,X680,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f9017,plain,(
  ( ! [X561,X559,X560,X558] : (~p(s(bool,X561)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X558),X559),s('type.closSem.state'(X558),X560)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X558),X559))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X558),X560)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X558),X559))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X558),X560)))))) | s('type.ffi.ffi_state'(X558),'const.closSem.state_ffi_1'(s('type.closSem.state'(X558),X559))) != s('type.ffi.ffi_state'(X558),'const.closSem.state_ffi_1'(s('type.closSem.state'(X558),X560))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X558),X559))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X558),X560))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X558),X559))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X558),X560)))))) ) | $spl138),
  inference(superposition,[],[f258,f7737])).
fof(f4587,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) ) | $spl88),
  inference(cnf_transformation,[],[f4587_D])).
fof(f4587_D,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9016,plain,(
  ( ! [X557,X555,X556,X554] : (p(s(bool,X557)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X554),X555),s('type.closSem.state'(X554),X556)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X554),X555))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X554),X556)))))) ) | $spl138),
  inference(superposition,[],[f255,f7737])).
fof(f8194,plain,(
  ( ! [X668,X670,X669] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) | $spl254),
  inference(cnf_transformation,[],[f8194_D])).
fof(f8194_D,plain,(
  ( ! [X668,X670,X669] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f9015,plain,(
  ( ! [X553,X551,X552,X550] : (~p(s(bool,X553)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X550),X551),s('type.closSem.state'(X550),X552)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X550),X551))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X550),X552)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X550),X551))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X550),X552)))))) | s('type.ffi.ffi_state'(X550),'const.closSem.state_ffi_1'(s('type.closSem.state'(X550),X551))) != s('type.ffi.ffi_state'(X550),'const.closSem.state_ffi_1'(s('type.closSem.state'(X550),X552))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X550),X551))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X550),X552))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X550),X551))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X550),X552)))))) ) | $spl138),
  inference(superposition,[],[f258,f7737])).
fof(f8446,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | $spl274),
  inference(cnf_transformation,[],[f8446_D])).
fof(f8446_D,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f9014,plain,(
  ( ! [X549] : (p(s(bool,X549)) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) ) | $spl138),
  inference(superposition,[],[f269,f7737])).
fof(f4331,plain,(
  ( ! [X368,X367,X369] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369)))))) ) | $spl80),
  inference(cnf_transformation,[],[f4331_D])).
fof(f4331_D,plain,(
  ( ! [X368,X367,X369] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9013,plain,(
  ( ! [X546,X545,X547,X548] : (p(s(bool,X548)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X545),X546),s('type.closSem.state'(X545),X547)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X545),X546))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X545),X547)))))) ) | $spl138),
  inference(superposition,[],[f257,f7737])).
fof(f8192,plain,(
  ( ! [X660,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) | $spl252),
  inference(cnf_transformation,[],[f8192_D])).
fof(f8192_D,plain,(
  ( ! [X660,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f9012,plain,(
  ( ! [X542,X541,X543,X544] : (~p(s(bool,X544)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X541),X542),s('type.closSem.state'(X541),X543)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X541),X542))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X541),X543)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X541),X542))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X541),X543)))))) | s('type.ffi.ffi_state'(X541),'const.closSem.state_ffi_1'(s('type.closSem.state'(X541),X542))) != s('type.ffi.ffi_state'(X541),'const.closSem.state_ffi_1'(s('type.closSem.state'(X541),X543))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X541),X542))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X541),X543))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X541),X542))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X541),X543)))))) ) | $spl138),
  inference(superposition,[],[f258,f7737])).
fof(f7776,plain,(
  ( ! [X94,X92,X93] : (s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X93))) = s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X94)))) ) | $spl146),
  inference(cnf_transformation,[],[f7776_D])).
fof(f7776_D,plain,(
  ( ! [X94,X92,X93] : (s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X93))) = s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X94)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f9525,plain,(
  ( ! [X540,X536,X538,X537,X539] : (s(bool,X539) != s(bool,X540) | s('type.ffi.ffi_state'(X536),'const.closSem.state_ffi_1'(s('type.closSem.state'(X536),X537))) = s('type.ffi.ffi_state'(X536),'const.closSem.state_ffi_1'(s('type.closSem.state'(X536),X538))) | ~p(s(bool,X540))) ) | ($spl70 | $spl138)),
  inference(subsumption_resolution,[],[f9011,f3144])).
fof(f3144,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),X61),s('type.closSem.state'(X60),X62))) | s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X61))) = s('type.ffi.ffi_state'(X60),'const.closSem.state_ffi_1'(s('type.closSem.state'(X60),X62)))) )),
  inference(resolution,[],[f254,f246])).
fof(f9011,plain,(
  ( ! [X540,X536,X538,X537,X539] : (s(bool,X539) != s(bool,X540) | s('type.ffi.ffi_state'(X536),'const.closSem.state_ffi_1'(s('type.closSem.state'(X536),X537))) = s('type.ffi.ffi_state'(X536),'const.closSem.state_ffi_1'(s('type.closSem.state'(X536),X538))) | ~p(s(bool,X540)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X536),X537),s('type.closSem.state'(X536),X538)))) ) | ($spl70 | $spl138)),
  inference(superposition,[],[f3186,f7737])).
fof(f9524,plain,(
  ( ! [X532,X534,X533,X535] : (s(bool,'T_0') != s(bool,X535) | s('type.ffi.ffi_state'(X532),'const.closSem.state_ffi_1'(s('type.closSem.state'(X532),X533))) = s('type.ffi.ffi_state'(X532),'const.closSem.state_ffi_1'(s('type.closSem.state'(X532),X534)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9010,f3144])).
fof(f9010,plain,(
  ( ! [X532,X534,X533,X535] : (s(bool,'T_0') != s(bool,X535) | s('type.ffi.ffi_state'(X532),'const.closSem.state_ffi_1'(s('type.closSem.state'(X532),X533))) = s('type.ffi.ffi_state'(X532),'const.closSem.state_ffi_1'(s('type.closSem.state'(X532),X534))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X532),X533),s('type.closSem.state'(X532),X534)))) ) | $spl138),
  inference(superposition,[],[f3144,f7737])).
fof(f9523,plain,(
  ( ! [X528,X530,X529,X531] : (s(bool,'F_0') = s(bool,X531) | s('type.ffi.ffi_state'(X528),'const.closSem.state_ffi_1'(s('type.closSem.state'(X528),X529))) = s('type.ffi.ffi_state'(X528),'const.closSem.state_ffi_1'(s('type.closSem.state'(X528),X530)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9009,f3144])).
fof(f9009,plain,(
  ( ! [X528,X530,X529,X531] : (s(bool,'F_0') = s(bool,X531) | s('type.ffi.ffi_state'(X528),'const.closSem.state_ffi_1'(s('type.closSem.state'(X528),X529))) = s('type.ffi.ffi_state'(X528),'const.closSem.state_ffi_1'(s('type.closSem.state'(X528),X530))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X528),X529),s('type.closSem.state'(X528),X530)))) ) | $spl138),
  inference(superposition,[],[f3143,f7737])).
fof(f3143,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X57),X58),s('type.closSem.state'(X57),X59))) | s('type.ffi.ffi_state'(X57),'const.closSem.state_ffi_1'(s('type.closSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.closSem.state_ffi_1'(s('type.closSem.state'(X57),X59)))) )),
  inference(resolution,[],[f254,f243])).
fof(f7778,plain,(
  ( ! [X99,X97,X98] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X99)))) ) | $spl148),
  inference(cnf_transformation,[],[f7778_D])).
fof(f7778_D,plain,(
  ( ! [X99,X97,X98] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X99)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f9522,plain,(
  ( ! [X524,X526,X525,X527,X523] : (s(bool,X526) != s(bool,X527) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X523),X524))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X523),X525))) | ~p(s(bool,X527))) ) | ($spl64 | $spl138)),
  inference(subsumption_resolution,[],[f9008,f2981])).
fof(f2981,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X60),X61),s('type.closSem.state'(X60),X62))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X60),X61))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X60),X62)))) )),
  inference(resolution,[],[f253,f246])).
fof(f9008,plain,(
  ( ! [X524,X526,X525,X527,X523] : (s(bool,X526) != s(bool,X527) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X523),X524))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X523),X525))) | ~p(s(bool,X527)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X523),X524),s('type.closSem.state'(X523),X525)))) ) | ($spl64 | $spl138)),
  inference(superposition,[],[f3029,f7737])).
fof(f9521,plain,(
  ( ! [X520,X522,X521,X519] : (s(bool,'T_0') != s(bool,X522) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X519),X520))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X519),X521)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9007,f2981])).
fof(f9007,plain,(
  ( ! [X520,X522,X521,X519] : (s(bool,'T_0') != s(bool,X522) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X519),X520))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X519),X521))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X519),X520),s('type.closSem.state'(X519),X521)))) ) | $spl138),
  inference(superposition,[],[f2981,f7737])).
fof(f9520,plain,(
  ( ! [X516,X518,X517,X515] : (s(bool,'F_0') = s(bool,X518) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X515),X516))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X515),X517)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9006,f2981])).
fof(f9006,plain,(
  ( ! [X516,X518,X517,X515] : (s(bool,'F_0') = s(bool,X518) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X515),X516))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X515),X517))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X515),X516),s('type.closSem.state'(X515),X517)))) ) | $spl138),
  inference(superposition,[],[f2980,f7737])).
fof(f2980,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X57),X58),s('type.closSem.state'(X57),X59))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X57),X58))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X57),X59)))) )),
  inference(resolution,[],[f253,f243])).
fof(f5005,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X21)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X20))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X21))) | s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X21))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X21))))))) ) | $spl102),
  inference(cnf_transformation,[],[f5005_D])).
fof(f5005_D,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X21)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X20))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X21))) | s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X21))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X21))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f9005,plain,(
  ( ! [X511,X512,X514,X513] : (p(s(bool,X514)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X512))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X513)))))) | s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X512))) != s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X513))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X512))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X513))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X511),X512),s('type.closSem.state'(X511),X513)))) ) | $spl138),
  inference(superposition,[],[f258,f7737])).
fof(f8184,plain,(
  ( ! [X617,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) | $spl250),
  inference(cnf_transformation,[],[f8184_D])).
fof(f8184_D,plain,(
  ( ! [X617,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f9004,plain,(
  ( ! [X509,X507,X508,X510] : (~p(s(bool,X510)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X507),X508))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X507),X509)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X507),X508),s('type.closSem.state'(X507),X509)))) ) | $spl138),
  inference(superposition,[],[f257,f7737])).
fof(f8183,plain,(
  ( ! [X613,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) | $spl248),
  inference(cnf_transformation,[],[f8183_D])).
fof(f8183_D,plain,(
  ( ! [X613,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f9003,plain,(
  ( ! [X505,X504,X506,X503] : (~p(s(bool,X506)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X503),X504))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X503),X505)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X503),X504),s('type.closSem.state'(X503),X505)))) ) | $spl138),
  inference(superposition,[],[f256,f7737])).
fof(f8182,plain,(
  ( ! [X606,X607,X608] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) | $spl246),
  inference(cnf_transformation,[],[f8182_D])).
fof(f8182_D,plain,(
  ( ! [X606,X607,X608] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f9002,plain,(
  ( ! [X499,X500,X502,X501] : (~p(s(bool,X502)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X499),X500))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X499),X501)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X499),X500),s('type.closSem.state'(X499),X501)))) ) | $spl138),
  inference(superposition,[],[f255,f7737])).
fof(f9519,plain,(
  ( ! [X496,X498,X495,X497] : (~p(s(bool,X498)) | s('type.ffi.ffi_state'(X495),'const.closSem.state_ffi_1'(s('type.closSem.state'(X495),X496))) = s('type.ffi.ffi_state'(X495),'const.closSem.state_ffi_1'(s('type.closSem.state'(X495),X497)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9001,f3144])).
fof(f9001,plain,(
  ( ! [X496,X498,X495,X497] : (~p(s(bool,X498)) | s('type.ffi.ffi_state'(X495),'const.closSem.state_ffi_1'(s('type.closSem.state'(X495),X496))) = s('type.ffi.ffi_state'(X495),'const.closSem.state_ffi_1'(s('type.closSem.state'(X495),X497))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X495),X496),s('type.closSem.state'(X495),X497)))) ) | $spl138),
  inference(superposition,[],[f254,f7737])).
fof(f9518,plain,(
  ( ! [X492,X494,X493,X491] : (~p(s(bool,X494)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X491),X492))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X491),X493)))) ) | $spl138),
  inference(subsumption_resolution,[],[f9000,f2981])).
fof(f9000,plain,(
  ( ! [X492,X494,X493,X491] : (~p(s(bool,X494)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X491),X492))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X491),X493))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X491),X492),s('type.closSem.state'(X491),X493)))) ) | $spl138),
  inference(superposition,[],[f253,f7737])).
fof(f2198,plain,(
  ( ! [X51] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X51)))) ) | $spl48),
  inference(cnf_transformation,[],[f2198_D])).
fof(f2198_D,plain,(
  ( ! [X51] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X51)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9517,plain,(
  ( ! [X490,X489] : (s(bool,'T_0') != s(bool,X490) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X489)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8999,f1072])).
fof(f8999,plain,(
  ( ! [X490,X489] : (s(bool,'T_0') != s(bool,X490) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X489))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X489)))) ) | $spl138),
  inference(superposition,[],[f1072,f7737])).
fof(f2200,plain,(
  ( ! [X55] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X55)))) ) | $spl50),
  inference(cnf_transformation,[],[f2200_D])).
fof(f2200_D,plain,(
  ( ! [X55] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X55)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9516,plain,(
  ( ! [X488,X487] : (s(bool,'T_0') != s(bool,X488) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X487)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8998,f968])).
fof(f8998,plain,(
  ( ! [X488,X487] : (s(bool,'T_0') != s(bool,X488) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X487))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X487)))) ) | $spl138),
  inference(superposition,[],[f968,f7737])).
fof(f9515,plain,(
  ( ! [X486,X485] : (~p(s(bool,X486)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X485)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8997,f968])).
fof(f8997,plain,(
  ( ! [X486,X485] : (~p(s(bool,X486)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X485))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X485)))) ) | $spl138),
  inference(superposition,[],[f544,f7737])).
fof(f9512,plain,(
  ( ! [X480,X479] : (s(bool,'T_0') != s(bool,X480) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X479)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8993,f1036])).
fof(f8993,plain,(
  ( ! [X480,X479] : (s(bool,'T_0') != s(bool,X480) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X479))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X479),s('type.closSem.state'(sK0),sK3)))) ) | $spl138),
  inference(superposition,[],[f1036,f7737])).
fof(f9511,plain,(
  ( ! [X478,X477] : (s(bool,'T_0') != s(bool,X478) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X477)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8992,f931])).
fof(f8992,plain,(
  ( ! [X478,X477] : (s(bool,'T_0') != s(bool,X478) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X477))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X477),s('type.closSem.state'(sK0),sK3)))) ) | $spl138),
  inference(superposition,[],[f931,f7737])).
fof(f9510,plain,(
  ( ! [X476,X475] : (~p(s(bool,X476)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X475)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8991,f931])).
fof(f8991,plain,(
  ( ! [X476,X475] : (~p(s(bool,X476)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X475))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X475),s('type.closSem.state'(sK0),sK3)))) ) | $spl138),
  inference(superposition,[],[f543,f7737])).
fof(f9501,plain,(
  ( ! [X379,X380,X378,X381] : (p(s(bool,X379)) | s(bool,X380) != s(bool,X381) | s(bool,'T_0') != s(bool,X381) | s(bool,'T_0') = s(bool,X378)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8957,f7737])).
fof(f8957,plain,(
  ( ! [X379,X380,X378,X381] : (p(s(bool,X379)) | s(bool,X378) != s(bool,X380) | s(bool,X380) != s(bool,X381) | s(bool,'T_0') != s(bool,X381) | s(bool,'T_0') = s(bool,X378)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f4702,f7737])).
fof(f9500,plain,(
  ( ! [X372,X370,X371] : (s(bool,X371) != s(bool,X372) | s(bool,'F_0') = s(bool,X372) | s(bool,'T_0') = s(bool,X370)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9499,f7737])).
fof(f9499,plain,(
  ( ! [X372,X370,X373,X371] : (s(bool,X371) != s(bool,X372) | s(bool,'F_0') = s(bool,X372) | s(bool,X370) != s(bool,X373) | s(bool,'T_0') = s(bool,X370)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f8955,f2882])).
fof(f8955,plain,(
  ( ! [X372,X370,X373,X371] : (s(bool,X371) != s(bool,X372) | s(bool,'F_0') = s(bool,X372) | s(bool,X370) != s(bool,X373) | s(bool,'T_0') = s(bool,X373) | s(bool,'T_0') = s(bool,X370)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f4623,f7737])).
fof(f9328,plain,(
  ( ! [X350,X352,X351,X353] : (s(bool,X351) = s(bool,X352) | s(bool,X352) = s(bool,X353) | s(bool,'T_0') = s(bool,X350) | s(bool,'F_0') = s(bool,X353)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8950])).
fof(f8950,plain,(
  ( ! [X350,X352,X351,X353] : (s(bool,X351) = s(bool,X352) | s(bool,X352) = s(bool,X353) | s(bool,'T_0') = s(bool,X350) | s(bool,'F_0') = s(bool,X353) | s(bool,'T_0') = s(bool,X350)) ) | $spl138),
  inference(superposition,[],[f3476,f7737])).
fof(f296,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) ) | $spl4),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9329,plain,(
  ( ! [X329,X328] : (s(bool,X329) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X328)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8944])).
fof(f8944,plain,(
  ( ! [X329,X328] : (s(bool,X329) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X328) | s(bool,'T_0') = s(bool,X328)) ) | $spl138),
  inference(superposition,[],[f3407,f7737])).
fof(f9497,plain,(
  ( ! [X294,X296,X293] : (~p(s(bool,X294)) | s(bool,'T_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f9496,f7737])).
fof(f9496,plain,(
  ( ! [X294,X296,X293,X295] : (~p(s(bool,X294)) | s(bool,X293) != s(bool,X295) | s(bool,'T_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8934,f4329])).
fof(f8934,plain,(
  ( ! [X294,X296,X293,X295] : (~p(s(bool,X294)) | s(bool,X293) != s(bool,X295) | s(bool,X295) != s(bool,X296) | s(bool,'T_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X293)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f3037,f7737])).
fof(f8146,plain,(
  ( ! [X370,X367,X369] : (s(bool,X367) != s(bool,X369) | p(s(bool,X370)) | s(bool,X369) != s(bool,X370)) ) | $spl240),
  inference(cnf_transformation,[],[f8146_D])).
fof(f8146_D,plain,(
  ( ! [X370,X367,X369] : (s(bool,X367) != s(bool,X369) | p(s(bool,X370)) | s(bool,X369) != s(bool,X370)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f9495,plain,(
  ( ! [X292,X290,X289,X291] : (~p(s(bool,X290)) | s(bool,X289) != s(bool,X291) | s(bool,X291) != s(bool,X292) | p(s(bool,X292))) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8933,f4702])).
fof(f8933,plain,(
  ( ! [X292,X290,X289,X291] : (~p(s(bool,X290)) | s(bool,X289) != s(bool,X291) | s(bool,X291) != s(bool,X292) | p(s(bool,X292)) | s(bool,'T_0') = s(bool,X289)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f3036,f7737])).
fof(f9494,plain,(
  ( ! [X282,X281,X283] : (s(bool,X281) != s(bool,X282) | s(bool,X282) = s(bool,X283) | p(s(bool,X283))) ) | $spl290),
  inference(cnf_transformation,[],[f9494_D])).
fof(f9494_D,plain,(
  ( ! [X282,X281,X283] : (s(bool,X281) != s(bool,X282) | s(bool,X282) = s(bool,X283) | p(s(bool,X283))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f9493,plain,(
  ( ! [X280,X282,X281,X283] : (s(bool,X281) != s(bool,X282) | p(s(bool,X283)) | p(s(bool,X280)) | s(bool,X282) = s(bool,X283)) ) | ($spl62 | $spl138)),
  inference(subsumption_resolution,[],[f8931,f246])).
fof(f8931,plain,(
  ( ! [X280,X282,X281,X283] : (s(bool,X281) != s(bool,X282) | p(s(bool,X283)) | p(s(bool,X280)) | s(bool,X282) = s(bool,X283) | s(bool,'T_0') = s(bool,X280)) ) | ($spl62 | $spl138)),
  inference(superposition,[],[f2964,f7737])).
fof(f9492,plain,(
  ( ! [X271,X270] : (s(bool,'T_0') != s(bool,X271) | s(bool,'T_0') = s(bool,X270)) ) | $spl138),
  inference(subsumption_resolution,[],[f9491,f7737])).
fof(f9491,plain,(
  ( ! [X271,X270,X272] : (s(bool,'T_0') != s(bool,X271) | s(bool,X270) != s(bool,X272) | s(bool,'T_0') = s(bool,X270)) ) | $spl138),
  inference(subsumption_resolution,[],[f8928,f2882])).
fof(f8928,plain,(
  ( ! [X271,X270,X272] : (s(bool,'T_0') != s(bool,X271) | s(bool,X270) != s(bool,X272) | s(bool,'T_0') = s(bool,X272) | s(bool,'T_0') = s(bool,X270)) ) | $spl138),
  inference(superposition,[],[f2882,f7737])).
fof(f9490,plain,(
  ( ! [X269,X267,X268] : (s(bool,'T_0') != s(bool,X268) | s(bool,X267) != s(bool,X269) | p(s(bool,X269))) ) | $spl138),
  inference(subsumption_resolution,[],[f8927,f2881])).
fof(f8927,plain,(
  ( ! [X269,X267,X268] : (s(bool,'T_0') != s(bool,X268) | s(bool,X267) != s(bool,X269) | p(s(bool,X269)) | s(bool,'T_0') = s(bool,X267)) ) | $spl138),
  inference(superposition,[],[f2881,f7737])).
fof(f9332,plain,(
  ( ! [X249,X248,X247] : (s(bool,X248) != s(bool,X249) | s(bool,'T_0') = s(bool,X247) | s(bool,'F_0') = s(bool,X249)) ) | ($spl6 | $spl138)),
  inference(duplicate_literal_removal,[],[f8920])).
fof(f8920,plain,(
  ( ! [X249,X248,X247] : (s(bool,X248) != s(bool,X249) | s(bool,'T_0') = s(bool,X247) | s(bool,'F_0') = s(bool,X249) | s(bool,'T_0') = s(bool,X247)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f828,f7737])).
fof(f9333,plain,(
  ( ! [X239,X240,X238] : (s(bool,X239) != s(bool,X240) | ~p(s(bool,X240)) | s(bool,'T_0') = s(bool,X238)) ) | ($spl6 | $spl138)),
  inference(duplicate_literal_removal,[],[f8917])).
fof(f8917,plain,(
  ( ! [X239,X240,X238] : (s(bool,X239) != s(bool,X240) | ~p(s(bool,X240)) | s(bool,'T_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X238)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f766,f7737])).
fof(f9488,plain,(
  ( ! [X233,X234,X232] : (s(bool,X233) != s(bool,X234) | p(s(bool,X232)) | ~p(s(bool,X234))) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f8915,f246])).
fof(f8915,plain,(
  ( ! [X233,X234,X232] : (s(bool,X233) != s(bool,X234) | p(s(bool,X232)) | ~p(s(bool,X234)) | s(bool,'T_0') = s(bool,X232)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f297,f7737])).
fof(f9487,plain,(
  ( ! [X231,X229,X230,X228] : (~p(s(bool,X229)) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X228) != s(bool,X231)) ) | $spl138),
  inference(subsumption_resolution,[],[f8914,f2881])).
fof(f8914,plain,(
  ( ! [X231,X229,X230,X228] : (~p(s(bool,X229)) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X228) != s(bool,X231) | s(bool,'T_0') = s(bool,X228)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9486,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,X217) != s(bool,X218) | p(s(bool,X219)) | ~p(s(bool,X218)) | p(s(bool,X216))) ) | $spl138),
  inference(subsumption_resolution,[],[f8911,f246])).
fof(f8911,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,X217) != s(bool,X218) | p(s(bool,X219)) | ~p(s(bool,X218)) | p(s(bool,X216)) | s(bool,'T_0') = s(bool,X216)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9334,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,X211) | s(bool,'T_0') = s(bool,X210)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8909])).
fof(f8909,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X210)) ) | $spl138),
  inference(superposition,[],[f261,f7737])).
fof(f9335,plain,(
  ( ! [X206,X207] : (~p(s(bool,X207)) | s(bool,'T_0') = s(bool,X206)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8907])).
fof(f8907,plain,(
  ( ! [X206,X207] : (~p(s(bool,X207)) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) ) | $spl138),
  inference(superposition,[],[f249,f7737])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f9485,plain,(
  ( ! [X200,X201] : (s(bool,'T_0') != s(bool,X201) | p(s(bool,X200))) ) | $spl138),
  inference(subsumption_resolution,[],[f8904,f246])).
fof(f8904,plain,(
  ( ! [X200,X201] : (s(bool,'T_0') != s(bool,X201) | p(s(bool,X200)) | s(bool,'T_0') = s(bool,X200)) ) | $spl138),
  inference(superposition,[],[f248,f7737])).
fof(f248,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f9336,plain,(
  ( ! [X198,X199] : (~p(s(bool,X199)) | s(bool,'T_0') = s(bool,X198)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8903])).
fof(f8903,plain,(
  ( ! [X198,X199] : (~p(s(bool,X199)) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,X198)) ) | $spl138),
  inference(superposition,[],[f247,f7737])).
fof(f9484,plain,(
  ( ! [X192,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X192))) ) | $spl138),
  inference(subsumption_resolution,[],[f8900,f246])).
fof(f8900,plain,(
  ( ! [X192,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X192)) | s(bool,'T_0') = s(bool,X192)) ) | $spl138),
  inference(superposition,[],[f246,f7737])).
fof(f9483,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') = s(bool,X189) | p(s(bool,X188))) ) | $spl138),
  inference(subsumption_resolution,[],[f8898,f246])).
fof(f8898,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') = s(bool,X189) | p(s(bool,X188)) | s(bool,'T_0') = s(bool,X188)) ) | $spl138),
  inference(superposition,[],[f243,f7737])).
fof(f9482,plain,(
  ( ! [X187,X186] : (~p(s(bool,X187)) | s(bool,'T_0') = s(bool,X186)) ) | $spl138),
  inference(subsumption_resolution,[],[f8897,f7737])).
fof(f8897,plain,(
  ( ! [X187,X186] : (~p(s(bool,X187)) | s(bool,'F_0') != s(bool,X186) | s(bool,'T_0') = s(bool,X186)) ) | $spl138),
  inference(superposition,[],[f242,f7737])).
fof(f9481,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | p(s(bool,X167)) | ~p(s(bool,X168)) | s(bool,X165) != s(bool,X167) | s(bool,X168) != s(bool,X169)) ) | $spl138),
  inference(subsumption_resolution,[],[f8888,f2881])).
fof(f8888,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | p(s(bool,X167)) | ~p(s(bool,X168)) | s(bool,X165) != s(bool,X167) | s(bool,X168) != s(bool,X169) | s(bool,'T_0') = s(bool,X165)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9479,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,X147) | p(s(bool,X145)) | ~p(s(bool,X148)) | ~p(s(bool,X147)) | s(bool,X148) != s(bool,X149)) ) | $spl138),
  inference(subsumption_resolution,[],[f8884,f246])).
fof(f8884,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,X147) | p(s(bool,X145)) | ~p(s(bool,X148)) | ~p(s(bool,X147)) | s(bool,X148) != s(bool,X149) | s(bool,'T_0') = s(bool,X145)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9476,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,X126)) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X125) != s(bool,X129)) ) | $spl138),
  inference(subsumption_resolution,[],[f8880,f2881])).
fof(f8880,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,X126)) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X125) != s(bool,X129) | s(bool,'T_0') = s(bool,X125)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9475,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,X111) != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | ~p(s(bool,X114)) | p(s(bool,X110))) ) | $spl138),
  inference(subsumption_resolution,[],[f8877,f246])).
fof(f8877,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,X111) != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | ~p(s(bool,X114)) | p(s(bool,X110)) | s(bool,'T_0') = s(bool,X110)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9474,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X100) != s(bool,X103)) ) | $spl138),
  inference(subsumption_resolution,[],[f8875,f2881])).
fof(f8875,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X100) != s(bool,X103) | s(bool,'T_0') = s(bool,X100)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9471,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X78) != s(bool,X79)) ) | $spl138),
  inference(subsumption_resolution,[],[f8870,f246])).
fof(f8870,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X78) != s(bool,X79) | s(bool,'T_0') = s(bool,X75)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9470,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X65) != s(bool,X67) | s(bool,X68) != s(bool,X69)) ) | $spl138),
  inference(subsumption_resolution,[],[f8868,f2881])).
fof(f8868,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X65) != s(bool,X67) | s(bool,X68) != s(bool,X69) | s(bool,'T_0') = s(bool,X65)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9468,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X51) != s(bool,X52) | p(s(bool,X50)) | ~p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X53) != s(bool,X54)) ) | $spl138),
  inference(subsumption_resolution,[],[f8865,f246])).
fof(f8865,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X51) != s(bool,X52) | p(s(bool,X50)) | ~p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X53) != s(bool,X54) | s(bool,'T_0') = s(bool,X50)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9465,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X30) != s(bool,X33)) ) | $spl138),
  inference(subsumption_resolution,[],[f8861,f2881])).
fof(f8861,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9462,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4)) ) | $spl138),
  inference(subsumption_resolution,[],[f8855,f246])).
fof(f8855,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X0)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9458,plain,(
  ( ! [X546,X545,X547] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X545),X546),s('type.closSem.state'(X545),X547))))) ) | $spl288),
  inference(cnf_transformation,[],[f9458_D])).
fof(f9458_D,plain,(
  ( ! [X546,X545,X547] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X545),X546),s('type.closSem.state'(X545),X547))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f9461,plain,(
  ( ! [X565,X563,X564,X562] : (p(s(bool,X565)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X562),X563),s('type.closSem.state'(X562),X564))))) ) | $spl138),
  inference(subsumption_resolution,[],[f8851,f246])).
fof(f8851,plain,(
  ( ! [X565,X563,X564,X562] : (p(s(bool,X565)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X562),X563),s('type.closSem.state'(X562),X564)))) | s(bool,'T_0') = s(bool,X565)) ) | $spl138),
  inference(superposition,[],[f256,f7737])).
fof(f9460,plain,(
  ( ! [X557,X555,X556,X554] : (p(s(bool,X557)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X554),X555),s('type.closSem.state'(X554),X556))))) ) | $spl138),
  inference(subsumption_resolution,[],[f8849,f246])).
fof(f8849,plain,(
  ( ! [X557,X555,X556,X554] : (p(s(bool,X557)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X554),X555),s('type.closSem.state'(X554),X556)))) | s(bool,'T_0') = s(bool,X557)) ) | $spl138),
  inference(superposition,[],[f255,f7737])).
fof(f9457,plain,(
  ( ! [X546,X545,X547,X548] : (p(s(bool,X548)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X545),X546),s('type.closSem.state'(X545),X547))))) ) | $spl138),
  inference(subsumption_resolution,[],[f8846,f246])).
fof(f8846,plain,(
  ( ! [X546,X545,X547,X548] : (p(s(bool,X548)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X545),X546),s('type.closSem.state'(X545),X547)))) | s(bool,'T_0') = s(bool,X548)) ) | $spl138),
  inference(superposition,[],[f257,f7737])).
fof(f9456,plain,(
  ( ! [X511,X512,X513] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X513)))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X512))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X513))) | s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X512))) != s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X513))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X512))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X513))))))) ) | $spl286),
  inference(cnf_transformation,[],[f9456_D])).
fof(f9456_D,plain,(
  ( ! [X511,X512,X513] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X513)))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X512))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X513))) | s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X512))) != s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X513))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X512))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X513))))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f9455,plain,(
  ( ! [X511,X512,X514,X513] : (p(s(bool,X514)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X512))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X513)))))) | s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X512))) != s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X513))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X512))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X513)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8838,f246])).
fof(f8838,plain,(
  ( ! [X511,X512,X514,X513] : (p(s(bool,X514)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X512))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X511),X513)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X512))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X511),X513)))))) | s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X512))) != s('type.ffi.ffi_state'(X511),'const.closSem.state_ffi_1'(s('type.closSem.state'(X511),X513))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X512))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X511),X513))) | s(bool,'T_0') = s(bool,X514)) ) | $spl138),
  inference(superposition,[],[f258,f7737])).
fof(f9452,plain,(
  ( ! [X472,X474,X473] : (s(bool,X473) != s(bool,X474) | s(bool,'T_0') != s(bool,X474) | s(bool,'T_0') = s(bool,X472)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8823,f7737])).
fof(f8823,plain,(
  ( ! [X472,X474,X473] : (s(bool,X472) != s(bool,X473) | s(bool,X473) != s(bool,X474) | s(bool,'T_0') != s(bool,X474) | s(bool,'T_0') = s(bool,X472)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f5870,f7737])).
fof(f9451,plain,(
  ( ! [X425,X427] : (s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,X425)) ) | $spl138),
  inference(subsumption_resolution,[],[f9450,f7737])).
fof(f9450,plain,(
  ( ! [X426,X425,X427] : (s(bool,X425) != s(bool,X426) | s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,X425)) ) | $spl138),
  inference(subsumption_resolution,[],[f8808,f4329])).
fof(f8808,plain,(
  ( ! [X426,X425,X427] : (s(bool,X425) != s(bool,X426) | s(bool,X426) != s(bool,X427) | s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,X425)) ) | $spl138),
  inference(superposition,[],[f2882,f7737])).
fof(f9440,plain,(
  ( ! [X397,X398] : (p(s(bool,X397)) | ~p(s(bool,X398))) ) | $spl138),
  inference(subsumption_resolution,[],[f8798,f246])).
fof(f8798,plain,(
  ( ! [X397,X398] : (p(s(bool,X397)) | ~p(s(bool,X398)) | s(bool,'T_0') = s(bool,X397)) ) | $spl138),
  inference(superposition,[],[f240,f7737])).
fof(f9439,plain,(
  ( ! [X395,X396] : (p(s(bool,X395)) | ~p(s(bool,X396))) ) | $spl138),
  inference(subsumption_resolution,[],[f8797,f246])).
fof(f8797,plain,(
  ( ! [X395,X396] : (p(s(bool,X395)) | ~p(s(bool,X396)) | s(bool,'T_0') = s(bool,X395)) ) | $spl138),
  inference(superposition,[],[f238,f7737])).
fof(f9438,plain,(
  ( ! [X393,X394] : (p(s(bool,X393)) | ~p(s(bool,X394))) ) | $spl138),
  inference(subsumption_resolution,[],[f8796,f246])).
fof(f8796,plain,(
  ( ! [X393,X394] : (p(s(bool,X393)) | ~p(s(bool,X394)) | s(bool,'T_0') = s(bool,X393)) ) | $spl138),
  inference(superposition,[],[f225,f7737])).
fof(f9437,plain,(
  ( ! [X391,X392] : (p(s(bool,X391)) | p(s(bool,X392))) ) | $spl138),
  inference(subsumption_resolution,[],[f8795,f246])).
fof(f8795,plain,(
  ( ! [X391,X392] : (p(s(bool,X391)) | p(s(bool,X392)) | s(bool,'T_0') = s(bool,X391)) ) | $spl138),
  inference(superposition,[],[f224,f7737])).
fof(f9436,plain,(
  ( ! [X389,X390] : (p(s(bool,X389)) | p(s(bool,X390))) ) | $spl138),
  inference(subsumption_resolution,[],[f8794,f246])).
fof(f8794,plain,(
  ( ! [X389,X390] : (p(s(bool,X389)) | p(s(bool,X390)) | s(bool,'T_0') = s(bool,X389)) ) | $spl138),
  inference(superposition,[],[f221,f7737])).
fof(f9435,plain,(
  ( ! [X387,X388] : (p(s(bool,X387)) | p(s(bool,X388))) ) | $spl138),
  inference(subsumption_resolution,[],[f8793,f246])).
fof(f8793,plain,(
  ( ! [X387,X388] : (p(s(bool,X387)) | p(s(bool,X388)) | s(bool,'T_0') = s(bool,X387)) ) | $spl138),
  inference(superposition,[],[f218,f7737])).
fof(f9434,plain,(
  ( ! [X385,X386] : (p(s(bool,X385)) | ~p(s(bool,X386))) ) | $spl138),
  inference(subsumption_resolution,[],[f8792,f246])).
fof(f8792,plain,(
  ( ! [X385,X386] : (p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,'T_0') = s(bool,X385)) ) | $spl138),
  inference(superposition,[],[f216,f7737])).
fof(f9433,plain,(
  ( ! [X383,X382,X384] : (s(bool,'F_0') != s(bool,X383) | s(bool,X382) != s(bool,X384) | s(bool,'T_0') != s(bool,X384)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8791,f2590])).
fof(f8791,plain,(
  ( ! [X383,X382,X384] : (s(bool,'F_0') != s(bool,X383) | s(bool,X382) != s(bool,X384) | s(bool,'T_0') != s(bool,X384) | s(bool,'T_0') = s(bool,X383)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f5870,f7737])).
fof(f9432,plain,(
  ( ! [X380,X378,X381] : (s(bool,X378) != s(bool,X380) | s(bool,'T_0') != s(bool,X381) | s(bool,X380) != s(bool,X381)) ) | $spl284),
  inference(cnf_transformation,[],[f9432_D])).
fof(f9432_D,plain,(
  ( ! [X380,X378,X381] : (s(bool,X378) != s(bool,X380) | s(bool,'T_0') != s(bool,X381) | s(bool,X380) != s(bool,X381)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f9431,plain,(
  ( ! [X379,X380,X378,X381] : (p(s(bool,X379)) | s(bool,X378) != s(bool,X380) | s(bool,X380) != s(bool,X381) | s(bool,'T_0') != s(bool,X381)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8790,f246])).
fof(f8790,plain,(
  ( ! [X379,X380,X378,X381] : (p(s(bool,X379)) | s(bool,X378) != s(bool,X380) | s(bool,X380) != s(bool,X381) | s(bool,'T_0') != s(bool,X381) | s(bool,'T_0') = s(bool,X379)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f4702,f7737])).
fof(f9430,plain,(
  ( ! [X377,X374,X375] : (s(bool,'F_0') = s(bool,X374) | s(bool,'T_0') = s(bool,X377) | s(bool,'T_0') = s(bool,X375)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f9429,f7737])).
fof(f9429,plain,(
  ( ! [X377,X374,X376,X375] : (s(bool,X375) != s(bool,X376) | s(bool,'F_0') = s(bool,X374) | s(bool,'T_0') = s(bool,X377) | s(bool,'T_0') = s(bool,X375)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f8789,f4329])).
fof(f8789,plain,(
  ( ! [X377,X374,X376,X375] : (s(bool,X375) != s(bool,X376) | s(bool,'F_0') = s(bool,X374) | s(bool,X376) != s(bool,X377) | s(bool,'T_0') = s(bool,X377) | s(bool,'T_0') = s(bool,X375)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f4623,f7737])).
fof(f9427,plain,(
  ( ! [X358,X360,X359,X361] : (p(s(bool,X359)) | s(bool,X358) = s(bool,X360) | s(bool,'T_0') = s(bool,X361)) ) | $spl138),
  inference(subsumption_resolution,[],[f9426,f7737])).
fof(f9426,plain,(
  ( ! [X358,X360,X359,X361] : (p(s(bool,X359)) | s(bool,X358) = s(bool,X360) | s(bool,X360) != s(bool,X361) | s(bool,'T_0') = s(bool,X361)) ) | $spl138),
  inference(subsumption_resolution,[],[f8785,f246])).
fof(f8785,plain,(
  ( ! [X358,X360,X359,X361] : (p(s(bool,X359)) | s(bool,X358) = s(bool,X360) | s(bool,X360) != s(bool,X361) | s(bool,'T_0') = s(bool,X361) | s(bool,'T_0') = s(bool,X359)) ) | $spl138),
  inference(superposition,[],[f4328,f7737])).
fof(f9338,plain,(
  ( ! [X333,X331,X332,X330] : (s(bool,'T_0') = s(bool,X331) | s(bool,X332) = s(bool,X333) | s(bool,X330) = s(bool,X333) | s(bool,'F_0') = s(bool,X332)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8778])).
fof(f8778,plain,(
  ( ! [X333,X331,X332,X330] : (s(bool,'T_0') = s(bool,X331) | s(bool,X332) = s(bool,X333) | s(bool,X330) = s(bool,X333) | s(bool,'F_0') = s(bool,X332) | s(bool,'T_0') = s(bool,X331)) ) | $spl138),
  inference(superposition,[],[f3476,f7737])).
fof(f9424,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'T_0') = s(bool,X324)) ) | ($spl70 | $spl138)),
  inference(subsumption_resolution,[],[f8776,f7737])).
fof(f8776,plain,(
  ( ! [X325,X327,X323,X324,X326] : (s(bool,X324) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X325),X326),s('type.closSem.state'(X325),X327))) | s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X326))) = s('type.ffi.ffi_state'(X325),'const.closSem.state_ffi_1'(s('type.closSem.state'(X325),X327))) | ~p(s(bool,X323)) | s(bool,'T_0') = s(bool,X324)) ) | ($spl70 | $spl138)),
  inference(superposition,[],[f3186,f7737])).
fof(f9423,plain,(
  ( ! [X321,X322,X319] : (s(bool,X319) != s(bool,X321) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) != s(bool,X322)) ) | $spl282),
  inference(cnf_transformation,[],[f9423_D])).
fof(f9423_D,plain,(
  ( ! [X321,X322,X319] : (s(bool,X319) != s(bool,X321) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) != s(bool,X322)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f9422,plain,(
  ( ! [X321,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) != s(bool,X321) | s(bool,X321) != s(bool,X322) | s(bool,'F_0') = s(bool,X322)) ) | ($spl68 | $spl138)),
  inference(subsumption_resolution,[],[f8775,f246])).
fof(f8775,plain,(
  ( ! [X321,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) != s(bool,X321) | s(bool,X321) != s(bool,X322) | s(bool,'F_0') = s(bool,X322) | s(bool,'T_0') = s(bool,X320)) ) | ($spl68 | $spl138)),
  inference(superposition,[],[f3130,f7737])).
fof(f9421,plain,(
  ( ! [X316,X318,X317] : (s(bool,'F_0') = s(bool,X318) | s(bool,X316) != s(bool,X318) | s(bool,'T_0') = s(bool,X317)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f8774,f261])).
fof(f8774,plain,(
  ( ! [X316,X318,X317] : (s(bool,'F_0') != s(bool,X317) | s(bool,'F_0') = s(bool,X318) | s(bool,X316) != s(bool,X318) | s(bool,'T_0') = s(bool,X317)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f3111,f7737])).
fof(f9419,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'T_0') = s(bool,X285)) ) | ($spl64 | $spl138)),
  inference(subsumption_resolution,[],[f8765,f7737])).
fof(f8765,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X285) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X286),X287),s('type.closSem.state'(X286),X288))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X287))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X286),X288))) | ~p(s(bool,X284)) | s(bool,'T_0') = s(bool,X285)) ) | ($spl64 | $spl138)),
  inference(superposition,[],[f3029,f7737])).
fof(f9418,plain,(
  ( ! [X277,X279,X276,X278] : (s(bool,X277) != s(bool,X278) | p(s(bool,X279)) | p(s(bool,X278)) | s(bool,X276) = s(bool,X279)) ) | ($spl62 | $spl138)),
  inference(subsumption_resolution,[],[f8763,f2881])).
fof(f8763,plain,(
  ( ! [X277,X279,X276,X278] : (s(bool,X277) != s(bool,X278) | p(s(bool,X279)) | p(s(bool,X278)) | s(bool,X276) = s(bool,X279) | s(bool,'T_0') = s(bool,X277)) ) | ($spl62 | $spl138)),
  inference(superposition,[],[f2964,f7737])).
fof(f9417,plain,(
  ( ! [X273,X275,X274] : (s(bool,X273) = s(bool,X275) | p(s(bool,X275)) | s(bool,'T_0') = s(bool,X274)) ) | $spl138),
  inference(subsumption_resolution,[],[f8762,f7737])).
fof(f8762,plain,(
  ( ! [X273,X275,X274] : (s(bool,'F_0') != s(bool,X274) | s(bool,X273) = s(bool,X275) | p(s(bool,X275)) | s(bool,'T_0') = s(bool,X274)) ) | $spl138),
  inference(superposition,[],[f2915,f7737])).
fof(f9416,plain,(
  ( ! [X265,X264,X266] : (p(s(bool,X265)) | s(bool,X264) != s(bool,X266) | s(bool,'F_0') = s(bool,X266)) ) | ($spl60 | $spl138)),
  inference(subsumption_resolution,[],[f8759,f246])).
fof(f8759,plain,(
  ( ! [X265,X264,X266] : (p(s(bool,X265)) | s(bool,X264) != s(bool,X266) | s(bool,'F_0') = s(bool,X266) | s(bool,'T_0') = s(bool,X265)) ) | ($spl60 | $spl138)),
  inference(superposition,[],[f2856,f7737])).
fof(f9414,plain,(
  ( ! [X259,X260,X258] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X258) = s(bool,X260)) ) | $spl138),
  inference(subsumption_resolution,[],[f8757,f246])).
fof(f8757,plain,(
  ( ! [X259,X260,X258] : (p(s(bool,X259)) | p(s(bool,X260)) | s(bool,X258) = s(bool,X260) | s(bool,'T_0') = s(bool,X259)) ) | $spl138),
  inference(superposition,[],[f2659,f7737])).
fof(f9413,plain,(
  ( ! [X255,X257,X256] : (p(s(bool,X256)) | p(s(bool,X257)) | s(bool,X255) = s(bool,X257)) ) | $spl138),
  inference(subsumption_resolution,[],[f8756,f246])).
fof(f8756,plain,(
  ( ! [X255,X257,X256] : (p(s(bool,X256)) | p(s(bool,X257)) | s(bool,X255) = s(bool,X257) | s(bool,'T_0') = s(bool,X256)) ) | $spl138),
  inference(superposition,[],[f2659,f7737])).
fof(f9412,plain,(
  ( ! [X253,X254,X252] : (s(bool,X252) != s(bool,X254) | ~p(s(bool,X254)) | s(bool,'T_0') = s(bool,X253)) ) | ($spl58 | $spl138)),
  inference(subsumption_resolution,[],[f8755,f7737])).
fof(f8755,plain,(
  ( ! [X253,X254,X252] : (s(bool,'F_0') != s(bool,X253) | s(bool,X252) != s(bool,X254) | ~p(s(bool,X254)) | s(bool,'T_0') = s(bool,X253)) ) | ($spl58 | $spl138)),
  inference(superposition,[],[f2634,f7737])).
fof(f9411,plain,(
  ( ! [X251,X250] : (s(bool,'F_0') != s(bool,X251) | s(bool,'T_0') != s(bool,X250)) ) | $spl138),
  inference(subsumption_resolution,[],[f8754,f2590])).
fof(f8754,plain,(
  ( ! [X251,X250] : (s(bool,'F_0') != s(bool,X251) | s(bool,'T_0') != s(bool,X250) | s(bool,'T_0') = s(bool,X251)) ) | $spl138),
  inference(superposition,[],[f2590,f7737])).
fof(f9410,plain,(
  ( ! [X244,X245] : (s(bool,'F_0') = s(bool,X244) | s(bool,'T_0') = s(bool,X245)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9409,f7737])).
fof(f9409,plain,(
  ( ! [X246,X244,X245] : (s(bool,X245) != s(bool,X246) | s(bool,'F_0') = s(bool,X244) | s(bool,'T_0') = s(bool,X245)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f8752,f2882])).
fof(f8752,plain,(
  ( ! [X246,X244,X245] : (s(bool,X245) != s(bool,X246) | s(bool,'T_0') = s(bool,X246) | s(bool,'F_0') = s(bool,X244) | s(bool,'T_0') = s(bool,X245)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f828,f7737])).
fof(f9408,plain,(
  ( ! [X241,X242] : (~p(s(bool,X241)) | s(bool,'T_0') = s(bool,X242)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f9407,f7737])).
fof(f9407,plain,(
  ( ! [X243,X241,X242] : (s(bool,X242) != s(bool,X243) | ~p(s(bool,X241)) | s(bool,'T_0') = s(bool,X242)) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f8751,f2882])).
fof(f8751,plain,(
  ( ! [X243,X241,X242] : (s(bool,X242) != s(bool,X243) | ~p(s(bool,X241)) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X242)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f766,f7737])).
fof(f9406,plain,(
  ( ! [X237,X235,X236] : (s(bool,X236) != s(bool,X237) | p(s(bool,X237)) | ~p(s(bool,X235))) ) | ($spl6 | $spl138)),
  inference(subsumption_resolution,[],[f8749,f2881])).
fof(f8749,plain,(
  ( ! [X237,X235,X236] : (s(bool,X236) != s(bool,X237) | p(s(bool,X237)) | ~p(s(bool,X235)) | s(bool,'T_0') = s(bool,X236)) ) | ($spl6 | $spl138)),
  inference(superposition,[],[f297,f7737])).
fof(f9404,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X220) != s(bool,X223)) ) | $spl138),
  inference(subsumption_resolution,[],[f8745,f246])).
fof(f8745,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X220) != s(bool,X223) | s(bool,'T_0') = s(bool,X221)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9403,plain,(
  ( ! [X212,X215,X213,X214] : (s(bool,X213) != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X212)) | p(s(bool,X214))) ) | $spl138),
  inference(subsumption_resolution,[],[f8743,f2881])).
fof(f8743,plain,(
  ( ! [X212,X215,X213,X214] : (s(bool,X213) != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X212)) | p(s(bool,X214)) | s(bool,'T_0') = s(bool,X213)) ) | $spl138),
  inference(superposition,[],[f262,f7737])).
fof(f9341,plain,(
  ( ! [X208,X209] : (s(bool,'T_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X208)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8741])).
fof(f8741,plain,(
  ( ! [X208,X209] : (s(bool,'T_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X208) | s(bool,'T_0') = s(bool,X209)) ) | $spl138),
  inference(superposition,[],[f261,f7737])).
fof(f9342,plain,(
  ( ! [X204,X205] : (s(bool,'T_0') = s(bool,X205) | ~p(s(bool,X204))) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8739])).
fof(f8739,plain,(
  ( ! [X204,X205] : (s(bool,'T_0') = s(bool,X205) | ~p(s(bool,X204)) | s(bool,'T_0') = s(bool,X205)) ) | $spl138),
  inference(superposition,[],[f249,f7737])).
fof(f9402,plain,(
  ( ! [X202,X203] : (p(s(bool,X203)) | s(bool,'T_0') != s(bool,X202)) ) | $spl138),
  inference(subsumption_resolution,[],[f8738,f246])).
fof(f8738,plain,(
  ( ! [X202,X203] : (p(s(bool,X203)) | s(bool,'T_0') != s(bool,X202) | s(bool,'T_0') = s(bool,X203)) ) | $spl138),
  inference(superposition,[],[f248,f7737])).
fof(f9343,plain,(
  ( ! [X196,X197] : (s(bool,'T_0') = s(bool,X197) | ~p(s(bool,X196))) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8735])).
fof(f8735,plain,(
  ( ! [X196,X197] : (s(bool,'T_0') = s(bool,X197) | ~p(s(bool,X196)) | s(bool,'T_0') = s(bool,X197)) ) | $spl138),
  inference(superposition,[],[f247,f7737])).
fof(f9401,plain,(
  ( ! [X194,X195] : (p(s(bool,X195)) | s(bool,'T_0') != s(bool,X194)) ) | $spl138),
  inference(subsumption_resolution,[],[f8734,f246])).
fof(f8734,plain,(
  ( ! [X194,X195] : (p(s(bool,X195)) | s(bool,'T_0') != s(bool,X194) | s(bool,'T_0') = s(bool,X195)) ) | $spl138),
  inference(superposition,[],[f246,f7737])).
fof(f9400,plain,(
  ( ! [X191,X190] : (p(s(bool,X191)) | s(bool,'F_0') = s(bool,X190)) ) | $spl138),
  inference(subsumption_resolution,[],[f8732,f246])).
fof(f8732,plain,(
  ( ! [X191,X190] : (p(s(bool,X191)) | s(bool,'F_0') = s(bool,X190) | s(bool,'T_0') = s(bool,X191)) ) | $spl138),
  inference(superposition,[],[f243,f7737])).
fof(f9399,plain,(
  ( ! [X185,X184] : (~p(s(bool,X184)) | s(bool,'T_0') = s(bool,X185)) ) | $spl138),
  inference(subsumption_resolution,[],[f8729,f7737])).
fof(f8729,plain,(
  ( ! [X185,X184] : (s(bool,'F_0') != s(bool,X185) | ~p(s(bool,X184)) | s(bool,'T_0') = s(bool,X185)) ) | $spl138),
  inference(superposition,[],[f242,f7737])).
fof(f9398,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X156)) | ~p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X155) != s(bool,X158) | s(bool,X157) != s(bool,X159)) ) | $spl138),
  inference(subsumption_resolution,[],[f8719,f246])).
fof(f8719,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X156)) | ~p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X155) != s(bool,X158) | s(bool,X157) != s(bool,X159) | s(bool,'T_0') = s(bool,X156)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9397,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X151) != s(bool,X152) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,X150)) | s(bool,X153) != s(bool,X154)) ) | $spl138),
  inference(subsumption_resolution,[],[f8718,f2881])).
fof(f8718,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X151) != s(bool,X152) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,X150)) | s(bool,X153) != s(bool,X154) | s(bool,'T_0') = s(bool,X151)) ) | $spl138),
  inference(superposition,[],[f203,f7737])).
fof(f9394,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X115) != s(bool,X118)) ) | $spl138),
  inference(subsumption_resolution,[],[f8711,f246])).
fof(f8711,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X115) != s(bool,X118) | s(bool,'T_0') = s(bool,X116)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9393,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X106) != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X105)) | ~p(s(bool,X109)) | p(s(bool,X107))) ) | $spl138),
  inference(subsumption_resolution,[],[f8709,f2881])).
fof(f8709,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X106) != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X105)) | ~p(s(bool,X109)) | p(s(bool,X107)) | s(bool,'T_0') = s(bool,X106)) ) | $spl138),
  inference(superposition,[],[f202,f7737])).
fof(f9391,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,X90) != s(bool,X93)) ) | $spl138),
  inference(subsumption_resolution,[],[f8706,f246])).
fof(f8706,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,X90) != s(bool,X93) | s(bool,'T_0') = s(bool,X91)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9389,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) != s(bool,X72) | p(s(bool,X73)) | ~p(s(bool,X70)) | p(s(bool,X72)) | s(bool,X73) != s(bool,X74)) ) | $spl138),
  inference(subsumption_resolution,[],[f8702,f2881])).
fof(f8702,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) != s(bool,X72) | p(s(bool,X73)) | ~p(s(bool,X70)) | p(s(bool,X72)) | s(bool,X73) != s(bool,X74) | s(bool,'T_0') = s(bool,X71)) ) | $spl138),
  inference(superposition,[],[f201,f7737])).
fof(f9388,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X55) != s(bool,X58) | s(bool,X57) != s(bool,X59)) ) | $spl138),
  inference(subsumption_resolution,[],[f8699,f246])).
fof(f8699,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X55) != s(bool,X58) | s(bool,X57) != s(bool,X59) | s(bool,'T_0') = s(bool,X56)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9387,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X45)) | s(bool,X48) != s(bool,X49)) ) | $spl138),
  inference(subsumption_resolution,[],[f8697,f2881])).
fof(f8697,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X45)) | s(bool,X48) != s(bool,X49) | s(bool,'T_0') = s(bool,X46)) ) | $spl138),
  inference(superposition,[],[f199,f7737])).
fof(f9384,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl138),
  inference(subsumption_resolution,[],[f8692,f246])).
fof(f8692,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23) | s(bool,'T_0') = s(bool,X21)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9382,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | ~p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X8) != s(bool,X9)) ) | $spl138),
  inference(subsumption_resolution,[],[f8689,f2881])).
fof(f8689,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | ~p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X6)) ) | $spl138),
  inference(superposition,[],[f197,f7737])).
fof(f9344,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8684])).
fof(f8684,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) ) | $spl138),
  inference(superposition,[],[f7737,f261])).
fof(f9379,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,X55) | p(s(bool,X54))) ) | $spl138),
  inference(subsumption_resolution,[],[f8683,f246])).
fof(f8683,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55) | p(s(bool,X54))) ) | $spl138),
  inference(superposition,[],[f7737,f243])).
fof(f9345,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8680])).
fof(f8680,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl138),
  inference(superposition,[],[f7737,f7737])).
fof(f9346,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,X28) = s(bool,X31) | s(bool,'F_0') = s(bool,X31)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8678])).
fof(f8678,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X31)) ) | $spl138),
  inference(superposition,[],[f7737,f3476])).
fof(f9349,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X48)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8664])).
fof(f8664,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X48)) ) | $spl138),
  inference(superposition,[],[f7737,f3476])).
fof(f9375,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'F_0') = s(bool,X85) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X83))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8647,f2981])).
fof(f8647,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X82),X83),s('type.closSem.state'(X82),X84))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X83))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X82),X84)))) ) | $spl138),
  inference(superposition,[],[f7737,f2980])).
fof(f9374,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'F_0') = s(bool,X81) | s('type.ffi.ffi_state'(X78),'const.closSem.state_ffi_1'(s('type.closSem.state'(X78),X79))) = s('type.ffi.ffi_state'(X78),'const.closSem.state_ffi_1'(s('type.closSem.state'(X78),X80)))) ) | $spl138),
  inference(subsumption_resolution,[],[f8646,f3144])).
fof(f8646,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X78),X79),s('type.closSem.state'(X78),X80))) | s('type.ffi.ffi_state'(X78),'const.closSem.state_ffi_1'(s('type.closSem.state'(X78),X79))) = s('type.ffi.ffi_state'(X78),'const.closSem.state_ffi_1'(s('type.closSem.state'(X78),X80)))) ) | $spl138),
  inference(superposition,[],[f7737,f3143])).
fof(f9352,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X52)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8634])).
fof(f8634,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl138),
  inference(superposition,[],[f7737,f7737])).
fof(f9353,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8631])).
fof(f8631,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X45)) ) | $spl138),
  inference(superposition,[],[f7737,f3476])).
fof(f9356,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8621])).
fof(f8621,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl138),
  inference(superposition,[],[f7737,f261])).
fof(f9371,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X0))) ) | $spl138),
  inference(subsumption_resolution,[],[f8617,f246])).
fof(f8617,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | $spl138),
  inference(superposition,[],[f7737,f243])).
fof(f9358,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X48)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8600])).
fof(f8600,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X48)) ) | $spl138),
  inference(superposition,[],[f7737,f7737])).
fof(f9359,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,X22) = s(bool,X25) | s(bool,'F_0') = s(bool,X24)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8594])).
fof(f8594,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,X22) = s(bool,X25) | s(bool,'F_0') = s(bool,X24)) ) | $spl138),
  inference(superposition,[],[f7737,f3476])).
fof(f9362,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8588])).
fof(f8588,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl138),
  inference(superposition,[],[f7737,f261])).
fof(f9363,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8587])).
fof(f8587,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) ) | $spl138),
  inference(superposition,[],[f7737,f249])).
fof(f9364,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl138),
  inference(duplicate_literal_removal,[],[f8586])).
fof(f8586,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl138),
  inference(superposition,[],[f7737,f247])).
fof(f8566,plain,(
  $false | $spl30),
  inference(equality_resolution,[],[f444])).
fof(f8499,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f8494])).
fof(f8494,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl28),
  inference(superposition,[],[f442,f285])).
fof(f8498,plain,(
  $false | $spl28),
  inference(equality_resolution,[],[f442])).
fof(f8451,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f8439,f2915])).
fof(f8439,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f269,f3476])).
fof(f8450,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl276),
  inference(cnf_transformation,[],[f8450_D])).
fof(f8450_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f8449,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f8438,f246])).
fof(f8438,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f269,f3476])).
fof(f8447,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f8432,f246])).
fof(f8432,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(superposition,[],[f269,f3074])).
fof(f8444,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f8431,f246])).
fof(f8431,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f269,f3074])).
fof(f1015,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) ) | $spl36),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8421,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f262])).
fof(f8413,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f269,f203])).
fof(f8410,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f202])).
fof(f8409,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f201])).
fof(f8408,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f269,f199])).
fof(f8406,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f197])).
fof(f7577,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f3476])).
fof(f8393,plain,(
  ( ! [X502,X504,X503] : (s(bool,'T_0') = s(bool,X502) | s(bool,X503) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X503) = s(bool,X504) | s(bool,'T_0') = s(bool,X504)) )),
  inference(subsumption_resolution,[],[f7521,f261])).
fof(f7521,plain,(
  ( ! [X502,X504,X503] : (s(bool,'F_0') != s(bool,X502) | s(bool,'T_0') = s(bool,X502) | s(bool,X503) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X503) = s(bool,X504) | s(bool,'T_0') = s(bool,X504)) )),
  inference(superposition,[],[f3407,f3476])).
fof(f7734,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X49) | s(bool,X48) = s(bool,X49)) ) | $spl134),
  inference(cnf_transformation,[],[f7734_D])).
fof(f7734_D,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X50) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X49) | s(bool,X48) = s(bool,X49)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f8384,plain,(
  ( ! [X432,X434,X431,X433,X435] : (s(bool,'T_0') = s(bool,X433) | s(bool,X431) = s(bool,X433) | s(bool,'T_0') = s(bool,X432) | s(bool,X431) = s(bool,X434) | s(bool,X434) = s(bool,X435) | s(bool,'T_0') = s(bool,X435)) )),
  inference(subsumption_resolution,[],[f7501,f261])).
fof(f7501,plain,(
  ( ! [X432,X434,X431,X433,X435] : (s(bool,'F_0') != s(bool,X432) | s(bool,'T_0') = s(bool,X433) | s(bool,X431) = s(bool,X433) | s(bool,'T_0') = s(bool,X432) | s(bool,X431) = s(bool,X434) | s(bool,X434) = s(bool,X435) | s(bool,'T_0') = s(bool,X435)) )),
  inference(superposition,[],[f4329,f3476])).
fof(f8372,plain,(
  ( ! [X350,X349,X351,X347] : (p(s(bool,X349)) | s(bool,'T_0') = s(bool,X351) | s(bool,X350) = s(bool,X351) | s(bool,X347) = s(bool,X350) | s(bool,X347) = s(bool,X349)) ) | $spl270),
  inference(cnf_transformation,[],[f8372_D])).
fof(f8372_D,plain,(
  ( ! [X350,X349,X351,X347] : (p(s(bool,X349)) | s(bool,'T_0') = s(bool,X351) | s(bool,X350) = s(bool,X351) | s(bool,X347) = s(bool,X350) | s(bool,X347) = s(bool,X349)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f8371,plain,(
  ( ! [X350,X349,X351,X347,X348] : (p(s(bool,X349)) | p(s(bool,X348)) | s(bool,X347) = s(bool,X349) | s(bool,X347) = s(bool,X350) | s(bool,X350) = s(bool,X351) | s(bool,'T_0') = s(bool,X351)) ) | $spl62),
  inference(subsumption_resolution,[],[f7483,f243])).
fof(f7483,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s(bool,'F_0') != s(bool,X348) | p(s(bool,X349)) | p(s(bool,X348)) | s(bool,X347) = s(bool,X349) | s(bool,X347) = s(bool,X350) | s(bool,X350) = s(bool,X351) | s(bool,'T_0') = s(bool,X351)) ) | $spl62),
  inference(superposition,[],[f2964,f3476])).
fof(f5660,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | $spl118),
  inference(cnf_transformation,[],[f5660_D])).
fof(f5660_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f8362,plain,(
  ( ! [X277,X273,X275,X276,X274] : (p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X273) != s(bool,X275) | s(bool,X273) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'T_0') = s(bool,X277)) )),
  inference(subsumption_resolution,[],[f7465,f250])).
fof(f7465,plain,(
  ( ! [X277,X273,X275,X276,X274] : (p(s(bool,'F_0')) | p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X273) != s(bool,X275) | s(bool,X273) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'T_0') = s(bool,X277)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8355,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (~p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X186) != s(bool,X188) | s(bool,X187) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(subsumption_resolution,[],[f7441,f250])).
fof(f7441,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,'F_0')) | ~p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X186) != s(bool,X188) | s(bool,X187) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8350,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X142) | s(bool,X142) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f7433,f250])).
fof(f7433,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X142) | s(bool,X142) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8347,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X109) != s(bool,X111) | s(bool,X108) != s(bool,X110) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f7428,f250])).
fof(f7428,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,'F_0')) | p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X109) != s(bool,X111) | s(bool,X108) != s(bool,X110) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8342,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X66) != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f7421,f250])).
fof(f7421,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,'F_0')) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X66) != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8337,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,X24) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f7414,f250])).
fof(f7414,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'F_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,X24) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8331,plain,(
  ( ! [X472,X474,X471,X475] : (s(bool,X471) != s(bool,X472) | s(bool,'T_0') = s(bool,X475) | s(bool,X474) = s(bool,X475) | s(bool,X471) = s(bool,X474) | s(bool,'F_0') = s(bool,X472)) ) | $spl268),
  inference(cnf_transformation,[],[f8331_D])).
fof(f8331_D,plain,(
  ( ! [X472,X474,X471,X475] : (s(bool,X471) != s(bool,X472) | s(bool,'T_0') = s(bool,X475) | s(bool,X474) = s(bool,X475) | s(bool,X471) = s(bool,X474) | s(bool,'F_0') = s(bool,X472)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f8330,plain,(
  ( ! [X472,X474,X471,X475,X473] : (s(bool,X471) != s(bool,X472) | s(bool,'F_0') = s(bool,X472) | s(bool,'T_0') = s(bool,X473) | s(bool,X471) = s(bool,X474) | s(bool,X474) = s(bool,X475) | s(bool,'T_0') = s(bool,X475)) ) | $spl60),
  inference(subsumption_resolution,[],[f7404,f261])).
fof(f7404,plain,(
  ( ! [X472,X474,X471,X475,X473] : (s(bool,X471) != s(bool,X472) | s(bool,'F_0') = s(bool,X472) | s(bool,'F_0') != s(bool,X473) | s(bool,'T_0') = s(bool,X473) | s(bool,X471) = s(bool,X474) | s(bool,X474) = s(bool,X475) | s(bool,'T_0') = s(bool,X475)) ) | $spl60),
  inference(superposition,[],[f4623,f3476])).
fof(f8320,plain,(
  ( ! [X383,X380,X382,X381] : (s(bool,'T_0') != s(bool,X380) | s(bool,'F_0') != s(bool,X381) | s(bool,X380) = s(bool,X382) | s(bool,X382) = s(bool,X383) | s(bool,'T_0') = s(bool,X383)) )),
  inference(subsumption_resolution,[],[f7386,f2590])).
fof(f7386,plain,(
  ( ! [X383,X380,X382,X381] : (s(bool,'T_0') != s(bool,X380) | s(bool,'F_0') != s(bool,X381) | s(bool,'T_0') = s(bool,X381) | s(bool,X380) = s(bool,X382) | s(bool,X382) = s(bool,X383) | s(bool,'T_0') = s(bool,X383)) )),
  inference(superposition,[],[f2882,f3476])).
fof(f8318,plain,(
  ( ! [X377,X379,X376,X378] : (s(bool,'T_0') != s(bool,X376) | p(s(bool,X377)) | s(bool,X376) = s(bool,X378) | s(bool,X378) = s(bool,X379) | s(bool,'T_0') = s(bool,X379)) )),
  inference(subsumption_resolution,[],[f7385,f243])).
fof(f7385,plain,(
  ( ! [X377,X379,X376,X378] : (s(bool,'T_0') != s(bool,X376) | s(bool,'F_0') != s(bool,X377) | p(s(bool,X377)) | s(bool,X376) = s(bool,X378) | s(bool,X378) = s(bool,X379) | s(bool,'T_0') = s(bool,X379)) )),
  inference(superposition,[],[f2881,f3476])).
fof(f8313,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | p(s(bool,X313)) | ~p(s(bool,X312)) | s(bool,X311) = s(bool,X314) | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(subsumption_resolution,[],[f7369,f250])).
fof(f7369,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | p(s(bool,X313)) | ~p(s(bool,X312)) | p(s(bool,'F_0')) | s(bool,X311) = s(bool,X314) | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8306,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | ~p(s(bool,X223)) | ~p(s(bool,X222)) | s(bool,X223) != s(bool,X224) | s(bool,X221) = s(bool,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(subsumption_resolution,[],[f7344,f250])).
fof(f7344,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | p(s(bool,'F_0')) | ~p(s(bool,X223)) | ~p(s(bool,X222)) | s(bool,X223) != s(bool,X224) | s(bool,X221) = s(bool,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8304,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,X180)) | ~p(s(bool,X182)) | s(bool,X179) = s(bool,X183) | s(bool,X183) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(subsumption_resolution,[],[f7337,f250])).
fof(f7337,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,X180)) | ~p(s(bool,X182)) | p(s(bool,'F_0')) | s(bool,X179) = s(bool,X183) | s(bool,X183) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8301,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X137) != s(bool,X138) | p(s(bool,X139)) | ~p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,X137) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(subsumption_resolution,[],[f7330,f250])).
fof(f7330,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X137) != s(bool,X138) | p(s(bool,X139)) | ~p(s(bool,X138)) | p(s(bool,'F_0')) | s(bool,X139) != s(bool,X140) | s(bool,X137) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8298,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,X107) != s(bool,X108) | ~p(s(bool,X109)) | ~p(s(bool,X108)) | s(bool,X109) != s(bool,X110) | s(bool,X107) = s(bool,X111) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f7325,f250])).
fof(f7325,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,X107) != s(bool,X108) | p(s(bool,'F_0')) | ~p(s(bool,X109)) | ~p(s(bool,X108)) | s(bool,X109) != s(bool,X110) | s(bool,X107) = s(bool,X111) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8295,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X49) != s(bool,X50) | s(bool,X47) = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f7315,f250])).
fof(f7315,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X48)) | p(s(bool,'F_0')) | s(bool,X49) != s(bool,X50) | s(bool,X47) = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8291,plain,(
  ( ! [X635,X631,X632,X634,X630] : (s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X631))) = s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X632))) | s(bool,'F_0') = s(bool,X634) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X630),X631),s('type.closSem.state'(X630),X632))) | s(bool,X634) = s(bool,X635)) ) | $spl264),
  inference(cnf_transformation,[],[f8291_D])).
fof(f8291_D,plain,(
  ( ! [X635,X631,X632,X634,X630] : (s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X631))) = s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X632))) | s(bool,'F_0') = s(bool,X634) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X630),X631),s('type.closSem.state'(X630),X632))) | s(bool,X634) = s(bool,X635)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f8290,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X631))) = s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X632))) | ~p(s(bool,X633)) | s(bool,X634) = s(bool,X635) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X630),X631),s('type.closSem.state'(X630),X632))) | s(bool,'F_0') = s(bool,X634)) ) | $spl70),
  inference(subsumption_resolution,[],[f7294,f247])).
fof(f7294,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (s(bool,'T_0') != s(bool,X633) | s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X631))) = s('type.ffi.ffi_state'(X630),'const.closSem.state_ffi_1'(s('type.closSem.state'(X630),X632))) | ~p(s(bool,X633)) | s(bool,X634) = s(bool,X635) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X630),X631),s('type.closSem.state'(X630),X632))) | s(bool,'F_0') = s(bool,X634)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f8288,plain,(
  ( ! [X619,X615,X616,X618,X614] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X615))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X616))) | s(bool,'F_0') = s(bool,X618) | s(bool,X619) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X614),X615),s('type.closSem.state'(X614),X616))) | s(bool,X618) = s(bool,X619)) ) | $spl262),
  inference(cnf_transformation,[],[f8288_D])).
fof(f8288_D,plain,(
  ( ! [X619,X615,X616,X618,X614] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X615))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X616))) | s(bool,'F_0') = s(bool,X618) | s(bool,X619) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X614),X615),s('type.closSem.state'(X614),X616))) | s(bool,X618) = s(bool,X619)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f8287,plain,(
  ( ! [X617,X619,X615,X616,X618,X614] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X615))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X616))) | ~p(s(bool,X617)) | s(bool,X618) = s(bool,X619) | s(bool,X619) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X614),X615),s('type.closSem.state'(X614),X616))) | s(bool,'F_0') = s(bool,X618)) ) | $spl64),
  inference(subsumption_resolution,[],[f7291,f247])).
fof(f7291,plain,(
  ( ! [X617,X619,X615,X616,X618,X614] : (s(bool,'T_0') != s(bool,X617) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X615))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X614),X616))) | ~p(s(bool,X617)) | s(bool,X618) = s(bool,X619) | s(bool,X619) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X614),X615),s('type.closSem.state'(X614),X616))) | s(bool,'F_0') = s(bool,X618)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f8278,plain,(
  ( ! [X549,X547,X548,X550] : (s(bool,X547) != s(bool,X548) | s(bool,'T_0') != s(bool,X548) | s(bool,X549) = s(bool,X550) | s(bool,'F_0') = s(bool,X550) | s(bool,'F_0') = s(bool,X549)) ) | $spl58),
  inference(subsumption_resolution,[],[f7274,f2882])).
fof(f7274,plain,(
  ( ! [X549,X547,X548,X550] : (s(bool,'T_0') != s(bool,X547) | s(bool,X547) != s(bool,X548) | s(bool,'T_0') != s(bool,X548) | s(bool,X549) = s(bool,X550) | s(bool,'F_0') = s(bool,X550) | s(bool,'F_0') = s(bool,X549)) ) | $spl58),
  inference(superposition,[],[f5870,f3476])).
fof(f8277,plain,(
  ( ! [X546,X545,X543,X544] : (s(bool,'F_0') = s(bool,X544) | s(bool,X543) != s(bool,X544) | s(bool,X545) = s(bool,X546) | s(bool,'F_0') = s(bool,X546) | s(bool,'F_0') = s(bool,X545)) ) | ($spl6 | $spl60)),
  inference(subsumption_resolution,[],[f7273,f828])).
fof(f7273,plain,(
  ( ! [X546,X545,X543,X544] : (s(bool,'T_0') != s(bool,X543) | s(bool,'F_0') = s(bool,X544) | s(bool,X543) != s(bool,X544) | s(bool,X545) = s(bool,X546) | s(bool,'F_0') = s(bool,X546) | s(bool,'F_0') = s(bool,X545)) ) | $spl60),
  inference(superposition,[],[f3111,f3476])).
fof(f8276,plain,(
  ( ! [X540,X542,X541,X539] : (s(bool,'T_0') != s(bool,X539) | s(bool,'T_0') = s(bool,X540) | s(bool,X541) = s(bool,X542) | s(bool,'F_0') = s(bool,X542) | s(bool,'F_0') = s(bool,X541)) )),
  inference(subsumption_resolution,[],[f7272,f2882])).
fof(f7272,plain,(
  ( ! [X540,X542,X541,X539] : (s(bool,'T_0') != s(bool,X539) | s(bool,'T_0') = s(bool,X540) | s(bool,X539) = s(bool,X540) | s(bool,X541) = s(bool,X542) | s(bool,'F_0') = s(bool,X542) | s(bool,'F_0') = s(bool,X541)) )),
  inference(superposition,[],[f3073,f3476])).
fof(f8275,plain,(
  ( ! [X536,X538,X537,X535] : (s(bool,'T_0') != s(bool,X535) | p(s(bool,X536)) | s(bool,X537) = s(bool,X538) | s(bool,'F_0') = s(bool,X538) | s(bool,'F_0') = s(bool,X537)) )),
  inference(subsumption_resolution,[],[f7271,f2881])).
fof(f7271,plain,(
  ( ! [X536,X538,X537,X535] : (s(bool,'T_0') != s(bool,X535) | s(bool,X535) = s(bool,X536) | p(s(bool,X536)) | s(bool,X537) = s(bool,X538) | s(bool,'F_0') = s(bool,X538) | s(bool,'F_0') = s(bool,X537)) )),
  inference(superposition,[],[f2915,f3476])).
fof(f8274,plain,(
  ( ! [X532,X534,X533,X531] : (s(bool,X531) != s(bool,X532) | ~p(s(bool,X532)) | s(bool,X533) = s(bool,X534) | s(bool,'F_0') = s(bool,X534) | s(bool,'F_0') = s(bool,X533)) ) | ($spl6 | $spl58)),
  inference(subsumption_resolution,[],[f7270,f766])).
fof(f7270,plain,(
  ( ! [X532,X534,X533,X531] : (s(bool,'T_0') != s(bool,X531) | s(bool,X531) != s(bool,X532) | ~p(s(bool,X532)) | s(bool,X533) = s(bool,X534) | s(bool,'F_0') = s(bool,X534) | s(bool,'F_0') = s(bool,X533)) ) | $spl58),
  inference(superposition,[],[f2634,f3476])).
fof(f7621,plain,(
  ( ! [X528,X530,X529] : (s(bool,'T_0') != s(bool,X528) | s(bool,X529) = s(bool,X530) | s(bool,'F_0') = s(bool,X530) | s(bool,'F_0') = s(bool,X529)) )),
  inference(duplicate_literal_removal,[],[f7269])).
fof(f7269,plain,(
  ( ! [X528,X530,X529] : (s(bool,'T_0') != s(bool,X528) | s(bool,'T_0') != s(bool,X528) | s(bool,X529) = s(bool,X530) | s(bool,'F_0') = s(bool,X530) | s(bool,'F_0') = s(bool,X529)) )),
  inference(superposition,[],[f2590,f3476])).
fof(f7623,plain,(
  ( ! [X524,X525,X523] : (s(bool,'T_0') = s(bool,X523) | s(bool,X524) = s(bool,X525) | s(bool,'F_0') = s(bool,X525) | s(bool,'F_0') = s(bool,X524)) )),
  inference(duplicate_literal_removal,[],[f7267])).
fof(f7267,plain,(
  ( ! [X524,X525,X523] : (s(bool,'T_0') = s(bool,X523) | s(bool,'T_0') = s(bool,X523) | s(bool,X524) = s(bool,X525) | s(bool,'F_0') = s(bool,X525) | s(bool,'F_0') = s(bool,X524)) )),
  inference(superposition,[],[f261,f3476])).
fof(f8272,plain,(
  ( ! [X520,X518,X519] : (p(s(bool,X518)) | s(bool,X519) = s(bool,X520) | s(bool,'F_0') = s(bool,X520) | s(bool,'F_0') = s(bool,X519)) )),
  inference(subsumption_resolution,[],[f7265,f246])).
fof(f7265,plain,(
  ( ! [X520,X518,X519] : (s(bool,'T_0') = s(bool,X518) | p(s(bool,X518)) | s(bool,X519) = s(bool,X520) | s(bool,'F_0') = s(bool,X520) | s(bool,'F_0') = s(bool,X519)) )),
  inference(superposition,[],[f243,f3476])).
fof(f8271,plain,(
  ( ! [X516,X517,X515] : (~p(s(bool,X515)) | s(bool,X516) = s(bool,X517) | s(bool,'F_0') = s(bool,X517) | s(bool,'F_0') = s(bool,X516)) )),
  inference(subsumption_resolution,[],[f7264,f247])).
fof(f7264,plain,(
  ( ! [X516,X517,X515] : (s(bool,'T_0') != s(bool,X515) | ~p(s(bool,X515)) | s(bool,X516) = s(bool,X517) | s(bool,'F_0') = s(bool,X517) | s(bool,'F_0') = s(bool,X516)) )),
  inference(superposition,[],[f242,f3476])).
fof(f8049,plain,(
  ( ! [X428,X429,X425,X427] : (s(bool,X425) = s(bool,X427) | s(bool,'F_0') = s(bool,X429) | s(bool,X425) = s(bool,X429) | s(bool,'T_0') = s(bool,X428) | s(bool,X427) != s(bool,X428)) ) | $spl224),
  inference(cnf_transformation,[],[f8049_D])).
fof(f8049_D,plain,(
  ( ! [X428,X429,X425,X427] : (s(bool,X425) = s(bool,X427) | s(bool,'F_0') = s(bool,X429) | s(bool,X425) = s(bool,X429) | s(bool,'T_0') = s(bool,X428) | s(bool,X427) != s(bool,X428)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8263,plain,(
  ( ! [X428,X430,X432,X429,X431] : (s(bool,'T_0') != s(bool,X429) | s(bool,X428) != s(bool,X430) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X432) | s(bool,X428) = s(bool,X432) | s(bool,'F_0') = s(bool,X431)) ) | $spl60),
  inference(subsumption_resolution,[],[f7241,f2590])).
fof(f7241,plain,(
  ( ! [X428,X430,X432,X429,X431] : (s(bool,'T_0') != s(bool,X429) | s(bool,'F_0') = s(bool,X429) | s(bool,X428) != s(bool,X430) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X432) | s(bool,X428) = s(bool,X432) | s(bool,'F_0') = s(bool,X431)) ) | $spl60),
  inference(superposition,[],[f4623,f3476])).
fof(f8258,plain,(
  ( ! [X350,X352,X354,X351,X353] : (s(bool,'T_0') != s(bool,X351) | p(s(bool,X352)) | p(s(bool,X350)) | s(bool,X353) = s(bool,X354) | s(bool,X350) = s(bool,X354) | s(bool,'F_0') = s(bool,X353)) ) | $spl62),
  inference(subsumption_resolution,[],[f7225,f2881])).
fof(f7225,plain,(
  ( ! [X350,X352,X354,X351,X353] : (s(bool,'T_0') != s(bool,X351) | p(s(bool,X352)) | p(s(bool,X350)) | s(bool,X351) = s(bool,X352) | s(bool,X353) = s(bool,X354) | s(bool,X350) = s(bool,X354) | s(bool,'F_0') = s(bool,X353)) ) | $spl62),
  inference(superposition,[],[f2964,f3476])).
fof(f8255,plain,(
  ( ! [X288,X287,X289,X286] : (s(bool,'T_0') != s(bool,X287) | p(s(bool,X286)) | s(bool,X288) = s(bool,X289) | s(bool,X286) = s(bool,X289) | s(bool,'F_0') = s(bool,X288)) ) | $spl6),
  inference(subsumption_resolution,[],[f7209,f246])).
fof(f7209,plain,(
  ( ! [X288,X287,X289,X286] : (s(bool,'T_0') != s(bool,X287) | p(s(bool,X286)) | ~p(s(bool,X287)) | s(bool,X288) = s(bool,X289) | s(bool,X286) = s(bool,X289) | s(bool,'F_0') = s(bool,X288)) ) | $spl6),
  inference(superposition,[],[f297,f3476])).
fof(f8199,plain,(
  ( ! [X101,X105,X102,X106] : (s(bool,X101) != s(bool,X102) | s(bool,'F_0') = s(bool,X105) | s(bool,X101) = s(bool,X106) | s(bool,X105) = s(bool,X106) | p(s(bool,X102))) ) | $spl258),
  inference(cnf_transformation,[],[f8199_D])).
fof(f8199_D,plain,(
  ( ! [X101,X105,X102,X106] : (s(bool,X101) != s(bool,X102) | s(bool,'F_0') = s(bool,X105) | s(bool,X101) = s(bool,X106) | s(bool,X105) = s(bool,X106) | p(s(bool,X102))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f8254,plain,(
  ( ! [X282,X285,X281,X283,X284] : (p(s(bool,X282)) | p(s(bool,X283)) | s(bool,X281) != s(bool,X283) | s(bool,X284) = s(bool,X285) | s(bool,X281) = s(bool,X285) | s(bool,'F_0') = s(bool,X284)) )),
  inference(subsumption_resolution,[],[f7208,f263])).
fof(f7208,plain,(
  ( ! [X282,X285,X281,X283,X284] : (~p(s(bool,'T_0')) | p(s(bool,X282)) | p(s(bool,X283)) | s(bool,X281) != s(bool,X283) | s(bool,X284) = s(bool,X285) | s(bool,X281) = s(bool,X285) | s(bool,'F_0') = s(bool,X284)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8253,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s(bool,'T_0') != s(bool,X267) | p(s(bool,X268)) | p(s(bool,X266)) | s(bool,X269) = s(bool,X270) | s(bool,X266) = s(bool,X270) | s(bool,'F_0') = s(bool,X269)) )),
  inference(subsumption_resolution,[],[f7205,f246])).
fof(f7205,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s(bool,'T_0') != s(bool,X267) | p(s(bool,X268)) | ~p(s(bool,X267)) | p(s(bool,X266)) | s(bool,X269) = s(bool,X270) | s(bool,X266) = s(bool,X270) | s(bool,'F_0') = s(bool,X269)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8250,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (p(s(bool,X199)) | ~p(s(bool,X200)) | s(bool,X198) != s(bool,X199) | s(bool,X200) != s(bool,X201) | s(bool,X202) = s(bool,X203) | s(bool,X198) = s(bool,X203) | s(bool,'F_0') = s(bool,X202)) )),
  inference(subsumption_resolution,[],[f7185,f263])).
fof(f7185,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X199)) | ~p(s(bool,X200)) | s(bool,X198) != s(bool,X199) | s(bool,X200) != s(bool,X201) | s(bool,X202) = s(bool,X203) | s(bool,X198) = s(bool,X203) | s(bool,'F_0') = s(bool,X202)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8248,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,'T_0') != s(bool,X175) | p(s(bool,X174)) | ~p(s(bool,X176)) | s(bool,X176) != s(bool,X177) | s(bool,X178) = s(bool,X179) | s(bool,X174) = s(bool,X179) | s(bool,'F_0') = s(bool,X178)) )),
  inference(subsumption_resolution,[],[f7181,f246])).
fof(f7181,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,'T_0') != s(bool,X175) | p(s(bool,X174)) | ~p(s(bool,X176)) | ~p(s(bool,X175)) | s(bool,X176) != s(bool,X177) | s(bool,X178) = s(bool,X179) | s(bool,X174) = s(bool,X179) | s(bool,'F_0') = s(bool,X178)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8245,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X150) != s(bool,X153) | s(bool,X154) = s(bool,X155) | s(bool,X150) = s(bool,X155) | s(bool,'F_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f7177,f263])).
fof(f7177,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X150) != s(bool,X153) | s(bool,X154) = s(bool,X155) | s(bool,X150) = s(bool,X155) | s(bool,'F_0') = s(bool,X154)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8244,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,X133) | p(s(bool,X134)) | ~p(s(bool,X135)) | p(s(bool,X132)) | s(bool,X136) = s(bool,X137) | s(bool,X132) = s(bool,X137) | s(bool,'F_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f7174,f246])).
fof(f7174,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,X133) | p(s(bool,X134)) | ~p(s(bool,X133)) | ~p(s(bool,X135)) | p(s(bool,X132)) | s(bool,X136) = s(bool,X137) | s(bool,X132) = s(bool,X137) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8243,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X120) != s(bool,X122) | s(bool,X124) = s(bool,X125) | s(bool,X120) = s(bool,X125) | s(bool,'F_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f7172,f263])).
fof(f7172,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,'T_0')) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X120) != s(bool,X122) | s(bool,X124) = s(bool,X125) | s(bool,X120) = s(bool,X125) | s(bool,'F_0') = s(bool,X124)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8240,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') != s(bool,X91) | p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X92) != s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,X90) = s(bool,X95) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f7167,f246])).
fof(f7167,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') != s(bool,X91) | p(s(bool,X92)) | ~p(s(bool,X91)) | p(s(bool,X90)) | s(bool,X92) != s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,X90) = s(bool,X95) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8239,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X78) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X82) = s(bool,X83) | s(bool,X78) = s(bool,X83) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f7165,f263])).
fof(f7165,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X78) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X82) = s(bool,X83) | s(bool,X78) = s(bool,X83) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8237,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X60)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,X60) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f7162,f246])).
fof(f7162,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X60)) | ~p(s(bool,X62)) | ~p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,X60) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8234,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38) | s(bool,X40) = s(bool,X41) | s(bool,X36) = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7158,f263])).
fof(f7158,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38) | s(bool,X40) = s(bool,X41) | s(bool,X36) = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8231,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7152,f246])).
fof(f7152,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8229,plain,(
  ( ! [X494,X496,X493,X495,X497] : (p(s(bool,X493)) | s(bool,X494) != s(bool,X495) | s(bool,'T_0') != s(bool,X495) | s(bool,X496) = s(bool,X497) | s(bool,X493) = s(bool,X497) | s(bool,'F_0') = s(bool,X496)) ) | $spl58),
  inference(subsumption_resolution,[],[f7149,f2882])).
fof(f7149,plain,(
  ( ! [X494,X496,X493,X495,X497] : (p(s(bool,X493)) | s(bool,'T_0') != s(bool,X494) | s(bool,X494) != s(bool,X495) | s(bool,'T_0') != s(bool,X495) | s(bool,X496) = s(bool,X497) | s(bool,X493) = s(bool,X497) | s(bool,'F_0') = s(bool,X496)) ) | $spl58),
  inference(superposition,[],[f4702,f3476])).
fof(f7633,plain,(
  ( ! [X478,X480,X482,X479,X481] : (s(bool,X478) != s(bool,X479) | s(bool,'T_0') = s(bool,X480) | s(bool,'T_0') = s(bool,X479) | s(bool,X481) = s(bool,X482) | s(bool,X478) = s(bool,X482) | s(bool,'F_0') = s(bool,X481)) )),
  inference(duplicate_literal_removal,[],[f7146])).
fof(f7146,plain,(
  ( ! [X478,X480,X482,X479,X481] : (s(bool,X478) != s(bool,X479) | s(bool,'T_0') = s(bool,X480) | s(bool,'T_0') = s(bool,X480) | s(bool,'T_0') = s(bool,X479) | s(bool,X481) = s(bool,X482) | s(bool,X478) = s(bool,X482) | s(bool,'F_0') = s(bool,X481)) )),
  inference(superposition,[],[f4329,f3476])).
fof(f8227,plain,(
  ( ! [X472,X474,X471,X473,X470] : (p(s(bool,X470)) | s(bool,'T_0') = s(bool,X471) | s(bool,'T_0') = s(bool,X472) | s(bool,X473) = s(bool,X474) | s(bool,X470) = s(bool,X474) | s(bool,'F_0') = s(bool,X473)) )),
  inference(subsumption_resolution,[],[f7144,f3074])).
fof(f7144,plain,(
  ( ! [X472,X474,X471,X473,X470] : (p(s(bool,X470)) | s(bool,'T_0') = s(bool,X471) | s(bool,X471) != s(bool,X472) | s(bool,'T_0') = s(bool,X472) | s(bool,X473) = s(bool,X474) | s(bool,X470) = s(bool,X474) | s(bool,'F_0') = s(bool,X473)) )),
  inference(superposition,[],[f4328,f3476])).
fof(f8226,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X455)) | s(bool,X456) != s(bool,X457) | s(bool,'F_0') = s(bool,X457) | s(bool,X458) = s(bool,X459) | s(bool,X455) = s(bool,X459) | s(bool,'F_0') = s(bool,X458)) ) | ($spl6 | $spl68)),
  inference(subsumption_resolution,[],[f7142,f828])).
fof(f7142,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X455)) | s(bool,'T_0') != s(bool,X456) | s(bool,X456) != s(bool,X457) | s(bool,'F_0') = s(bool,X457) | s(bool,X458) = s(bool,X459) | s(bool,X455) = s(bool,X459) | s(bool,'F_0') = s(bool,X458)) ) | $spl68),
  inference(superposition,[],[f3130,f3476])).
fof(f8225,plain,(
  ( ! [X453,X449,X451,X452,X450] : (p(s(bool,X449)) | s(bool,'T_0') != s(bool,X450) | s(bool,'T_0') = s(bool,X451) | s(bool,X452) = s(bool,X453) | s(bool,X449) = s(bool,X453) | s(bool,'F_0') = s(bool,X452)) ) | $spl66),
  inference(subsumption_resolution,[],[f7141,f2882])).
fof(f7141,plain,(
  ( ! [X453,X449,X451,X452,X450] : (p(s(bool,X449)) | s(bool,'T_0') != s(bool,X450) | s(bool,X450) = s(bool,X451) | s(bool,'T_0') = s(bool,X451) | s(bool,X452) = s(bool,X453) | s(bool,X449) = s(bool,X453) | s(bool,'F_0') = s(bool,X452)) ) | $spl66),
  inference(superposition,[],[f3095,f3476])).
fof(f8224,plain,(
  ( ! [X420,X422,X418,X421,X419] : (s(bool,X418) != s(bool,X419) | p(s(bool,X420)) | p(s(bool,X419)) | s(bool,X421) = s(bool,X422) | s(bool,X418) = s(bool,X422) | s(bool,'F_0') = s(bool,X421)) ) | $spl62),
  inference(subsumption_resolution,[],[f7135,f246])).
fof(f7135,plain,(
  ( ! [X420,X422,X418,X421,X419] : (s(bool,X418) != s(bool,X419) | p(s(bool,X420)) | p(s(bool,X419)) | s(bool,'T_0') = s(bool,X420) | s(bool,X421) = s(bool,X422) | s(bool,X418) = s(bool,X422) | s(bool,'F_0') = s(bool,X421)) ) | $spl62),
  inference(superposition,[],[f2964,f3476])).
fof(f8223,plain,(
  ( ! [X408,X405,X407,X406] : (p(s(bool,X405)) | s(bool,'T_0') != s(bool,X406) | s(bool,X407) = s(bool,X408) | s(bool,X405) = s(bool,X408) | s(bool,'F_0') = s(bool,X407)) ) | $spl60),
  inference(subsumption_resolution,[],[f7132,f2590])).
fof(f7132,plain,(
  ( ! [X408,X405,X407,X406] : (p(s(bool,X405)) | s(bool,'T_0') != s(bool,X406) | s(bool,'F_0') = s(bool,X406) | s(bool,X407) = s(bool,X408) | s(bool,X405) = s(bool,X408) | s(bool,'F_0') = s(bool,X407)) ) | $spl60),
  inference(superposition,[],[f2856,f3476])).
fof(f7635,plain,(
  ( ! [X401,X403,X404,X402] : (p(s(bool,X401)) | s(bool,'T_0') = s(bool,X402) | s(bool,X403) = s(bool,X404) | s(bool,X401) = s(bool,X404) | s(bool,'F_0') = s(bool,X403)) )),
  inference(duplicate_literal_removal,[],[f7131])).
fof(f7131,plain,(
  ( ! [X401,X403,X404,X402] : (p(s(bool,X401)) | s(bool,'T_0') = s(bool,X402) | s(bool,'T_0') = s(bool,X402) | s(bool,X403) = s(bool,X404) | s(bool,X401) = s(bool,X404) | s(bool,'F_0') = s(bool,X403)) )),
  inference(superposition,[],[f2660,f3476])).
fof(f8222,plain,(
  ( ! [X397,X399,X398,X400] : (p(s(bool,X397)) | p(s(bool,X398)) | s(bool,X399) = s(bool,X400) | s(bool,X397) = s(bool,X400) | s(bool,'F_0') = s(bool,X399)) )),
  inference(subsumption_resolution,[],[f7130,f246])).
fof(f7130,plain,(
  ( ! [X397,X399,X398,X400] : (p(s(bool,X397)) | p(s(bool,X398)) | s(bool,'T_0') = s(bool,X398) | s(bool,X399) = s(bool,X400) | s(bool,X397) = s(bool,X400) | s(bool,'F_0') = s(bool,X399)) )),
  inference(superposition,[],[f2659,f3476])).
fof(f8221,plain,(
  ( ! [X393,X395,X396,X394] : (p(s(bool,X393)) | p(s(bool,X394)) | s(bool,X395) = s(bool,X396) | s(bool,X393) = s(bool,X396) | s(bool,'F_0') = s(bool,X395)) )),
  inference(subsumption_resolution,[],[f7129,f246])).
fof(f7129,plain,(
  ( ! [X393,X395,X396,X394] : (p(s(bool,X393)) | p(s(bool,X394)) | s(bool,'T_0') = s(bool,X394) | s(bool,X395) = s(bool,X396) | s(bool,X393) = s(bool,X396) | s(bool,'F_0') = s(bool,X395)) )),
  inference(superposition,[],[f2659,f3476])).
fof(f8217,plain,(
  ( ! [X356,X358,X360,X357,X359] : (p(s(bool,X356)) | p(s(bool,X357)) | s(bool,'T_0') != s(bool,X358) | s(bool,X359) = s(bool,X360) | s(bool,X356) = s(bool,X360) | s(bool,'F_0') = s(bool,X359)) )),
  inference(subsumption_resolution,[],[f7120,f246])).
fof(f7120,plain,(
  ( ! [X356,X358,X360,X357,X359] : (p(s(bool,X356)) | p(s(bool,X357)) | ~p(s(bool,X358)) | s(bool,'T_0') != s(bool,X358) | s(bool,X359) = s(bool,X360) | s(bool,X356) = s(bool,X360) | s(bool,'F_0') = s(bool,X359)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8216,plain,(
  ( ! [X350,X346,X349,X347,X348] : (s(bool,X346) != s(bool,X347) | p(s(bool,X348)) | p(s(bool,X347)) | s(bool,X349) = s(bool,X350) | s(bool,X346) = s(bool,X350) | s(bool,'F_0') = s(bool,X349)) )),
  inference(subsumption_resolution,[],[f7118,f263])).
fof(f7118,plain,(
  ( ! [X350,X346,X349,X347,X348] : (s(bool,X346) != s(bool,X347) | p(s(bool,X348)) | ~p(s(bool,'T_0')) | p(s(bool,X347)) | s(bool,X349) = s(bool,X350) | s(bool,X346) = s(bool,X350) | s(bool,'F_0') = s(bool,X349)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8213,plain,(
  ( ! [X282,X285,X281,X283,X284,X286] : (p(s(bool,X281)) | ~p(s(bool,X282)) | s(bool,'T_0') != s(bool,X283) | s(bool,X282) != s(bool,X284) | s(bool,X285) = s(bool,X286) | s(bool,X281) = s(bool,X286) | s(bool,'F_0') = s(bool,X285)) )),
  inference(subsumption_resolution,[],[f7100,f246])).
fof(f7100,plain,(
  ( ! [X282,X285,X281,X283,X284,X286] : (p(s(bool,X281)) | ~p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,'T_0') != s(bool,X283) | s(bool,X282) != s(bool,X284) | s(bool,X285) = s(bool,X286) | s(bool,X281) = s(bool,X286) | s(bool,'F_0') = s(bool,X285)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8212,plain,(
  ( ! [X280,X277,X279,X275,X276,X278] : (s(bool,X275) != s(bool,X276) | p(s(bool,X276)) | ~p(s(bool,X277)) | s(bool,X277) != s(bool,X278) | s(bool,X279) = s(bool,X280) | s(bool,X275) = s(bool,X280) | s(bool,'F_0') = s(bool,X279)) )),
  inference(subsumption_resolution,[],[f7099,f263])).
fof(f7099,plain,(
  ( ! [X280,X277,X279,X275,X276,X278] : (s(bool,X275) != s(bool,X276) | p(s(bool,X276)) | ~p(s(bool,X277)) | ~p(s(bool,'T_0')) | s(bool,X277) != s(bool,X278) | s(bool,X279) = s(bool,X280) | s(bool,X275) = s(bool,X280) | s(bool,'F_0') = s(bool,X279)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8210,plain,(
  ( ! [X237,X235,X233,X238,X236,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | ~p(s(bool,X236)) | s(bool,'T_0') != s(bool,X235) | s(bool,X237) = s(bool,X238) | s(bool,X233) = s(bool,X238) | s(bool,'F_0') = s(bool,X237)) )),
  inference(subsumption_resolution,[],[f7092,f246])).
fof(f7092,plain,(
  ( ! [X237,X235,X233,X238,X236,X234] : (p(s(bool,X233)) | p(s(bool,X234)) | ~p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,'T_0') != s(bool,X235) | s(bool,X237) = s(bool,X238) | s(bool,X233) = s(bool,X238) | s(bool,'F_0') = s(bool,X237)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8209,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | p(s(bool,X223)) | ~p(s(bool,X224)) | p(s(bool,X222)) | s(bool,X225) = s(bool,X226) | s(bool,X221) = s(bool,X226) | s(bool,'F_0') = s(bool,X225)) )),
  inference(subsumption_resolution,[],[f7090,f263])).
fof(f7090,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | p(s(bool,X223)) | ~p(s(bool,'T_0')) | ~p(s(bool,X224)) | p(s(bool,X222)) | s(bool,X225) = s(bool,X226) | s(bool,X221) = s(bool,X226) | s(bool,'F_0') = s(bool,X225)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8208,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (p(s(bool,X203)) | p(s(bool,X204)) | s(bool,X204) != s(bool,X206) | s(bool,'T_0') != s(bool,X205) | s(bool,X207) = s(bool,X208) | s(bool,X203) = s(bool,X208) | s(bool,'F_0') = s(bool,X207)) )),
  inference(subsumption_resolution,[],[f7087,f246])).
fof(f7087,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (p(s(bool,X203)) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,'T_0') != s(bool,X205) | s(bool,X207) = s(bool,X208) | s(bool,X203) = s(bool,X208) | s(bool,'F_0') = s(bool,X207)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8206,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | p(s(bool,X180)) | s(bool,X181) != s(bool,X182) | s(bool,X183) = s(bool,X184) | s(bool,X179) = s(bool,X184) | s(bool,'F_0') = s(bool,X183)) )),
  inference(subsumption_resolution,[],[f7083,f263])).
fof(f7083,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,'T_0')) | p(s(bool,X180)) | s(bool,X181) != s(bool,X182) | s(bool,X183) = s(bool,X184) | s(bool,X179) = s(bool,X184) | s(bool,'F_0') = s(bool,X183)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8204,plain,(
  ( ! [X165,X163,X161,X166,X164,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,'T_0') != s(bool,X163) | s(bool,X162) != s(bool,X164) | s(bool,X165) = s(bool,X166) | s(bool,X161) = s(bool,X166) | s(bool,'F_0') = s(bool,X165)) )),
  inference(subsumption_resolution,[],[f7080,f246])).
fof(f7080,plain,(
  ( ! [X165,X163,X161,X166,X164,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | ~p(s(bool,X163)) | s(bool,'T_0') != s(bool,X163) | s(bool,X162) != s(bool,X164) | s(bool,X165) = s(bool,X166) | s(bool,X161) = s(bool,X166) | s(bool,'F_0') = s(bool,X165)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8203,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,X149) != s(bool,X150) | p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X151) != s(bool,X152) | s(bool,X153) = s(bool,X154) | s(bool,X149) = s(bool,X154) | s(bool,'F_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f7078,f263])).
fof(f7078,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,X149) != s(bool,X150) | p(s(bool,X150)) | ~p(s(bool,X151)) | ~p(s(bool,'T_0')) | s(bool,X151) != s(bool,X152) | s(bool,X153) = s(bool,X154) | s(bool,X149) = s(bool,X154) | s(bool,'F_0') = s(bool,X153)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8201,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,X122) | s(bool,'T_0') != s(bool,X121) | s(bool,X123) = s(bool,X124) | s(bool,X119) = s(bool,X124) | s(bool,'F_0') = s(bool,X123)) )),
  inference(subsumption_resolution,[],[f7073,f246])).
fof(f7073,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (p(s(bool,X119)) | p(s(bool,X120)) | ~p(s(bool,X121)) | s(bool,X120) != s(bool,X122) | s(bool,'T_0') != s(bool,X121) | s(bool,X123) = s(bool,X124) | s(bool,X119) = s(bool,X124) | s(bool,'F_0') = s(bool,X123)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8198,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (s(bool,X101) != s(bool,X102) | p(s(bool,X103)) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | s(bool,X105) = s(bool,X106) | s(bool,X101) = s(bool,X106) | s(bool,'F_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f7070,f263])).
fof(f7070,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (s(bool,X101) != s(bool,X102) | p(s(bool,X103)) | ~p(s(bool,'T_0')) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | s(bool,X105) = s(bool,X106) | s(bool,X101) = s(bool,X106) | s(bool,'F_0') = s(bool,X105)) )),
  inference(superposition,[],[f197,f3476])).
fof(f7056,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X26) = s(bool,X27) | s(bool,X24) = s(bool,X27) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f240,f3476])).
fof(f7055,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X22) = s(bool,X23) | s(bool,X20) = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f238,f3476])).
fof(f7054,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f225,f3476])).
fof(f7053,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) = s(bool,X15) | s(bool,X12) = s(bool,X15) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f224,f3476])).
fof(f7052,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X10) = s(bool,X11) | s(bool,X8) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f221,f3476])).
fof(f7051,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f218,f3476])).
fof(f7050,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f216,f3476])).
fof(f8197,plain,(
  ( ! [X685,X687,X683,X684,X686] : (p(s(bool,X686)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,X686) = s(bool,X687) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685)))))) )),
  inference(subsumption_resolution,[],[f7046,f2915])).
fof(f7046,plain,(
  ( ! [X685,X687,X683,X684,X686] : (p(s(bool,X686)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,X686) = s(bool,X687) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685))))) | s(bool,'F_0') = s(bool,X687)) )),
  inference(superposition,[],[f256,f3476])).
fof(f7752,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl142),
  inference(cnf_transformation,[],[f7752_D])).
fof(f7752_D,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7045,plain,(
  ( ! [X681,X679,X680,X682,X678] : (~p(s(bool,X681)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s(bool,X681) = s(bool,X682) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s(bool,'F_0') = s(bool,X682)) )),
  inference(superposition,[],[f258,f3476])).
fof(f8195,plain,(
  ( ! [X677,X673,X675,X676,X674] : (p(s(bool,X676)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,X676) = s(bool,X677) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675)))))) )),
  inference(subsumption_resolution,[],[f7044,f2915])).
fof(f7044,plain,(
  ( ! [X677,X673,X675,X676,X674] : (p(s(bool,X676)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,X676) = s(bool,X677) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675))))) | s(bool,'F_0') = s(bool,X677)) )),
  inference(superposition,[],[f255,f3476])).
fof(f7043,plain,(
  ( ! [X668,X670,X669,X671,X672] : (~p(s(bool,X671)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s(bool,'F_0') = s(bool,X672)) )),
  inference(superposition,[],[f258,f3476])).
fof(f8193,plain,(
  ( ! [X666,X664,X667,X663,X665] : (p(s(bool,X666)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,X666) = s(bool,X667) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665)))))) )),
  inference(subsumption_resolution,[],[f7042,f2915])).
fof(f7042,plain,(
  ( ! [X666,X664,X667,X663,X665] : (p(s(bool,X666)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,X666) = s(bool,X667) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665))))) | s(bool,'F_0') = s(bool,X667)) )),
  inference(superposition,[],[f257,f3476])).
fof(f7041,plain,(
  ( ! [X660,X662,X658,X661,X659] : (~p(s(bool,X661)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s(bool,X661) = s(bool,X662) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s(bool,'F_0') = s(bool,X662)) )),
  inference(superposition,[],[f258,f3476])).
fof(f8191,plain,(
  ( ! [X656,X652,X654,X657,X653,X655] : (s(bool,X655) != s(bool,X656) | s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X653))) = s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X654))) | ~p(s(bool,X656)) | s(bool,X655) = s(bool,X657) | s(bool,'F_0') = s(bool,X657)) ) | $spl70),
  inference(subsumption_resolution,[],[f7040,f3144])).
fof(f7040,plain,(
  ( ! [X656,X652,X654,X657,X653,X655] : (s(bool,X655) != s(bool,X656) | s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X653))) = s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X654))) | ~p(s(bool,X656)) | s(bool,X655) = s(bool,X657) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X652),X653),s('type.closSem.state'(X652),X654))) | s(bool,'F_0') = s(bool,X657)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f8190,plain,(
  ( ! [X648,X650,X649,X651,X647] : (s(bool,'T_0') != s(bool,X650) | s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X648))) = s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X649))) | s(bool,X650) = s(bool,X651) | s(bool,'F_0') = s(bool,X651)) )),
  inference(subsumption_resolution,[],[f7039,f3144])).
fof(f7039,plain,(
  ( ! [X648,X650,X649,X651,X647] : (s(bool,'T_0') != s(bool,X650) | s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X648))) = s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X649))) | s(bool,X650) = s(bool,X651) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X647),X648),s('type.closSem.state'(X647),X649))) | s(bool,'F_0') = s(bool,X651)) )),
  inference(superposition,[],[f3144,f3476])).
fof(f8189,plain,(
  ( ! [X644,X646,X642,X645,X643] : (s(bool,'F_0') = s(bool,X645) | s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X643))) = s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X644))) | s(bool,X645) = s(bool,X646) | s(bool,'F_0') = s(bool,X646)) )),
  inference(subsumption_resolution,[],[f7038,f3144])).
fof(f7038,plain,(
  ( ! [X644,X646,X642,X645,X643] : (s(bool,'F_0') = s(bool,X645) | s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X643))) = s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X644))) | s(bool,X645) = s(bool,X646) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X642),X643),s('type.closSem.state'(X642),X644))) | s(bool,'F_0') = s(bool,X646)) )),
  inference(superposition,[],[f3143,f3476])).
fof(f8188,plain,(
  ( ! [X639,X640,X636,X638,X641,X637] : (s(bool,X639) != s(bool,X640) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X637))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X638))) | ~p(s(bool,X640)) | s(bool,X639) = s(bool,X641) | s(bool,'F_0') = s(bool,X641)) ) | $spl64),
  inference(subsumption_resolution,[],[f7037,f2981])).
fof(f7037,plain,(
  ( ! [X639,X640,X636,X638,X641,X637] : (s(bool,X639) != s(bool,X640) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X637))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X638))) | ~p(s(bool,X640)) | s(bool,X639) = s(bool,X641) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X636),X637),s('type.closSem.state'(X636),X638))) | s(bool,'F_0') = s(bool,X641)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f8187,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'T_0') != s(bool,X634) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X632))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X633))) | s(bool,X634) = s(bool,X635) | s(bool,'F_0') = s(bool,X635)) )),
  inference(subsumption_resolution,[],[f7036,f2981])).
fof(f7036,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'T_0') != s(bool,X634) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X632))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X633))) | s(bool,X634) = s(bool,X635) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X631),X632),s('type.closSem.state'(X631),X633))) | s(bool,'F_0') = s(bool,X635)) )),
  inference(superposition,[],[f2981,f3476])).
fof(f8186,plain,(
  ( ! [X629,X627,X628,X630,X626] : (s(bool,'F_0') = s(bool,X629) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X627))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X628))) | s(bool,X629) = s(bool,X630) | s(bool,'F_0') = s(bool,X630)) )),
  inference(subsumption_resolution,[],[f7035,f2981])).
fof(f7035,plain,(
  ( ! [X629,X627,X628,X630,X626] : (s(bool,'F_0') = s(bool,X629) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X627))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X628))) | s(bool,X629) = s(bool,X630) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X626),X627),s('type.closSem.state'(X626),X628))) | s(bool,'F_0') = s(bool,X630)) )),
  inference(superposition,[],[f2980,f3476])).
fof(f8185,plain,(
  ( ! [X625,X621,X623,X624,X622] : (p(s(bool,X624)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s(bool,X624) = s(bool,X625) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623)))) )),
  inference(subsumption_resolution,[],[f7034,f2915])).
fof(f7034,plain,(
  ( ! [X625,X621,X623,X624,X622] : (p(s(bool,X624)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s(bool,X624) = s(bool,X625) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623))) | s(bool,'F_0') = s(bool,X625)) )),
  inference(superposition,[],[f258,f3476])).
fof(f7033,plain,(
  ( ! [X617,X619,X620,X616,X618] : (~p(s(bool,X619)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,X619) = s(bool,X620) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618))) | s(bool,'F_0') = s(bool,X620)) )),
  inference(superposition,[],[f257,f3476])).
fof(f7032,plain,(
  ( ! [X613,X615,X611,X612,X614] : (~p(s(bool,X614)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,X614) = s(bool,X615) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613))) | s(bool,'F_0') = s(bool,X615)) )),
  inference(superposition,[],[f256,f3476])).
fof(f7031,plain,(
  ( ! [X606,X609,X607,X608,X610] : (~p(s(bool,X609)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,X609) = s(bool,X610) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608))) | s(bool,'F_0') = s(bool,X610)) )),
  inference(superposition,[],[f255,f3476])).
fof(f8181,plain,(
  ( ! [X602,X605,X601,X603,X604] : (~p(s(bool,X604)) | s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X602))) = s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X603))) | s(bool,X604) = s(bool,X605) | s(bool,'F_0') = s(bool,X605)) )),
  inference(subsumption_resolution,[],[f7030,f3144])).
fof(f7030,plain,(
  ( ! [X602,X605,X601,X603,X604] : (~p(s(bool,X604)) | s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X602))) = s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X603))) | s(bool,X604) = s(bool,X605) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X601),X602),s('type.closSem.state'(X601),X603))) | s(bool,'F_0') = s(bool,X605)) )),
  inference(superposition,[],[f254,f3476])).
fof(f8180,plain,(
  ( ! [X600,X596,X598,X597,X599] : (~p(s(bool,X599)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X597))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X598))) | s(bool,X599) = s(bool,X600) | s(bool,'F_0') = s(bool,X600)) )),
  inference(subsumption_resolution,[],[f7029,f2981])).
fof(f7029,plain,(
  ( ! [X600,X596,X598,X597,X599] : (~p(s(bool,X599)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X597))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X598))) | s(bool,X599) = s(bool,X600) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X596),X597),s('type.closSem.state'(X596),X598))) | s(bool,'F_0') = s(bool,X600)) )),
  inference(superposition,[],[f253,f3476])).
fof(f8179,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,X594) = s(bool,X595) | s(bool,'F_0') = s(bool,X595)) )),
  inference(subsumption_resolution,[],[f7028,f1072])).
fof(f7028,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,X594) = s(bool,X595) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593))) | s(bool,'F_0') = s(bool,X595)) )),
  inference(superposition,[],[f1072,f3476])).
fof(f8178,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,X591) = s(bool,X592) | s(bool,'F_0') = s(bool,X592)) )),
  inference(subsumption_resolution,[],[f7027,f968])).
fof(f7027,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,X591) = s(bool,X592) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X590))) | s(bool,'F_0') = s(bool,X592)) )),
  inference(superposition,[],[f968,f3476])).
fof(f8177,plain,(
  ( ! [X588,X589,X587] : (~p(s(bool,X588)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X588) = s(bool,X589) | s(bool,'F_0') = s(bool,X589)) )),
  inference(subsumption_resolution,[],[f7026,f968])).
fof(f7026,plain,(
  ( ! [X588,X589,X587] : (~p(s(bool,X588)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X588) = s(bool,X589) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587))) | s(bool,'F_0') = s(bool,X589)) )),
  inference(superposition,[],[f544,f3476])).
fof(f8173,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,X578) = s(bool,X579) | s(bool,'F_0') = s(bool,X579)) )),
  inference(subsumption_resolution,[],[f7022,f1036])).
fof(f7022,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,X578) = s(bool,X579) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X579)) )),
  inference(superposition,[],[f1036,f3476])).
fof(f8172,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,X575) = s(bool,X576) | s(bool,'F_0') = s(bool,X576)) )),
  inference(subsumption_resolution,[],[f7021,f931])).
fof(f7021,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,X575) = s(bool,X576) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X574),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X576)) )),
  inference(superposition,[],[f931,f3476])).
fof(f8171,plain,(
  ( ! [X573,X571,X572] : (~p(s(bool,X572)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X572) = s(bool,X573) | s(bool,'F_0') = s(bool,X573)) )),
  inference(subsumption_resolution,[],[f7020,f931])).
fof(f7020,plain,(
  ( ! [X573,X571,X572] : (~p(s(bool,X572)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X572) = s(bool,X573) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X573)) )),
  inference(superposition,[],[f543,f3476])).
fof(f5903,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') != s(bool,X22) | s(bool,X21) != s(bool,X22)) ) | $spl124),
  inference(cnf_transformation,[],[f5903_D])).
fof(f5903_D,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') != s(bool,X22) | s(bool,X21) != s(bool,X22)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f8159,plain,(
  ( ! [X453,X451,X452,X454,X450] : (p(s(bool,X451)) | s(bool,X450) != s(bool,X452) | s(bool,X452) != s(bool,X453) | s(bool,'T_0') != s(bool,X453) | s(bool,X451) = s(bool,X454) | s(bool,'T_0') = s(bool,X450)) ) | $spl58),
  inference(subsumption_resolution,[],[f6988,f2915])).
fof(f6988,plain,(
  ( ! [X453,X451,X452,X454,X450] : (p(s(bool,X451)) | s(bool,X450) != s(bool,X452) | s(bool,X452) != s(bool,X453) | s(bool,'T_0') != s(bool,X453) | s(bool,X451) = s(bool,X454) | s(bool,'T_0') = s(bool,X450) | s(bool,'F_0') = s(bool,X454)) ) | $spl58),
  inference(superposition,[],[f4702,f3476])).
fof(f8024,plain,(
  ( ! [X306,X305] : (s(bool,X305) != s(bool,X306) | s(bool,'T_0') = s(bool,X305)) ) | $spl210),
  inference(cnf_transformation,[],[f8024_D])).
fof(f8024_D,plain,(
  ( ! [X306,X305] : (s(bool,X305) != s(bool,X306) | s(bool,'T_0') = s(bool,X305)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f8158,plain,(
  ( ! [X441,X443,X444,X440,X442] : (s(bool,X441) != s(bool,X442) | s(bool,'F_0') = s(bool,X442) | s(bool,X440) != s(bool,X443) | s(bool,X441) = s(bool,X444) | s(bool,'T_0') = s(bool,X440) | s(bool,'F_0') = s(bool,X444)) ) | $spl60),
  inference(subsumption_resolution,[],[f6986,f2882])).
fof(f6986,plain,(
  ( ! [X441,X443,X444,X440,X442] : (s(bool,X441) != s(bool,X442) | s(bool,'F_0') = s(bool,X442) | s(bool,X440) != s(bool,X443) | s(bool,'T_0') = s(bool,X443) | s(bool,X441) = s(bool,X444) | s(bool,'T_0') = s(bool,X440) | s(bool,'F_0') = s(bool,X444)) ) | $spl60),
  inference(superposition,[],[f4623,f3476])).
fof(f8042,plain,(
  ( ! [X397,X399,X401,X400] : (s(bool,X397) != s(bool,X399) | s(bool,'F_0') = s(bool,X401) | s(bool,X397) = s(bool,X401) | s(bool,'T_0') = s(bool,X400) | s(bool,X399) = s(bool,X400)) ) | $spl218),
  inference(cnf_transformation,[],[f8042_D])).
fof(f8042_D,plain,(
  ( ! [X397,X399,X401,X400] : (s(bool,X397) != s(bool,X399) | s(bool,'F_0') = s(bool,X401) | s(bool,X397) = s(bool,X401) | s(bool,'T_0') = s(bool,X400) | s(bool,X399) = s(bool,X400)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7640,plain,(
  ( ! [X430,X432,X434,X431,X433] : (s(bool,X431) != s(bool,X432) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) = s(bool,X433) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X434) | s(bool,'F_0') = s(bool,X434)) )),
  inference(duplicate_literal_removal,[],[f6984])).
fof(f6984,plain,(
  ( ! [X430,X432,X434,X431,X433] : (s(bool,X431) != s(bool,X432) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) = s(bool,X433) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X434) | s(bool,'T_0') = s(bool,X430) | s(bool,'F_0') = s(bool,X434)) )),
  inference(superposition,[],[f4329,f3476])).
fof(f8155,plain,(
  ( ! [X420,X422,X424,X421,X423] : (~p(s(bool,X421)) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X422) | s(bool,X421) = s(bool,X424) | s(bool,'T_0') = s(bool,X420) | s(bool,'F_0') = s(bool,X424)) )),
  inference(subsumption_resolution,[],[f6982,f4329])).
fof(f6982,plain,(
  ( ! [X420,X422,X424,X421,X423] : (~p(s(bool,X421)) | s(bool,X422) = s(bool,X423) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X422) | s(bool,X421) = s(bool,X424) | s(bool,'T_0') = s(bool,X420) | s(bool,'F_0') = s(bool,X424)) )),
  inference(superposition,[],[f3477,f3476])).
fof(f3477,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X11) = s(bool,X12)) )),
  inference(superposition,[],[f3074,f249])).
fof(f8154,plain,(
  ( ! [X418,X419] : (s(bool,X418) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X419) | s(bool,X418) = s(bool,X419)) ) | $spl242),
  inference(cnf_transformation,[],[f8154_D])).
fof(f8154_D,plain,(
  ( ! [X418,X419] : (s(bool,X418) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X419) | s(bool,X418) = s(bool,X419)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f7641,plain,(
  ( ! [X418,X417,X419] : (s(bool,X418) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X417) | s(bool,X418) = s(bool,X419) | s(bool,'F_0') = s(bool,X419)) )),
  inference(duplicate_literal_removal,[],[f6981])).
fof(f6981,plain,(
  ( ! [X418,X417,X419] : (s(bool,X418) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X417) | s(bool,X418) = s(bool,X419) | s(bool,'T_0') = s(bool,X417) | s(bool,'F_0') = s(bool,X419)) )),
  inference(superposition,[],[f3407,f3476])).
fof(f8147,plain,(
  ( ! [X372,X374,X376,X373,X375] : (~p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,'T_0') = s(bool,X375) | s(bool,X373) = s(bool,X376) | s(bool,'T_0') = s(bool,X372) | s(bool,'F_0') = s(bool,X376)) ) | $spl58),
  inference(subsumption_resolution,[],[f6971,f4329])).
fof(f6971,plain,(
  ( ! [X372,X374,X376,X373,X375] : (~p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') = s(bool,X375) | s(bool,X373) = s(bool,X376) | s(bool,'T_0') = s(bool,X372) | s(bool,'F_0') = s(bool,X376)) ) | $spl58),
  inference(superposition,[],[f3037,f3476])).
fof(f8145,plain,(
  ( ! [X368,X370,X367,X369,X371] : (~p(s(bool,X368)) | s(bool,X367) != s(bool,X369) | s(bool,X369) != s(bool,X370) | p(s(bool,X370)) | s(bool,X368) = s(bool,X371) | s(bool,'F_0') = s(bool,X371)) ) | $spl58),
  inference(subsumption_resolution,[],[f6970,f4702])).
fof(f6970,plain,(
  ( ! [X368,X370,X367,X369,X371] : (~p(s(bool,X368)) | s(bool,X367) != s(bool,X369) | s(bool,X369) != s(bool,X370) | p(s(bool,X370)) | s(bool,X368) = s(bool,X371) | s(bool,'T_0') = s(bool,X367) | s(bool,'F_0') = s(bool,X371)) ) | $spl58),
  inference(superposition,[],[f3036,f3476])).
fof(f8144,plain,(
  ( ! [X358,X360,X357,X359] : (s(bool,X357) != s(bool,X358) | s(bool,'F_0') = s(bool,X360) | s(bool,X357) = s(bool,X360) | s(bool,X358) = s(bool,X359) | p(s(bool,X359))) ) | $spl238),
  inference(cnf_transformation,[],[f8144_D])).
fof(f8144_D,plain,(
  ( ! [X358,X360,X357,X359] : (s(bool,X357) != s(bool,X358) | s(bool,'F_0') = s(bool,X360) | s(bool,X357) = s(bool,X360) | s(bool,X358) = s(bool,X359) | p(s(bool,X359))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f8143,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,X357) != s(bool,X358) | p(s(bool,X359)) | p(s(bool,X356)) | s(bool,X358) = s(bool,X359) | s(bool,X357) = s(bool,X360) | s(bool,'F_0') = s(bool,X360)) ) | $spl62),
  inference(subsumption_resolution,[],[f6968,f246])).
fof(f6968,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,X357) != s(bool,X358) | p(s(bool,X359)) | p(s(bool,X356)) | s(bool,X358) = s(bool,X359) | s(bool,X357) = s(bool,X360) | s(bool,'T_0') = s(bool,X356) | s(bool,'F_0') = s(bool,X360)) ) | $spl62),
  inference(superposition,[],[f2964,f3476])).
fof(f8140,plain,(
  ( ! [X344,X346,X343,X345] : (s(bool,'T_0') != s(bool,X344) | s(bool,X343) != s(bool,X345) | s(bool,X344) = s(bool,X346) | s(bool,'T_0') = s(bool,X343) | s(bool,'F_0') = s(bool,X346)) )),
  inference(subsumption_resolution,[],[f6965,f2882])).
fof(f6965,plain,(
  ( ! [X344,X346,X343,X345] : (s(bool,'T_0') != s(bool,X344) | s(bool,X343) != s(bool,X345) | s(bool,'T_0') = s(bool,X345) | s(bool,X344) = s(bool,X346) | s(bool,'T_0') = s(bool,X343) | s(bool,'F_0') = s(bool,X346)) )),
  inference(superposition,[],[f2882,f3476])).
fof(f8139,plain,(
  ( ! [X341,X339,X340,X342] : (s(bool,'T_0') != s(bool,X340) | s(bool,X339) != s(bool,X341) | p(s(bool,X341)) | s(bool,X340) = s(bool,X342) | s(bool,'F_0') = s(bool,X342)) )),
  inference(subsumption_resolution,[],[f6964,f2881])).
fof(f6964,plain,(
  ( ! [X341,X339,X340,X342] : (s(bool,'T_0') != s(bool,X340) | s(bool,X339) != s(bool,X341) | p(s(bool,X341)) | s(bool,X340) = s(bool,X342) | s(bool,'T_0') = s(bool,X339) | s(bool,'F_0') = s(bool,X342)) )),
  inference(superposition,[],[f2881,f3476])).
fof(f7644,plain,(
  ( ! [X313,X315,X312,X314] : (s(bool,X313) != s(bool,X314) | s(bool,'T_0') = s(bool,X312) | s(bool,'F_0') = s(bool,X314) | s(bool,X313) = s(bool,X315) | s(bool,'F_0') = s(bool,X315)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6957])).
fof(f6957,plain,(
  ( ! [X313,X315,X312,X314] : (s(bool,X313) != s(bool,X314) | s(bool,'T_0') = s(bool,X312) | s(bool,'F_0') = s(bool,X314) | s(bool,X313) = s(bool,X315) | s(bool,'T_0') = s(bool,X312) | s(bool,'F_0') = s(bool,X315)) ) | $spl6),
  inference(superposition,[],[f828,f3476])).
fof(f7645,plain,(
  ( ! [X300,X302,X301,X303] : (s(bool,X301) != s(bool,X302) | ~p(s(bool,X302)) | s(bool,'T_0') = s(bool,X300) | s(bool,X301) = s(bool,X303) | s(bool,'F_0') = s(bool,X303)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6954])).
fof(f6954,plain,(
  ( ! [X300,X302,X301,X303] : (s(bool,X301) != s(bool,X302) | ~p(s(bool,X302)) | s(bool,'T_0') = s(bool,X300) | s(bool,X301) = s(bool,X303) | s(bool,'T_0') = s(bool,X300) | s(bool,'F_0') = s(bool,X303)) ) | $spl6),
  inference(superposition,[],[f766,f3476])).
fof(f8130,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X293) != s(bool,X294) | p(s(bool,X292)) | ~p(s(bool,X294)) | s(bool,X293) = s(bool,X295) | s(bool,'F_0') = s(bool,X295)) ) | $spl6),
  inference(subsumption_resolution,[],[f6952,f246])).
fof(f6952,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X293) != s(bool,X294) | p(s(bool,X292)) | ~p(s(bool,X294)) | s(bool,X293) = s(bool,X295) | s(bool,'T_0') = s(bool,X292) | s(bool,'F_0') = s(bool,X295)) ) | $spl6),
  inference(superposition,[],[f297,f3476])).
fof(f8129,plain,(
  ( ! [X288,X290,X287,X289,X291] : (~p(s(bool,X288)) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X287) != s(bool,X290) | s(bool,X288) = s(bool,X291) | s(bool,'F_0') = s(bool,X291)) )),
  inference(subsumption_resolution,[],[f6951,f2881])).
fof(f6951,plain,(
  ( ! [X288,X290,X287,X289,X291] : (~p(s(bool,X288)) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X287) != s(bool,X290) | s(bool,X288) = s(bool,X291) | s(bool,'T_0') = s(bool,X287) | s(bool,'F_0') = s(bool,X291)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8126,plain,(
  ( ! [X273,X275,X276,X272,X274] : (s(bool,X273) != s(bool,X274) | p(s(bool,X275)) | ~p(s(bool,X274)) | p(s(bool,X272)) | s(bool,X273) = s(bool,X276) | s(bool,'F_0') = s(bool,X276)) )),
  inference(subsumption_resolution,[],[f6948,f246])).
fof(f6948,plain,(
  ( ! [X273,X275,X276,X272,X274] : (s(bool,X273) != s(bool,X274) | p(s(bool,X275)) | ~p(s(bool,X274)) | p(s(bool,X272)) | s(bool,X273) = s(bool,X276) | s(bool,'T_0') = s(bool,X272) | s(bool,'F_0') = s(bool,X276)) )),
  inference(superposition,[],[f262,f3476])).
fof(f7646,plain,(
  ( ! [X265,X264,X266] : (s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264) | s(bool,X265) = s(bool,X266) | s(bool,'F_0') = s(bool,X266)) )),
  inference(duplicate_literal_removal,[],[f6946])).
fof(f6946,plain,(
  ( ! [X265,X264,X266] : (s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264) | s(bool,X265) = s(bool,X266) | s(bool,'T_0') = s(bool,X264) | s(bool,'F_0') = s(bool,X266)) )),
  inference(superposition,[],[f261,f3476])).
fof(f7647,plain,(
  ( ! [X259,X260,X258] : (~p(s(bool,X259)) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260)) )),
  inference(duplicate_literal_removal,[],[f6944])).
fof(f6944,plain,(
  ( ! [X259,X260,X258] : (~p(s(bool,X259)) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X258) | s(bool,'F_0') = s(bool,X260)) )),
  inference(superposition,[],[f249,f3476])).
fof(f8122,plain,(
  ( ! [X251,X249,X250] : (s(bool,'T_0') != s(bool,X250) | p(s(bool,X249)) | s(bool,X250) = s(bool,X251) | s(bool,'F_0') = s(bool,X251)) )),
  inference(subsumption_resolution,[],[f6941,f246])).
fof(f6941,plain,(
  ( ! [X251,X249,X250] : (s(bool,'T_0') != s(bool,X250) | p(s(bool,X249)) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X249) | s(bool,'F_0') = s(bool,X251)) )),
  inference(superposition,[],[f248,f3476])).
fof(f7648,plain,(
  ( ! [X246,X248,X247] : (~p(s(bool,X247)) | s(bool,'T_0') = s(bool,X246) | s(bool,X247) = s(bool,X248) | s(bool,'F_0') = s(bool,X248)) )),
  inference(duplicate_literal_removal,[],[f6940])).
fof(f6940,plain,(
  ( ! [X246,X248,X247] : (~p(s(bool,X247)) | s(bool,'T_0') = s(bool,X246) | s(bool,X247) = s(bool,X248) | s(bool,'T_0') = s(bool,X246) | s(bool,'F_0') = s(bool,X248)) )),
  inference(superposition,[],[f247,f3476])).
fof(f8120,plain,(
  ( ! [X239,X237,X238] : (s(bool,'T_0') != s(bool,X238) | p(s(bool,X237)) | s(bool,X238) = s(bool,X239) | s(bool,'F_0') = s(bool,X239)) )),
  inference(subsumption_resolution,[],[f6937,f246])).
fof(f6937,plain,(
  ( ! [X239,X237,X238] : (s(bool,'T_0') != s(bool,X238) | p(s(bool,X237)) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X237) | s(bool,'F_0') = s(bool,X239)) )),
  inference(superposition,[],[f246,f3476])).
fof(f8118,plain,(
  ( ! [X231,X233,X232] : (s(bool,'F_0') = s(bool,X232) | p(s(bool,X231)) | s(bool,X232) = s(bool,X233) | s(bool,'F_0') = s(bool,X233)) )),
  inference(subsumption_resolution,[],[f6935,f246])).
fof(f6935,plain,(
  ( ! [X231,X233,X232] : (s(bool,'F_0') = s(bool,X232) | p(s(bool,X231)) | s(bool,X232) = s(bool,X233) | s(bool,'T_0') = s(bool,X231) | s(bool,'F_0') = s(bool,X233)) )),
  inference(superposition,[],[f243,f3476])).
fof(f8117,plain,(
  ( ! [X229,X230,X228] : (~p(s(bool,X229)) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f6934,f261])).
fof(f6934,plain,(
  ( ! [X229,X230,X228] : (~p(s(bool,X229)) | s(bool,'F_0') != s(bool,X228) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230)) )),
  inference(superposition,[],[f242,f3476])).
fof(f8115,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X198) != s(bool,X200) | s(bool,X201) != s(bool,X202) | s(bool,X199) = s(bool,X203) | s(bool,'F_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f6925,f2881])).
fof(f6925,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X198) != s(bool,X200) | s(bool,X201) != s(bool,X202) | s(bool,X199) = s(bool,X203) | s(bool,'T_0') = s(bool,X198) | s(bool,'F_0') = s(bool,X203)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8111,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,X175) != s(bool,X176) | p(s(bool,X174)) | ~p(s(bool,X177)) | ~p(s(bool,X176)) | s(bool,X177) != s(bool,X178) | s(bool,X175) = s(bool,X179) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f6921,f246])).
fof(f6921,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,X175) != s(bool,X176) | p(s(bool,X174)) | ~p(s(bool,X177)) | ~p(s(bool,X176)) | s(bool,X177) != s(bool,X178) | s(bool,X175) = s(bool,X179) | s(bool,'T_0') = s(bool,X174) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8108,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,X151)) | p(s(bool,X152)) | ~p(s(bool,X153)) | p(s(bool,X154)) | s(bool,X150) != s(bool,X154) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f6917,f2881])).
fof(f6917,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,X151)) | p(s(bool,X152)) | ~p(s(bool,X153)) | p(s(bool,X154)) | s(bool,X150) != s(bool,X154) | s(bool,X151) = s(bool,X155) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8105,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,X133) != s(bool,X134) | p(s(bool,X135)) | ~p(s(bool,X134)) | ~p(s(bool,X136)) | p(s(bool,X132)) | s(bool,X133) = s(bool,X137) | s(bool,'F_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f6914,f246])).
fof(f6914,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,X133) != s(bool,X134) | p(s(bool,X135)) | ~p(s(bool,X134)) | ~p(s(bool,X136)) | p(s(bool,X132)) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X132) | s(bool,'F_0') = s(bool,X137)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8103,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X122) != s(bool,X124) | s(bool,X120) != s(bool,X123) | s(bool,X121) = s(bool,X125) | s(bool,'F_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f6912,f2881])).
fof(f6912,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X122) != s(bool,X124) | s(bool,X120) != s(bool,X123) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X125)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8098,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f6907,f246])).
fof(f6907,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'T_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f201,f3476])).
fof(f8096,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X78) != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83) | s(bool,'F_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f6905,f2881])).
fof(f6905,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X78) != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X83)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8093,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X60)) | ~p(s(bool,X63)) | ~p(s(bool,X62)) | s(bool,X63) != s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f6902,f246])).
fof(f6902,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X60)) | ~p(s(bool,X63)) | ~p(s(bool,X62)) | s(bool,X63) != s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f199,f3476])).
fof(f8089,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,X37)) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X36) != s(bool,X39) | s(bool,X37) = s(bool,X41) | s(bool,'F_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f6898,f2881])).
fof(f6898,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,X37)) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X36) != s(bool,X39) | s(bool,X37) = s(bool,X41) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8083,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,X1) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6892,f246])).
fof(f6892,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,X1) = s(bool,X5) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f197,f3476])).
fof(f8082,plain,(
  ( ! [X685,X687,X683,X684] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,'F_0') = s(bool,X687) | s(bool,X687) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685)))))) ) | $spl236),
  inference(cnf_transformation,[],[f8082_D])).
fof(f8082_D,plain,(
  ( ! [X685,X687,X683,X684] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,'F_0') = s(bool,X687) | s(bool,X687) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685)))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f8081,plain,(
  ( ! [X685,X687,X683,X684,X686] : (p(s(bool,X686)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,X687) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685))))) | s(bool,'F_0') = s(bool,X687)) )),
  inference(subsumption_resolution,[],[f6888,f246])).
fof(f6888,plain,(
  ( ! [X685,X687,X683,X684,X686] : (p(s(bool,X686)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X683),X684),s('type.closSem.state'(X683),X685)))) | s(bool,X687) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X684))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X683),X685))))) | s(bool,'T_0') = s(bool,X686) | s(bool,'F_0') = s(bool,X687)) )),
  inference(superposition,[],[f256,f3476])).
fof(f8080,plain,(
  ( ! [X677,X673,X675,X674] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,'F_0') = s(bool,X677) | s(bool,X677) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675)))))) ) | $spl234),
  inference(cnf_transformation,[],[f8080_D])).
fof(f8080_D,plain,(
  ( ! [X677,X673,X675,X674] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,'F_0') = s(bool,X677) | s(bool,X677) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f8079,plain,(
  ( ! [X677,X673,X675,X676,X674] : (p(s(bool,X676)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,X677) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675))))) | s(bool,'F_0') = s(bool,X677)) )),
  inference(subsumption_resolution,[],[f6886,f246])).
fof(f6886,plain,(
  ( ! [X677,X673,X675,X676,X674] : (p(s(bool,X676)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X673),X674),s('type.closSem.state'(X673),X675)))) | s(bool,X677) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X674))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X673),X675))))) | s(bool,'T_0') = s(bool,X676) | s(bool,'F_0') = s(bool,X677)) )),
  inference(superposition,[],[f255,f3476])).
fof(f8078,plain,(
  ( ! [X664,X667,X663,X665] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,'F_0') = s(bool,X667) | s(bool,X667) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665)))))) ) | $spl232),
  inference(cnf_transformation,[],[f8078_D])).
fof(f8078_D,plain,(
  ( ! [X664,X667,X663,X665] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,'F_0') = s(bool,X667) | s(bool,X667) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665)))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f8077,plain,(
  ( ! [X666,X664,X667,X663,X665] : (p(s(bool,X666)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,X667) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665))))) | s(bool,'F_0') = s(bool,X667)) )),
  inference(subsumption_resolution,[],[f6884,f246])).
fof(f6884,plain,(
  ( ! [X666,X664,X667,X663,X665] : (p(s(bool,X666)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X663),X664),s('type.closSem.state'(X663),X665)))) | s(bool,X667) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X664))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X663),X665))))) | s(bool,'T_0') = s(bool,X666) | s(bool,'F_0') = s(bool,X667)) )),
  inference(superposition,[],[f257,f3476])).
fof(f8076,plain,(
  ( ! [X625,X621,X623,X622] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | s(bool,'F_0') = s(bool,X625) | s(bool,X625) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623))))))) ) | $spl230),
  inference(cnf_transformation,[],[f8076_D])).
fof(f8076_D,plain,(
  ( ! [X625,X621,X623,X622] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | s(bool,'F_0') = s(bool,X625) | s(bool,X625) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623))))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f8075,plain,(
  ( ! [X625,X621,X623,X624,X622] : (p(s(bool,X624)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s(bool,X625) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623))) | s(bool,'F_0') = s(bool,X625)) )),
  inference(subsumption_resolution,[],[f6876,f246])).
fof(f6876,plain,(
  ( ! [X625,X621,X623,X624,X622] : (p(s(bool,X624)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X622))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X621),X623)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X622))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X621),X623)))))) | s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X622))) != s('type.ffi.ffi_state'(X621),'const.closSem.state_ffi_1'(s('type.closSem.state'(X621),X623))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X622))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X621),X623))) | s(bool,X625) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X621),X622),s('type.closSem.state'(X621),X623))) | s(bool,'T_0') = s(bool,X624) | s(bool,'F_0') = s(bool,X625)) )),
  inference(superposition,[],[f258,f3476])).
fof(f7649,plain,(
  ( ! [X586,X585] : (s(bool,'T_0') = s(bool,X585) | s(bool,X586) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X586)) )),
  inference(duplicate_literal_removal,[],[f6867])).
fof(f6867,plain,(
  ( ! [X586,X585] : (s(bool,'T_0') = s(bool,X585) | s(bool,X586) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X585) | s(bool,'F_0') = s(bool,X586)) )),
  inference(superposition,[],[f285,f3476])).
fof(f8074,plain,(
  ( ! [X584,X582,X583] : (s(bool,X582) != s(bool,X583) | s(bool,X584) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X582) | s(bool,'F_0') = s(bool,X584)) )),
  inference(subsumption_resolution,[],[f6866,f2882])).
fof(f6866,plain,(
  ( ! [X584,X582,X583] : (s(bool,X582) != s(bool,X583) | s(bool,'T_0') = s(bool,X583) | s(bool,X584) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X582) | s(bool,'F_0') = s(bool,X584)) )),
  inference(superposition,[],[f3407,f3476])).
fof(f7650,plain,(
  ( ! [X569,X567,X568,X570] : (s(bool,X567) != s(bool,X568) | s(bool,X568) != s(bool,X569) | s(bool,'T_0') != s(bool,X569) | s(bool,'F_0') = s(bool,X570) | s(bool,'T_0') = s(bool,X567)) ) | $spl58),
  inference(duplicate_literal_removal,[],[f6861])).
fof(f6861,plain,(
  ( ! [X569,X567,X568,X570] : (s(bool,X567) != s(bool,X568) | s(bool,X568) != s(bool,X569) | s(bool,'T_0') != s(bool,X569) | s(bool,'F_0') = s(bool,X570) | s(bool,'T_0') = s(bool,X567) | s(bool,'F_0') = s(bool,X570)) ) | $spl58),
  inference(superposition,[],[f5870,f3476])).
fof(f8056,plain,(
  ( ! [X455,X457,X458] : (s(bool,X455) != s(bool,X457) | s(bool,'F_0') = s(bool,X458) | s(bool,X455) = s(bool,X458) | s(bool,'T_0') != s(bool,X457)) ) | $spl228),
  inference(cnf_transformation,[],[f8056_D])).
fof(f8056_D,plain,(
  ( ! [X455,X457,X458] : (s(bool,X455) != s(bool,X457) | s(bool,'F_0') = s(bool,X458) | s(bool,X455) = s(bool,X458) | s(bool,'T_0') != s(bool,X457)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8055,plain,(
  ( ! [X455,X457,X456,X458] : (s(bool,'F_0') != s(bool,X456) | s(bool,X455) != s(bool,X457) | s(bool,'T_0') != s(bool,X457) | s(bool,X455) = s(bool,X458) | s(bool,'F_0') = s(bool,X458)) ) | $spl58),
  inference(subsumption_resolution,[],[f6831,f2590])).
fof(f6831,plain,(
  ( ! [X455,X457,X456,X458] : (s(bool,'F_0') != s(bool,X456) | s(bool,X455) != s(bool,X457) | s(bool,'T_0') != s(bool,X457) | s(bool,X455) = s(bool,X458) | s(bool,'T_0') = s(bool,X456) | s(bool,'F_0') = s(bool,X458)) ) | $spl58),
  inference(superposition,[],[f5870,f3476])).
fof(f8054,plain,(
  ( ! [X453,X452,X454,X450] : (s(bool,X450) != s(bool,X452) | s(bool,'F_0') = s(bool,X454) | s(bool,X450) = s(bool,X454) | s(bool,'T_0') != s(bool,X453) | s(bool,X452) != s(bool,X453)) ) | $spl226),
  inference(cnf_transformation,[],[f8054_D])).
fof(f8054_D,plain,(
  ( ! [X453,X452,X454,X450] : (s(bool,X450) != s(bool,X452) | s(bool,'F_0') = s(bool,X454) | s(bool,X450) = s(bool,X454) | s(bool,'T_0') != s(bool,X453) | s(bool,X452) != s(bool,X453)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8053,plain,(
  ( ! [X453,X451,X452,X454,X450] : (p(s(bool,X451)) | s(bool,X450) != s(bool,X452) | s(bool,X452) != s(bool,X453) | s(bool,'T_0') != s(bool,X453) | s(bool,X450) = s(bool,X454) | s(bool,'F_0') = s(bool,X454)) ) | $spl58),
  inference(subsumption_resolution,[],[f6830,f246])).
fof(f6830,plain,(
  ( ! [X453,X451,X452,X454,X450] : (p(s(bool,X451)) | s(bool,X450) != s(bool,X452) | s(bool,X452) != s(bool,X453) | s(bool,'T_0') != s(bool,X453) | s(bool,X450) = s(bool,X454) | s(bool,'T_0') = s(bool,X451) | s(bool,'F_0') = s(bool,X454)) ) | $spl58),
  inference(superposition,[],[f4702,f3476])).
fof(f8052,plain,(
  ( ! [X445,X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'F_0') = s(bool,X445) | s(bool,'T_0') = s(bool,X448) | s(bool,X445) = s(bool,X449) | s(bool,'T_0') = s(bool,X446) | s(bool,'F_0') = s(bool,X449)) ) | $spl60),
  inference(subsumption_resolution,[],[f6829,f4329])).
fof(f6829,plain,(
  ( ! [X445,X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'F_0') = s(bool,X445) | s(bool,X447) != s(bool,X448) | s(bool,'T_0') = s(bool,X448) | s(bool,X445) = s(bool,X449) | s(bool,'T_0') = s(bool,X446) | s(bool,'F_0') = s(bool,X449)) ) | $spl60),
  inference(superposition,[],[f4623,f3476])).
fof(f8048,plain,(
  ( ! [X428,X426,X429,X425,X427] : (p(s(bool,X426)) | s(bool,X425) = s(bool,X427) | s(bool,X427) != s(bool,X428) | s(bool,'T_0') = s(bool,X428) | s(bool,X425) = s(bool,X429) | s(bool,'F_0') = s(bool,X429)) )),
  inference(subsumption_resolution,[],[f6825,f246])).
fof(f6825,plain,(
  ( ! [X428,X426,X429,X425,X427] : (p(s(bool,X426)) | s(bool,X425) = s(bool,X427) | s(bool,X427) != s(bool,X428) | s(bool,'T_0') = s(bool,X428) | s(bool,X425) = s(bool,X429) | s(bool,'T_0') = s(bool,X426) | s(bool,'F_0') = s(bool,X429)) )),
  inference(superposition,[],[f4328,f3476])).
fof(f8046,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'T_0') = s(bool,X412) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) | $spl222),
  inference(cnf_transformation,[],[f8046_D])).
fof(f8046_D,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'T_0') = s(bool,X412) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f6822,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415))) | ~p(s(bool,X411)) | s(bool,X411) = s(bool,X416) | s(bool,'T_0') = s(bool,X412) | s(bool,'F_0') = s(bool,X416)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f8045,plain,(
  ( ! [X408,X410,X409,X406] : (s(bool,X406) != s(bool,X408) | s(bool,'F_0') = s(bool,X410) | s(bool,X406) = s(bool,X410) | s(bool,'F_0') = s(bool,X409) | s(bool,X408) != s(bool,X409)) ) | $spl220),
  inference(cnf_transformation,[],[f8045_D])).
fof(f8045_D,plain,(
  ( ! [X408,X410,X409,X406] : (s(bool,X406) != s(bool,X408) | s(bool,'F_0') = s(bool,X410) | s(bool,X406) = s(bool,X410) | s(bool,'F_0') = s(bool,X409) | s(bool,X408) != s(bool,X409)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f8044,plain,(
  ( ! [X408,X410,X407,X409,X406] : (p(s(bool,X407)) | s(bool,X406) != s(bool,X408) | s(bool,X408) != s(bool,X409) | s(bool,'F_0') = s(bool,X409) | s(bool,X406) = s(bool,X410) | s(bool,'F_0') = s(bool,X410)) ) | $spl68),
  inference(subsumption_resolution,[],[f6821,f246])).
fof(f6821,plain,(
  ( ! [X408,X410,X407,X409,X406] : (p(s(bool,X407)) | s(bool,X406) != s(bool,X408) | s(bool,X408) != s(bool,X409) | s(bool,'F_0') = s(bool,X409) | s(bool,X406) = s(bool,X410) | s(bool,'T_0') = s(bool,X407) | s(bool,'F_0') = s(bool,X410)) ) | $spl68),
  inference(superposition,[],[f3130,f3476])).
fof(f8043,plain,(
  ( ! [X405,X403,X404,X402] : (s(bool,'F_0') = s(bool,X404) | s(bool,X402) != s(bool,X404) | s(bool,X402) = s(bool,X405) | s(bool,'T_0') = s(bool,X403) | s(bool,'F_0') = s(bool,X405)) ) | $spl60),
  inference(subsumption_resolution,[],[f6820,f261])).
fof(f6820,plain,(
  ( ! [X405,X403,X404,X402] : (s(bool,'F_0') != s(bool,X403) | s(bool,'F_0') = s(bool,X404) | s(bool,X402) != s(bool,X404) | s(bool,X402) = s(bool,X405) | s(bool,'T_0') = s(bool,X403) | s(bool,'F_0') = s(bool,X405)) ) | $spl60),
  inference(superposition,[],[f3111,f3476])).
fof(f8041,plain,(
  ( ! [X397,X399,X401,X398,X400] : (p(s(bool,X398)) | s(bool,X397) != s(bool,X399) | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400) | s(bool,X397) = s(bool,X401) | s(bool,'F_0') = s(bool,X401)) ) | $spl66),
  inference(subsumption_resolution,[],[f6819,f246])).
fof(f6819,plain,(
  ( ! [X397,X399,X401,X398,X400] : (p(s(bool,X398)) | s(bool,X397) != s(bool,X399) | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400) | s(bool,X397) = s(bool,X401) | s(bool,'T_0') = s(bool,X398) | s(bool,'F_0') = s(bool,X401)) ) | $spl66),
  inference(superposition,[],[f3095,f3476])).
fof(f8037,plain,(
  ( ! [X364,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'T_0') = s(bool,X362) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) | $spl216),
  inference(cnf_transformation,[],[f8037_D])).
fof(f8037_D,plain,(
  ( ! [X364,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'T_0') = s(bool,X362) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f6811,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365))) | ~p(s(bool,X361)) | s(bool,X361) = s(bool,X366) | s(bool,'T_0') = s(bool,X362) | s(bool,'F_0') = s(bool,X366)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f8035,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s(bool,X352) != s(bool,X353) | p(s(bool,X354)) | p(s(bool,X353)) | s(bool,X351) = s(bool,X354) | s(bool,X351) = s(bool,X355) | s(bool,'F_0') = s(bool,X355)) ) | $spl62),
  inference(subsumption_resolution,[],[f6809,f2881])).
fof(f6809,plain,(
  ( ! [X352,X354,X351,X353,X355] : (s(bool,X352) != s(bool,X353) | p(s(bool,X354)) | p(s(bool,X353)) | s(bool,X351) = s(bool,X354) | s(bool,X351) = s(bool,X355) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,X355)) ) | $spl62),
  inference(superposition,[],[f2964,f3476])).
fof(f8034,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,X347) = s(bool,X349) | p(s(bool,X349)) | s(bool,X347) = s(bool,X350) | s(bool,'T_0') = s(bool,X348) | s(bool,'F_0') = s(bool,X350)) )),
  inference(subsumption_resolution,[],[f6808,f261])).
fof(f6808,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,'F_0') != s(bool,X348) | s(bool,X347) = s(bool,X349) | p(s(bool,X349)) | s(bool,X347) = s(bool,X350) | s(bool,'T_0') = s(bool,X348) | s(bool,'F_0') = s(bool,X350)) )),
  inference(superposition,[],[f2915,f3476])).
fof(f8032,plain,(
  ( ! [X335,X337,X336,X338] : (p(s(bool,X336)) | s(bool,X335) != s(bool,X337) | s(bool,'F_0') = s(bool,X337) | s(bool,X335) = s(bool,X338) | s(bool,'F_0') = s(bool,X338)) ) | $spl60),
  inference(subsumption_resolution,[],[f6805,f246])).
fof(f6805,plain,(
  ( ! [X335,X337,X336,X338] : (p(s(bool,X336)) | s(bool,X335) != s(bool,X337) | s(bool,'F_0') = s(bool,X337) | s(bool,X335) = s(bool,X338) | s(bool,'T_0') = s(bool,X336) | s(bool,'F_0') = s(bool,X338)) ) | $spl60),
  inference(superposition,[],[f2856,f3476])).
fof(f8030,plain,(
  ( ! [X327,X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X327) = s(bool,X329) | s(bool,X327) = s(bool,X330) | s(bool,'F_0') = s(bool,X330)) )),
  inference(subsumption_resolution,[],[f6803,f246])).
fof(f6803,plain,(
  ( ! [X327,X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X327) = s(bool,X329) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X328) | s(bool,'F_0') = s(bool,X330)) )),
  inference(superposition,[],[f2659,f3476])).
fof(f8028,plain,(
  ( ! [X325,X323,X324,X326] : (p(s(bool,X324)) | p(s(bool,X325)) | s(bool,X323) = s(bool,X325) | s(bool,X323) = s(bool,X326) | s(bool,'F_0') = s(bool,X326)) )),
  inference(subsumption_resolution,[],[f6802,f246])).
fof(f6802,plain,(
  ( ! [X325,X323,X324,X326] : (p(s(bool,X324)) | p(s(bool,X325)) | s(bool,X323) = s(bool,X325) | s(bool,X323) = s(bool,X326) | s(bool,'T_0') = s(bool,X324) | s(bool,'F_0') = s(bool,X326)) )),
  inference(superposition,[],[f2659,f3476])).
fof(f8027,plain,(
  ( ! [X321,X320,X322,X319] : (s(bool,X319) != s(bool,X321) | ~p(s(bool,X321)) | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X320) | s(bool,'F_0') = s(bool,X322)) ) | $spl58),
  inference(subsumption_resolution,[],[f6801,f261])).
fof(f6801,plain,(
  ( ! [X321,X320,X322,X319] : (s(bool,'F_0') != s(bool,X320) | s(bool,X319) != s(bool,X321) | ~p(s(bool,X321)) | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X320) | s(bool,'F_0') = s(bool,X322)) ) | $spl58),
  inference(superposition,[],[f2634,f3476])).
fof(f8026,plain,(
  ( ! [X316,X318,X317] : (s(bool,'F_0') != s(bool,X317) | s(bool,'T_0') != s(bool,X316) | s(bool,X316) = s(bool,X318) | s(bool,'F_0') = s(bool,X318)) )),
  inference(subsumption_resolution,[],[f6800,f2590])).
fof(f6800,plain,(
  ( ! [X316,X318,X317] : (s(bool,'F_0') != s(bool,X317) | s(bool,'T_0') != s(bool,X316) | s(bool,X316) = s(bool,X318) | s(bool,'T_0') = s(bool,X317) | s(bool,'F_0') = s(bool,X318)) )),
  inference(superposition,[],[f2590,f3476])).
fof(f8025,plain,(
  ( ! [X308,X310,X309,X311] : (s(bool,X309) != s(bool,X310) | s(bool,'F_0') = s(bool,X308) | s(bool,X308) = s(bool,X311) | s(bool,'T_0') = s(bool,X309) | s(bool,'F_0') = s(bool,X311)) ) | $spl6),
  inference(subsumption_resolution,[],[f6798,f2882])).
fof(f6798,plain,(
  ( ! [X308,X310,X309,X311] : (s(bool,X309) != s(bool,X310) | s(bool,'T_0') = s(bool,X310) | s(bool,'F_0') = s(bool,X308) | s(bool,X308) = s(bool,X311) | s(bool,'T_0') = s(bool,X309) | s(bool,'F_0') = s(bool,X311)) ) | $spl6),
  inference(superposition,[],[f828,f3476])).
fof(f8023,plain,(
  ( ! [X304,X306,X305,X307] : (s(bool,X305) != s(bool,X306) | ~p(s(bool,X304)) | s(bool,X304) = s(bool,X307) | s(bool,'T_0') = s(bool,X305) | s(bool,'F_0') = s(bool,X307)) ) | $spl6),
  inference(subsumption_resolution,[],[f6797,f2882])).
fof(f6797,plain,(
  ( ! [X304,X306,X305,X307] : (s(bool,X305) != s(bool,X306) | ~p(s(bool,X304)) | s(bool,'T_0') = s(bool,X306) | s(bool,X304) = s(bool,X307) | s(bool,'T_0') = s(bool,X305) | s(bool,'F_0') = s(bool,X307)) ) | $spl6),
  inference(superposition,[],[f766,f3476])).
fof(f8022,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,X297) != s(bool,X298) | p(s(bool,X298)) | ~p(s(bool,X296)) | s(bool,X296) = s(bool,X299) | s(bool,'F_0') = s(bool,X299)) ) | $spl6),
  inference(subsumption_resolution,[],[f6795,f2881])).
fof(f6795,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,X297) != s(bool,X298) | p(s(bool,X298)) | ~p(s(bool,X296)) | s(bool,X296) = s(bool,X299) | s(bool,'T_0') = s(bool,X297) | s(bool,'F_0') = s(bool,X299)) ) | $spl6),
  inference(superposition,[],[f297,f3476])).
fof(f8019,plain,(
  ( ! [X280,X277,X279,X281,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | ~p(s(bool,X280)) | s(bool,X277) != s(bool,X280) | s(bool,X277) = s(bool,X281) | s(bool,'F_0') = s(bool,X281)) )),
  inference(subsumption_resolution,[],[f6791,f246])).
fof(f6791,plain,(
  ( ! [X280,X277,X279,X281,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | ~p(s(bool,X280)) | s(bool,X277) != s(bool,X280) | s(bool,X277) = s(bool,X281) | s(bool,'T_0') = s(bool,X278) | s(bool,'F_0') = s(bool,X281)) )),
  inference(superposition,[],[f262,f3476])).
fof(f8017,plain,(
  ( ! [X269,X271,X267,X268,X270] : (s(bool,X268) != s(bool,X269) | p(s(bool,X270)) | ~p(s(bool,X267)) | p(s(bool,X269)) | s(bool,X267) = s(bool,X271) | s(bool,'F_0') = s(bool,X271)) )),
  inference(subsumption_resolution,[],[f6789,f2881])).
fof(f6789,plain,(
  ( ! [X269,X271,X267,X268,X270] : (s(bool,X268) != s(bool,X269) | p(s(bool,X270)) | ~p(s(bool,X267)) | p(s(bool,X269)) | s(bool,X267) = s(bool,X271) | s(bool,'T_0') = s(bool,X268) | s(bool,'F_0') = s(bool,X271)) )),
  inference(superposition,[],[f262,f3476])).
fof(f7663,plain,(
  ( ! [X261,X263,X262] : (s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X261) | s(bool,X261) = s(bool,X263) | s(bool,'F_0') = s(bool,X263)) )),
  inference(duplicate_literal_removal,[],[f6787])).
fof(f6787,plain,(
  ( ! [X261,X263,X262] : (s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X261) | s(bool,X261) = s(bool,X263) | s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X263)) )),
  inference(superposition,[],[f261,f3476])).
fof(f7664,plain,(
  ( ! [X255,X257,X256] : (s(bool,'T_0') = s(bool,X256) | ~p(s(bool,X255)) | s(bool,X255) = s(bool,X257) | s(bool,'F_0') = s(bool,X257)) )),
  inference(duplicate_literal_removal,[],[f6785])).
fof(f6785,plain,(
  ( ! [X255,X257,X256] : (s(bool,'T_0') = s(bool,X256) | ~p(s(bool,X255)) | s(bool,X255) = s(bool,X257) | s(bool,'T_0') = s(bool,X256) | s(bool,'F_0') = s(bool,X257)) )),
  inference(superposition,[],[f249,f3476])).
fof(f8016,plain,(
  ( ! [X253,X254,X252] : (p(s(bool,X253)) | s(bool,'T_0') != s(bool,X252) | s(bool,X252) = s(bool,X254) | s(bool,'F_0') = s(bool,X254)) )),
  inference(subsumption_resolution,[],[f6784,f246])).
fof(f6784,plain,(
  ( ! [X253,X254,X252] : (p(s(bool,X253)) | s(bool,'T_0') != s(bool,X252) | s(bool,X252) = s(bool,X254) | s(bool,'T_0') = s(bool,X253) | s(bool,'F_0') = s(bool,X254)) )),
  inference(superposition,[],[f248,f3476])).
fof(f7665,plain,(
  ( ! [X243,X244,X245] : (s(bool,'T_0') = s(bool,X244) | ~p(s(bool,X243)) | s(bool,X243) = s(bool,X245) | s(bool,'F_0') = s(bool,X245)) )),
  inference(duplicate_literal_removal,[],[f6781])).
fof(f6781,plain,(
  ( ! [X243,X244,X245] : (s(bool,'T_0') = s(bool,X244) | ~p(s(bool,X243)) | s(bool,X243) = s(bool,X245) | s(bool,'T_0') = s(bool,X244) | s(bool,'F_0') = s(bool,X245)) )),
  inference(superposition,[],[f247,f3476])).
fof(f8014,plain,(
  ( ! [X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,X240) | s(bool,X240) = s(bool,X242) | s(bool,'F_0') = s(bool,X242)) )),
  inference(subsumption_resolution,[],[f6780,f246])).
fof(f6780,plain,(
  ( ! [X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,X240) | s(bool,X240) = s(bool,X242) | s(bool,'T_0') = s(bool,X241) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f246,f3476])).
fof(f8013,plain,(
  ( ! [X235,X236,X234] : (p(s(bool,X235)) | s(bool,'F_0') = s(bool,X234) | s(bool,X234) = s(bool,X236) | s(bool,'F_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f6778,f246])).
fof(f6778,plain,(
  ( ! [X235,X236,X234] : (p(s(bool,X235)) | s(bool,'F_0') = s(bool,X234) | s(bool,X234) = s(bool,X236) | s(bool,'T_0') = s(bool,X235) | s(bool,'F_0') = s(bool,X236)) )),
  inference(superposition,[],[f243,f3476])).
fof(f8011,plain,(
  ( ! [X227,X225,X226] : (~p(s(bool,X225)) | s(bool,X225) = s(bool,X227) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,X227)) )),
  inference(subsumption_resolution,[],[f6775,f261])).
fof(f6775,plain,(
  ( ! [X227,X225,X226] : (s(bool,'F_0') != s(bool,X226) | ~p(s(bool,X225)) | s(bool,X225) = s(bool,X227) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,X227)) )),
  inference(superposition,[],[f242,f3476])).
fof(f8010,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~p(s(bool,X189)) | s(bool,X186) != s(bool,X189) | s(bool,X188) != s(bool,X190) | s(bool,X186) = s(bool,X191) | s(bool,'F_0') = s(bool,X191)) )),
  inference(subsumption_resolution,[],[f6765,f246])).
fof(f6765,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~p(s(bool,X189)) | s(bool,X186) != s(bool,X189) | s(bool,X188) != s(bool,X190) | s(bool,X186) = s(bool,X191) | s(bool,'T_0') = s(bool,X187) | s(bool,'F_0') = s(bool,X191)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8009,plain,(
  ( ! [X185,X182,X180,X184,X183,X181] : (s(bool,X181) != s(bool,X182) | p(s(bool,X182)) | ~p(s(bool,X183)) | ~p(s(bool,X180)) | s(bool,X183) != s(bool,X184) | s(bool,X180) = s(bool,X185) | s(bool,'F_0') = s(bool,X185)) )),
  inference(subsumption_resolution,[],[f6764,f2881])).
fof(f6764,plain,(
  ( ! [X185,X182,X180,X184,X183,X181] : (s(bool,X181) != s(bool,X182) | p(s(bool,X182)) | ~p(s(bool,X183)) | ~p(s(bool,X180)) | s(bool,X183) != s(bool,X184) | s(bool,X180) = s(bool,X185) | s(bool,'T_0') = s(bool,X181) | s(bool,'F_0') = s(bool,X185)) )),
  inference(superposition,[],[f203,f3476])).
fof(f8004,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X138) != s(bool,X141) | s(bool,X138) = s(bool,X143) | s(bool,'F_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f6757,f246])).
fof(f6757,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X138) != s(bool,X141) | s(bool,X138) = s(bool,X143) | s(bool,'T_0') = s(bool,X139) | s(bool,'F_0') = s(bool,X143)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8002,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,X127) != s(bool,X128) | p(s(bool,X129)) | ~p(s(bool,X126)) | ~p(s(bool,X130)) | p(s(bool,X128)) | s(bool,X126) = s(bool,X131) | s(bool,'F_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f6755,f2881])).
fof(f6755,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,X127) != s(bool,X128) | p(s(bool,X129)) | ~p(s(bool,X126)) | ~p(s(bool,X130)) | p(s(bool,X128)) | s(bool,X126) = s(bool,X131) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X131)) )),
  inference(superposition,[],[f202,f3476])).
fof(f8000,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X108) != s(bool,X111) | s(bool,X108) = s(bool,X113) | s(bool,'F_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f6752,f246])).
fof(f6752,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X108) != s(bool,X111) | s(bool,X108) = s(bool,X113) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X113)) )),
  inference(superposition,[],[f201,f3476])).
fof(f7996,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X84)) | p(s(bool,X86)) | s(bool,X87) != s(bool,X88) | s(bool,X84) = s(bool,X89) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f6748,f2881])).
fof(f6748,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X84)) | p(s(bool,X86)) | s(bool,X87) != s(bool,X88) | s(bool,X84) = s(bool,X89) | s(bool,'T_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f201,f3476])).
fof(f7995,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X66) != s(bool,X69) | s(bool,X68) != s(bool,X70) | s(bool,X66) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f6745,f246])).
fof(f6745,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X66) != s(bool,X69) | s(bool,X68) != s(bool,X70) | s(bool,X66) = s(bool,X71) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f199,f3476])).
fof(f7993,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X54)) | s(bool,X57) != s(bool,X58) | s(bool,X54) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f6743,f2881])).
fof(f6743,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X54)) | s(bool,X57) != s(bool,X58) | s(bool,X54) = s(bool,X59) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f199,f3476])).
fof(f7988,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X24) != s(bool,X27) | s(bool,X24) = s(bool,X29) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f6738,f246])).
fof(f6738,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X24) != s(bool,X27) | s(bool,X24) = s(bool,X29) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f197,f3476])).
fof(f7985,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,X7) != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X6)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f6735,f2881])).
fof(f6735,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,X7) != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X6)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f197,f3476])).
fof(f7983,plain,(
  ( ! [X679,X680,X682,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'T_0') = s(bool,X682) | s(bool,X682) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) | $spl204),
  inference(cnf_transformation,[],[f7983_D])).
fof(f7983_D,plain,(
  ( ! [X679,X680,X682,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | s(bool,'T_0') = s(bool,X682) | s(bool,X682) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680))))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7982,plain,(
  ( ! [X681,X679,X680,X682,X678] : (~p(s(bool,X681)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s(bool,X682) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s(bool,'T_0') = s(bool,X682)) )),
  inference(subsumption_resolution,[],[f6729,f242])).
fof(f6729,plain,(
  ( ! [X681,X679,X680,X682,X678] : (~p(s(bool,X681)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s(bool,X682) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680))))) | s(bool,'T_0') = s(bool,X682) | s(bool,'F_0') = s(bool,X681)) )),
  inference(superposition,[],[f258,f3476])).
fof(f7981,plain,(
  ( ! [X668,X670,X669,X672] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'T_0') = s(bool,X672) | s(bool,X672) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) | $spl202),
  inference(cnf_transformation,[],[f7981_D])).
fof(f7981_D,plain,(
  ( ! [X668,X670,X669,X672] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | s(bool,'T_0') = s(bool,X672) | s(bool,X672) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670))))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7980,plain,(
  ( ! [X668,X670,X669,X671,X672] : (~p(s(bool,X671)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s(bool,X672) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s(bool,'T_0') = s(bool,X672)) )),
  inference(subsumption_resolution,[],[f6727,f242])).
fof(f6727,plain,(
  ( ! [X668,X670,X669,X671,X672] : (~p(s(bool,X671)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s(bool,X672) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670))))) | s(bool,'T_0') = s(bool,X672) | s(bool,'F_0') = s(bool,X671)) )),
  inference(superposition,[],[f258,f3476])).
fof(f7979,plain,(
  ( ! [X660,X662,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'T_0') = s(bool,X662) | s(bool,X662) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) | $spl200),
  inference(cnf_transformation,[],[f7979_D])).
fof(f7979_D,plain,(
  ( ! [X660,X662,X658,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | s(bool,'T_0') = s(bool,X662) | s(bool,X662) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660))))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7978,plain,(
  ( ! [X660,X662,X658,X661,X659] : (~p(s(bool,X661)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s(bool,X662) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s(bool,'T_0') = s(bool,X662)) )),
  inference(subsumption_resolution,[],[f6725,f242])).
fof(f6725,plain,(
  ( ! [X660,X662,X658,X661,X659] : (~p(s(bool,X661)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s(bool,X662) = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660))))) | s(bool,'T_0') = s(bool,X662) | s(bool,'F_0') = s(bool,X661)) )),
  inference(superposition,[],[f258,f3476])).
fof(f6144,plain,(
  ( ! [X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X60) | s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X57))) = s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X58))) | s(bool,X60) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X56),X57),s('type.closSem.state'(X56),X58)))) ) | $spl130),
  inference(cnf_transformation,[],[f6144_D])).
fof(f6144_D,plain,(
  ( ! [X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X60) | s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X57))) = s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X58))) | s(bool,X60) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X56),X57),s('type.closSem.state'(X56),X58)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7977,plain,(
  ( ! [X656,X652,X654,X657,X653,X655] : (s(bool,X655) != s(bool,X656) | s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X653))) = s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X654))) | s(bool,X657) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X652),X653),s('type.closSem.state'(X652),X654))) | s(bool,'T_0') = s(bool,X657) | s(bool,'F_0') = s(bool,X655)) ) | ($spl60 | $spl70)),
  inference(subsumption_resolution,[],[f6724,f2856])).
fof(f6724,plain,(
  ( ! [X656,X652,X654,X657,X653,X655] : (s(bool,X655) != s(bool,X656) | s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X653))) = s('type.ffi.ffi_state'(X652),'const.closSem.state_ffi_1'(s('type.closSem.state'(X652),X654))) | ~p(s(bool,X656)) | s(bool,X657) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X652),X653),s('type.closSem.state'(X652),X654))) | s(bool,'T_0') = s(bool,X657) | s(bool,'F_0') = s(bool,X655)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f7976,plain,(
  ( ! [X648,X650,X649,X651,X647] : (s(bool,'T_0') != s(bool,X650) | s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X648))) = s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X649))) | s(bool,X651) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X647),X648),s('type.closSem.state'(X647),X649))) | s(bool,'T_0') = s(bool,X651)) )),
  inference(subsumption_resolution,[],[f6723,f2590])).
fof(f6723,plain,(
  ( ! [X648,X650,X649,X651,X647] : (s(bool,'T_0') != s(bool,X650) | s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X648))) = s('type.ffi.ffi_state'(X647),'const.closSem.state_ffi_1'(s('type.closSem.state'(X647),X649))) | s(bool,X651) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X647),X648),s('type.closSem.state'(X647),X649))) | s(bool,'T_0') = s(bool,X651) | s(bool,'F_0') = s(bool,X650)) )),
  inference(superposition,[],[f3144,f3476])).
fof(f7666,plain,(
  ( ! [X644,X646,X642,X645,X643] : (s(bool,'F_0') = s(bool,X645) | s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X643))) = s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X644))) | s(bool,X646) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X642),X643),s('type.closSem.state'(X642),X644))) | s(bool,'T_0') = s(bool,X646)) )),
  inference(duplicate_literal_removal,[],[f6722])).
fof(f6722,plain,(
  ( ! [X644,X646,X642,X645,X643] : (s(bool,'F_0') = s(bool,X645) | s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X643))) = s('type.ffi.ffi_state'(X642),'const.closSem.state_ffi_1'(s('type.closSem.state'(X642),X644))) | s(bool,X646) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X642),X643),s('type.closSem.state'(X642),X644))) | s(bool,'T_0') = s(bool,X646) | s(bool,'F_0') = s(bool,X645)) )),
  inference(superposition,[],[f3143,f3476])).
fof(f6082,plain,(
  ( ! [X30,X33,X31,X29] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31)))) ) | $spl128),
  inference(cnf_transformation,[],[f6082_D])).
fof(f6082_D,plain,(
  ( ! [X30,X33,X31,X29] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7974,plain,(
  ( ! [X639,X640,X636,X638,X641,X637] : (s(bool,X639) != s(bool,X640) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X637))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X638))) | s(bool,X641) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X636),X637),s('type.closSem.state'(X636),X638))) | s(bool,'T_0') = s(bool,X641) | s(bool,'F_0') = s(bool,X639)) ) | ($spl60 | $spl64)),
  inference(subsumption_resolution,[],[f6721,f2856])).
fof(f6721,plain,(
  ( ! [X639,X640,X636,X638,X641,X637] : (s(bool,X639) != s(bool,X640) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X637))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X636),X638))) | ~p(s(bool,X640)) | s(bool,X641) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X636),X637),s('type.closSem.state'(X636),X638))) | s(bool,'T_0') = s(bool,X641) | s(bool,'F_0') = s(bool,X639)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f7973,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'T_0') != s(bool,X634) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X632))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X633))) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X631),X632),s('type.closSem.state'(X631),X633))) | s(bool,'T_0') = s(bool,X635)) )),
  inference(subsumption_resolution,[],[f6720,f2590])).
fof(f6720,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'T_0') != s(bool,X634) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X632))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X631),X633))) | s(bool,X635) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X631),X632),s('type.closSem.state'(X631),X633))) | s(bool,'T_0') = s(bool,X635) | s(bool,'F_0') = s(bool,X634)) )),
  inference(superposition,[],[f2981,f3476])).
fof(f7667,plain,(
  ( ! [X629,X627,X628,X630,X626] : (s(bool,'F_0') = s(bool,X629) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X627))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X628))) | s(bool,X630) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X626),X627),s('type.closSem.state'(X626),X628))) | s(bool,'T_0') = s(bool,X630)) )),
  inference(duplicate_literal_removal,[],[f6719])).
fof(f6719,plain,(
  ( ! [X629,X627,X628,X630,X626] : (s(bool,'F_0') = s(bool,X629) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X627))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X626),X628))) | s(bool,X630) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X626),X627),s('type.closSem.state'(X626),X628))) | s(bool,'T_0') = s(bool,X630) | s(bool,'F_0') = s(bool,X629)) )),
  inference(superposition,[],[f2980,f3476])).
fof(f7972,plain,(
  ( ! [X617,X620,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'T_0') = s(bool,X620) | s(bool,X620) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) | $spl196),
  inference(cnf_transformation,[],[f7972_D])).
fof(f7972_D,plain,(
  ( ! [X617,X620,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,'T_0') = s(bool,X620) | s(bool,X620) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7971,plain,(
  ( ! [X617,X619,X620,X616,X618] : (~p(s(bool,X619)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,X620) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618))) | s(bool,'T_0') = s(bool,X620)) )),
  inference(subsumption_resolution,[],[f6717,f242])).
fof(f6717,plain,(
  ( ! [X617,X619,X620,X616,X618] : (~p(s(bool,X619)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,X620) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618))) | s(bool,'T_0') = s(bool,X620) | s(bool,'F_0') = s(bool,X619)) )),
  inference(superposition,[],[f257,f3476])).
fof(f7970,plain,(
  ( ! [X613,X615,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'T_0') = s(bool,X615) | s(bool,X615) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) | $spl194),
  inference(cnf_transformation,[],[f7970_D])).
fof(f7970_D,plain,(
  ( ! [X613,X615,X611,X612] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,'T_0') = s(bool,X615) | s(bool,X615) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7969,plain,(
  ( ! [X613,X615,X611,X612,X614] : (~p(s(bool,X614)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,X615) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613))) | s(bool,'T_0') = s(bool,X615)) )),
  inference(subsumption_resolution,[],[f6716,f242])).
fof(f6716,plain,(
  ( ! [X613,X615,X611,X612,X614] : (~p(s(bool,X614)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,X615) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613))) | s(bool,'T_0') = s(bool,X615) | s(bool,'F_0') = s(bool,X614)) )),
  inference(superposition,[],[f256,f3476])).
fof(f7968,plain,(
  ( ! [X606,X607,X608,X610] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'T_0') = s(bool,X610) | s(bool,X610) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) | $spl192),
  inference(cnf_transformation,[],[f7968_D])).
fof(f7968_D,plain,(
  ( ! [X606,X607,X608,X610] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,'T_0') = s(bool,X610) | s(bool,X610) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7967,plain,(
  ( ! [X606,X609,X607,X608,X610] : (~p(s(bool,X609)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,X610) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608))) | s(bool,'T_0') = s(bool,X610)) )),
  inference(subsumption_resolution,[],[f6715,f242])).
fof(f6715,plain,(
  ( ! [X606,X609,X607,X608,X610] : (~p(s(bool,X609)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,X610) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608))) | s(bool,'T_0') = s(bool,X610) | s(bool,'F_0') = s(bool,X609)) )),
  inference(superposition,[],[f255,f3476])).
fof(f7966,plain,(
  ( ! [X602,X605,X601,X603,X604] : (~p(s(bool,X604)) | s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X602))) = s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X603))) | s(bool,X605) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X601),X602),s('type.closSem.state'(X601),X603))) | s(bool,'T_0') = s(bool,X605)) )),
  inference(subsumption_resolution,[],[f6714,f242])).
fof(f6714,plain,(
  ( ! [X602,X605,X601,X603,X604] : (~p(s(bool,X604)) | s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X602))) = s('type.ffi.ffi_state'(X601),'const.closSem.state_ffi_1'(s('type.closSem.state'(X601),X603))) | s(bool,X605) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X601),X602),s('type.closSem.state'(X601),X603))) | s(bool,'T_0') = s(bool,X605) | s(bool,'F_0') = s(bool,X604)) )),
  inference(superposition,[],[f254,f3476])).
fof(f7965,plain,(
  ( ! [X600,X596,X598,X597,X599] : (~p(s(bool,X599)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X597))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X598))) | s(bool,X600) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X596),X597),s('type.closSem.state'(X596),X598))) | s(bool,'T_0') = s(bool,X600)) )),
  inference(subsumption_resolution,[],[f6713,f242])).
fof(f6713,plain,(
  ( ! [X600,X596,X598,X597,X599] : (~p(s(bool,X599)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X597))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X596),X598))) | s(bool,X600) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X596),X597),s('type.closSem.state'(X596),X598))) | s(bool,'T_0') = s(bool,X600) | s(bool,'F_0') = s(bool,X599)) )),
  inference(superposition,[],[f253,f3476])).
fof(f7964,plain,(
  ( ! [X593,X595] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,'T_0') = s(bool,X595) | s(bool,X595) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) ) | $spl190),
  inference(cnf_transformation,[],[f7964_D])).
fof(f7964_D,plain,(
  ( ! [X593,X595] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,'T_0') = s(bool,X595) | s(bool,X595) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7963,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,X595) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593))) | s(bool,'T_0') = s(bool,X595)) )),
  inference(subsumption_resolution,[],[f6712,f2590])).
fof(f6712,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,X595) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593))) | s(bool,'T_0') = s(bool,X595) | s(bool,'F_0') = s(bool,X594)) )),
  inference(superposition,[],[f1072,f3476])).
fof(f7961,plain,(
  ( ! [X589,X587] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,'T_0') = s(bool,X589) | s(bool,X589) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) ) | $spl188),
  inference(cnf_transformation,[],[f7961_D])).
fof(f7961_D,plain,(
  ( ! [X589,X587] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,'T_0') = s(bool,X589) | s(bool,X589) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7962,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,X592) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X590))) | s(bool,'T_0') = s(bool,X592)) )),
  inference(subsumption_resolution,[],[f6711,f2590])).
fof(f6711,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,X592) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X590))) | s(bool,'T_0') = s(bool,X592) | s(bool,'F_0') = s(bool,X591)) )),
  inference(superposition,[],[f968,f3476])).
fof(f7960,plain,(
  ( ! [X588,X589,X587] : (~p(s(bool,X588)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X589) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587))) | s(bool,'T_0') = s(bool,X589)) )),
  inference(subsumption_resolution,[],[f6710,f242])).
fof(f6710,plain,(
  ( ! [X588,X589,X587] : (~p(s(bool,X588)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X589) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587))) | s(bool,'T_0') = s(bool,X589) | s(bool,'F_0') = s(bool,X588)) )),
  inference(superposition,[],[f544,f3476])).
fof(f7959,plain,(
  ( ! [X581] : (s(bool,X581) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X581)) ) | $spl186),
  inference(cnf_transformation,[],[f7959_D])).
fof(f7959_D,plain,(
  ( ! [X581] : (s(bool,X581) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X581)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7958,plain,(
  ( ! [X580,X581] : (s(bool,'T_0') != s(bool,X580) | s(bool,X581) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X581)) )),
  inference(subsumption_resolution,[],[f6707,f2590])).
fof(f6707,plain,(
  ( ! [X580,X581] : (s(bool,'T_0') != s(bool,X580) | s(bool,X581) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,'T_0') = s(bool,X581) | s(bool,'F_0') = s(bool,X580)) )),
  inference(superposition,[],[f321,f3476])).
fof(f7957,plain,(
  ( ! [X577,X579] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,'T_0') = s(bool,X579) | s(bool,X579) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) ) | $spl184),
  inference(cnf_transformation,[],[f7957_D])).
fof(f7957_D,plain,(
  ( ! [X577,X579] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,'T_0') = s(bool,X579) | s(bool,X579) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7956,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,X579) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X579)) )),
  inference(subsumption_resolution,[],[f6706,f2590])).
fof(f6706,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,X579) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X579) | s(bool,'F_0') = s(bool,X578)) )),
  inference(superposition,[],[f1036,f3476])).
fof(f7954,plain,(
  ( ! [X573,X571] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,'T_0') = s(bool,X573) | s(bool,X573) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) ) | $spl182),
  inference(cnf_transformation,[],[f7954_D])).
fof(f7954_D,plain,(
  ( ! [X573,X571] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,'T_0') = s(bool,X573) | s(bool,X573) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7955,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,X576) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X574),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X576)) )),
  inference(subsumption_resolution,[],[f6705,f2590])).
fof(f6705,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,X576) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X574),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,X575)) )),
  inference(superposition,[],[f931,f3476])).
fof(f7953,plain,(
  ( ! [X573,X571,X572] : (~p(s(bool,X572)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X573) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X573)) )),
  inference(subsumption_resolution,[],[f6704,f242])).
fof(f6704,plain,(
  ( ! [X573,X571,X572] : (~p(s(bool,X572)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X573) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X573) | s(bool,'F_0') = s(bool,X572)) )),
  inference(superposition,[],[f543,f3476])).
fof(f7673,plain,(
  ( ! [X496,X498,X497] : (s(bool,X496) = s(bool,X497) | s(bool,'F_0') = s(bool,X497) | s(bool,'T_0') = s(bool,X498) | s(bool,'F_0') = s(bool,X496)) )),
  inference(duplicate_literal_removal,[],[f6687])).
fof(f6687,plain,(
  ( ! [X496,X498,X497] : (s(bool,X496) = s(bool,X497) | s(bool,'F_0') = s(bool,X497) | s(bool,'T_0') = s(bool,X498) | s(bool,'T_0') = s(bool,X498) | s(bool,'F_0') = s(bool,X496)) )),
  inference(superposition,[],[f261,f3476])).
fof(f7676,plain,(
  ( ! [X490,X489,X491] : (s(bool,X489) = s(bool,X490) | ~p(s(bool,X490)) | s(bool,'T_0') = s(bool,X491) | s(bool,'F_0') = s(bool,X489)) )),
  inference(duplicate_literal_removal,[],[f6684])).
fof(f6684,plain,(
  ( ! [X490,X489,X491] : (s(bool,X489) = s(bool,X490) | ~p(s(bool,X490)) | s(bool,'T_0') = s(bool,X491) | s(bool,'T_0') = s(bool,X491) | s(bool,'F_0') = s(bool,X489)) )),
  inference(superposition,[],[f249,f3476])).
fof(f7678,plain,(
  ( ! [X484,X485,X483] : (s(bool,X483) = s(bool,X484) | ~p(s(bool,X484)) | s(bool,'T_0') = s(bool,X485) | s(bool,'F_0') = s(bool,X483)) )),
  inference(duplicate_literal_removal,[],[f6682])).
fof(f6682,plain,(
  ( ! [X484,X485,X483] : (s(bool,X483) = s(bool,X484) | ~p(s(bool,X484)) | s(bool,'T_0') = s(bool,X485) | s(bool,'T_0') = s(bool,X485) | s(bool,'F_0') = s(bool,X483)) )),
  inference(superposition,[],[f247,f3476])).
fof(f7940,plain,(
  ( ! [X420,X422,X424,X421,X423] : (~p(s(bool,X421)) | s(bool,X422) = s(bool,X423) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X422) | s(bool,X420) = s(bool,X424) | s(bool,'T_0') = s(bool,X424)) )),
  inference(subsumption_resolution,[],[f6666,f242])).
fof(f6666,plain,(
  ( ! [X420,X422,X424,X421,X423] : (~p(s(bool,X421)) | s(bool,X422) = s(bool,X423) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X422) | s(bool,X420) = s(bool,X424) | s(bool,'T_0') = s(bool,X424) | s(bool,'F_0') = s(bool,X421)) )),
  inference(superposition,[],[f3477,f3476])).
fof(f7938,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415))) | ~p(s(bool,X411)) | s(bool,'T_0') = s(bool,X416) | s(bool,'F_0') = s(bool,X412)) ) | ($spl6 | $spl70)),
  inference(subsumption_resolution,[],[f6664,f766])).
fof(f6664,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415))) | ~p(s(bool,X411)) | s(bool,X411) = s(bool,X416) | s(bool,'T_0') = s(bool,X416) | s(bool,'F_0') = s(bool,X412)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f7857,plain,(
  ( ! [X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'T_0') = s(bool,X449) | s(bool,X446) = s(bool,X449) | s(bool,'T_0') = s(bool,X448) | s(bool,X447) != s(bool,X448)) ) | $spl154),
  inference(cnf_transformation,[],[f7857_D])).
fof(f7857_D,plain,(
  ( ! [X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'T_0') = s(bool,X449) | s(bool,X446) = s(bool,X449) | s(bool,'T_0') = s(bool,X448) | s(bool,X447) != s(bool,X448)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7935,plain,(
  ( ! [X372,X374,X376,X373,X375] : (~p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') = s(bool,X375) | s(bool,X372) = s(bool,X376) | s(bool,'T_0') = s(bool,X376)) ) | $spl58),
  inference(subsumption_resolution,[],[f6655,f242])).
fof(f6655,plain,(
  ( ! [X372,X374,X376,X373,X375] : (~p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') = s(bool,X375) | s(bool,X372) = s(bool,X376) | s(bool,'T_0') = s(bool,X376) | s(bool,'F_0') = s(bool,X373)) ) | $spl58),
  inference(superposition,[],[f3037,f3476])).
fof(f5308,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl106),
  inference(cnf_transformation,[],[f5308_D])).
fof(f5308_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7934,plain,(
  ( ! [X368,X370,X367,X369,X371] : (~p(s(bool,X368)) | s(bool,X367) != s(bool,X369) | s(bool,X369) != s(bool,X370) | p(s(bool,X370)) | s(bool,X367) = s(bool,X371) | s(bool,'T_0') = s(bool,X371)) ) | $spl58),
  inference(subsumption_resolution,[],[f6654,f242])).
fof(f6654,plain,(
  ( ! [X368,X370,X367,X369,X371] : (~p(s(bool,X368)) | s(bool,X367) != s(bool,X369) | s(bool,X369) != s(bool,X370) | p(s(bool,X370)) | s(bool,X367) = s(bool,X371) | s(bool,'T_0') = s(bool,X371) | s(bool,'F_0') = s(bool,X368)) ) | $spl58),
  inference(superposition,[],[f3036,f3476])).
fof(f7932,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365))) | ~p(s(bool,X361)) | s(bool,'T_0') = s(bool,X366) | s(bool,'F_0') = s(bool,X362)) ) | ($spl6 | $spl64)),
  inference(subsumption_resolution,[],[f6653,f766])).
fof(f6653,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365))) | ~p(s(bool,X361)) | s(bool,X361) = s(bool,X366) | s(bool,'T_0') = s(bool,X366) | s(bool,'F_0') = s(bool,X362)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f7891,plain,(
  ( ! [X681,X679,X680,X682,X678] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s(bool,X681) = s(bool,X682) | s(bool,'T_0') = s(bool,X682) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680)))))) )),
  inference(subsumption_resolution,[],[f6571,f2660])).
fof(f6571,plain,(
  ( ! [X681,X679,X680,X682,X678] : (~p(s(bool,X681)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X678),X679),s('type.closSem.state'(X678),X680)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X678),X680)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X679))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X678),X680)))))) | s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X679))) != s('type.ffi.ffi_state'(X678),'const.closSem.state_ffi_1'(s('type.closSem.state'(X678),X680))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X679))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X678),X680))) | s(bool,X681) = s(bool,X682) | s(bool,'T_0') = s(bool,X682) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X679))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X678),X680)))))) )),
  inference(superposition,[],[f258,f3476])).
fof(f7889,plain,(
  ( ! [X668,X670,X669,X671,X672] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670)))))) )),
  inference(subsumption_resolution,[],[f6569,f2660])).
fof(f6569,plain,(
  ( ! [X668,X670,X669,X671,X672] : (~p(s(bool,X671)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X668),X669),s('type.closSem.state'(X668),X670)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X668),X670)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X669))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X668),X670)))))) | s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X669))) != s('type.ffi.ffi_state'(X668),'const.closSem.state_ffi_1'(s('type.closSem.state'(X668),X670))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X669))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X668),X670))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X669))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X668),X670)))))) )),
  inference(superposition,[],[f258,f3476])).
fof(f7887,plain,(
  ( ! [X660,X662,X658,X661,X659] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s(bool,X661) = s(bool,X662) | s(bool,'T_0') = s(bool,X662) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660)))))) )),
  inference(subsumption_resolution,[],[f6567,f2660])).
fof(f6567,plain,(
  ( ! [X660,X662,X658,X661,X659] : (~p(s(bool,X661)) | p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X658),X659),s('type.closSem.state'(X658),X660)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X659))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X658),X660)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X658),X660)))))) | s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X659))) != s('type.ffi.ffi_state'(X658),'const.closSem.state_ffi_1'(s('type.closSem.state'(X658),X660))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X659))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X658),X660))) | s(bool,X661) = s(bool,X662) | s(bool,'T_0') = s(bool,X662) | s(bool,'F_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X659))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X658),X660)))))) )),
  inference(superposition,[],[f258,f3476])).
fof(f7881,plain,(
  ( ! [X617,X619,X620,X616,X618] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,X619) = s(bool,X620) | s(bool,'T_0') = s(bool,X620) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) )),
  inference(subsumption_resolution,[],[f6559,f2660])).
fof(f6559,plain,(
  ( ! [X617,X619,X620,X616,X618] : (~p(s(bool,X619)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X617))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X616),X618)))))) | s(bool,X619) = s(bool,X620) | s(bool,'T_0') = s(bool,X620) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X616),X617),s('type.closSem.state'(X616),X618)))) )),
  inference(superposition,[],[f257,f3476])).
fof(f7879,plain,(
  ( ! [X613,X615,X611,X612,X614] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,X614) = s(bool,X615) | s(bool,'T_0') = s(bool,X615) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) )),
  inference(subsumption_resolution,[],[f6558,f2660])).
fof(f6558,plain,(
  ( ! [X613,X615,X611,X612,X614] : (~p(s(bool,X614)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X612))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X611),X613)))))) | s(bool,X614) = s(bool,X615) | s(bool,'T_0') = s(bool,X615) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X611),X612),s('type.closSem.state'(X611),X613)))) )),
  inference(superposition,[],[f256,f3476])).
fof(f7877,plain,(
  ( ! [X606,X609,X607,X608,X610] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,X609) = s(bool,X610) | s(bool,'T_0') = s(bool,X610) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) )),
  inference(subsumption_resolution,[],[f6557,f2660])).
fof(f6557,plain,(
  ( ! [X606,X609,X607,X608,X610] : (~p(s(bool,X609)) | p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X607))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X606),X608)))))) | s(bool,X609) = s(bool,X610) | s(bool,'T_0') = s(bool,X610) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X606),X607),s('type.closSem.state'(X606),X608)))) )),
  inference(superposition,[],[f255,f3476])).
fof(f7875,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,'T_0') = s(bool,X595) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) )),
  inference(subsumption_resolution,[],[f6554,f2882])).
fof(f6554,plain,(
  ( ! [X594,X593,X595] : (s(bool,'T_0') != s(bool,X594) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X593))) | s(bool,X594) = s(bool,X595) | s(bool,'T_0') = s(bool,X595) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X593)))) )),
  inference(superposition,[],[f1072,f3476])).
fof(f7874,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,'T_0') = s(bool,X592) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X590)))) )),
  inference(subsumption_resolution,[],[f6553,f2882])).
fof(f6553,plain,(
  ( ! [X592,X590,X591] : (s(bool,'T_0') != s(bool,X591) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X590))) | s(bool,X591) = s(bool,X592) | s(bool,'T_0') = s(bool,X592) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X590)))) )),
  inference(superposition,[],[f968,f3476])).
fof(f7872,plain,(
  ( ! [X588,X589,X587] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X588) = s(bool,X589) | s(bool,'T_0') = s(bool,X589) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) )),
  inference(subsumption_resolution,[],[f6552,f2660])).
fof(f6552,plain,(
  ( ! [X588,X589,X587] : (~p(s(bool,X588)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X587))) | s(bool,X588) = s(bool,X589) | s(bool,'T_0') = s(bool,X589) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X587)))) )),
  inference(superposition,[],[f544,f3476])).
fof(f7869,plain,(
  ( ! [X580,X581] : (s(bool,'T_0') != s(bool,X580) | s(bool,'T_0') = s(bool,X581) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(subsumption_resolution,[],[f6549,f2882])).
fof(f6549,plain,(
  ( ! [X580,X581] : (s(bool,'T_0') != s(bool,X580) | s(bool,X580) = s(bool,X581) | s(bool,'T_0') = s(bool,X581) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(superposition,[],[f321,f3476])).
fof(f7867,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,'T_0') = s(bool,X579) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f6548,f2882])).
fof(f6548,plain,(
  ( ! [X578,X577,X579] : (s(bool,'T_0') != s(bool,X578) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X577))) | s(bool,X578) = s(bool,X579) | s(bool,'T_0') = s(bool,X579) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X577),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f1036,f3476])).
fof(f7866,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X574),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f6547,f2882])).
fof(f6547,plain,(
  ( ! [X575,X576,X574] : (s(bool,'T_0') != s(bool,X575) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X574))) | s(bool,X575) = s(bool,X576) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X574),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f931,f3476])).
fof(f7864,plain,(
  ( ! [X573,X571,X572] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X572) = s(bool,X573) | s(bool,'T_0') = s(bool,X573) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f6546,f2660])).
fof(f6546,plain,(
  ( ! [X573,X571,X572] : (~p(s(bool,X572)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X571))) | s(bool,X572) = s(bool,X573) | s(bool,'T_0') = s(bool,X573) | s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X571),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f543,f3476])).
fof(f7689,plain,(
  ( ! [X445,X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'F_0') = s(bool,X445) | s(bool,X447) != s(bool,X448) | s(bool,'T_0') = s(bool,X448) | s(bool,X446) = s(bool,X449) | s(bool,'T_0') = s(bool,X449)) ) | $spl60),
  inference(duplicate_literal_removal,[],[f6513])).
fof(f6513,plain,(
  ( ! [X445,X447,X449,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'F_0') = s(bool,X445) | s(bool,X447) != s(bool,X448) | s(bool,'T_0') = s(bool,X448) | s(bool,X446) = s(bool,X449) | s(bool,'T_0') = s(bool,X449) | s(bool,'F_0') = s(bool,X445)) ) | $spl60),
  inference(superposition,[],[f4623,f3476])).
fof(f7853,plain,(
  ( ! [X412,X414,X416,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'T_0') = s(bool,X416) | s(bool,X412) = s(bool,X416) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) | $spl152),
  inference(cnf_transformation,[],[f7853_D])).
fof(f7853_D,plain,(
  ( ! [X412,X414,X416,X413,X415] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s(bool,'T_0') = s(bool,X416) | s(bool,X412) = s(bool,X416) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f7852,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415))) | ~p(s(bool,X411)) | s(bool,X412) = s(bool,X416) | s(bool,'T_0') = s(bool,X416)) ) | $spl70),
  inference(subsumption_resolution,[],[f6506,f242])).
fof(f6506,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (s(bool,X412) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X413),X414),s('type.closSem.state'(X413),X415))) | s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X414))) = s('type.ffi.ffi_state'(X413),'const.closSem.state_ffi_1'(s('type.closSem.state'(X413),X415))) | ~p(s(bool,X411)) | s(bool,X412) = s(bool,X416) | s(bool,'T_0') = s(bool,X416) | s(bool,'F_0') = s(bool,X411)) ) | $spl70),
  inference(superposition,[],[f3186,f3476])).
fof(f7845,plain,(
  ( ! [X364,X366,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'T_0') = s(bool,X366) | s(bool,X362) = s(bool,X366) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) | $spl150),
  inference(cnf_transformation,[],[f7845_D])).
fof(f7845_D,plain,(
  ( ! [X364,X366,X362,X365,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s(bool,'T_0') = s(bool,X366) | s(bool,X362) = s(bool,X366) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f7844,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365))) | ~p(s(bool,X361)) | s(bool,X362) = s(bool,X366) | s(bool,'T_0') = s(bool,X366)) ) | $spl64),
  inference(subsumption_resolution,[],[f6495,f242])).
fof(f6495,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X362) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X363),X364),s('type.closSem.state'(X363),X365))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X364))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X363),X365))) | ~p(s(bool,X361)) | s(bool,X362) = s(bool,X366) | s(bool,'T_0') = s(bool,X366) | s(bool,'F_0') = s(bool,X361)) ) | $spl64),
  inference(superposition,[],[f3029,f3476])).
fof(f7795,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f6415,f4329])).
fof(f6415,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7693,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(duplicate_literal_removal,[],[f6412])).
fof(f6412,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7694,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X65)) )),
  inference(duplicate_literal_removal,[],[f6402])).
fof(f6402,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7790,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f6398,f4329])).
fof(f6398,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,X45) | s(bool,X45) = s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7695,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f6382])).
fof(f6382,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f3476,f261])).
fof(f7785,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | p(s(bool,X56))) )),
  inference(subsumption_resolution,[],[f6381,f246])).
fof(f6381,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | p(s(bool,X56))) )),
  inference(superposition,[],[f3476,f243])).
fof(f7784,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X52) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f6380,f4329])).
fof(f6380,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7751,plain,(
  ( ! [X61,X64,X62,X65] : (s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X62)) ) | $spl140),
  inference(cnf_transformation,[],[f7751_D])).
fof(f7751_D,plain,(
  ( ! [X61,X64,X62,X65] : (s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X62)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7696,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X35) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(duplicate_literal_removal,[],[f6378])).
fof(f6378,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X35) = s(bool,X39) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7699,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,X59) = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X58)) )),
  inference(duplicate_literal_removal,[],[f6367])).
fof(f6367,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,X59) = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7777,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s(bool,'F_0') = s(bool,X100) | s(bool,X100) = s(bool,X101) | s(bool,'F_0') = s(bool,X101) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X99)))) )),
  inference(subsumption_resolution,[],[f6351,f2981])).
fof(f6351,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s(bool,'F_0') = s(bool,X100) | s(bool,X100) = s(bool,X101) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X97),X98),s('type.closSem.state'(X97),X99))) | s(bool,'F_0') = s(bool,X101) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X98))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X97),X99)))) )),
  inference(superposition,[],[f3476,f2980])).
fof(f7775,plain,(
  ( ! [X94,X92,X95,X93,X96] : (s(bool,'F_0') = s(bool,X95) | s(bool,X95) = s(bool,X96) | s(bool,'F_0') = s(bool,X96) | s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X93))) = s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X94)))) )),
  inference(subsumption_resolution,[],[f6350,f3144])).
fof(f6350,plain,(
  ( ! [X94,X92,X95,X93,X96] : (s(bool,'F_0') = s(bool,X95) | s(bool,X95) = s(bool,X96) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X92),X93),s('type.closSem.state'(X92),X94))) | s(bool,'F_0') = s(bool,X96) | s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X93))) = s('type.ffi.ffi_state'(X92),'const.closSem.state_ffi_1'(s('type.closSem.state'(X92),X94)))) )),
  inference(superposition,[],[f3476,f3143])).
fof(f7702,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,'F_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f6339])).
fof(f6339,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7770,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f6335,f4329])).
fof(f6335,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7705,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X14)) )),
  inference(duplicate_literal_removal,[],[f6329])).
fof(f6329,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f3476,f261])).
fof(f7763,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f6325,f246])).
fof(f6325,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X0))) )),
  inference(superposition,[],[f3476,f243])).
fof(f7706,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,X95)) )),
  inference(duplicate_literal_removal,[],[f6321])).
fof(f6321,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f3476,f285])).
fof(f7709,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X84) | s(bool,X83) = s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(duplicate_literal_removal,[],[f6318])).
fof(f6318,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,X83) = s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7757,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f6309,f4329])).
fof(f6309,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7710,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,X31) = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(duplicate_literal_removal,[],[f6306])).
fof(f6306,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,X31) = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7713,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X9)) )),
  inference(duplicate_literal_removal,[],[f6300])).
fof(f6300,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f3476,f261])).
fof(f7714,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(duplicate_literal_removal,[],[f6299])).
fof(f6299,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f3476,f249])).
fof(f7715,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) )),
  inference(duplicate_literal_removal,[],[f6298])).
fof(f6298,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f3476,f247])).
fof(f7716,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s(bool,'F_0') = s(bool,X104) | s(bool,X105) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X101),X102),s('type.closSem.state'(X101),X103))) | s(bool,'T_0') = s(bool,X105) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X101),X102))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X101),X103)))) )),
  inference(duplicate_literal_removal,[],[f6295])).
fof(f6295,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s(bool,'F_0') = s(bool,X104) | s(bool,X105) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X101),X102),s('type.closSem.state'(X101),X103))) | s(bool,'T_0') = s(bool,X105) | s(bool,'F_0') = s(bool,X104) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X101),X102))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X101),X103)))) )),
  inference(superposition,[],[f3476,f2980])).
fof(f7717,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,'F_0') = s(bool,X99) | s(bool,X100) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X96),X97),s('type.closSem.state'(X96),X98))) | s(bool,'T_0') = s(bool,X100) | s('type.ffi.ffi_state'(X96),'const.closSem.state_ffi_1'(s('type.closSem.state'(X96),X97))) = s('type.ffi.ffi_state'(X96),'const.closSem.state_ffi_1'(s('type.closSem.state'(X96),X98)))) )),
  inference(duplicate_literal_removal,[],[f6294])).
fof(f6294,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,'F_0') = s(bool,X99) | s(bool,X100) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X96),X97),s('type.closSem.state'(X96),X98))) | s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,X99) | s('type.ffi.ffi_state'(X96),'const.closSem.state_ffi_1'(s('type.closSem.state'(X96),X97))) = s('type.ffi.ffi_state'(X96),'const.closSem.state_ffi_1'(s('type.closSem.state'(X96),X98)))) )),
  inference(superposition,[],[f3476,f3143])).
fof(f7718,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,X81) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81) | ~p(s(bool,X80))) )),
  inference(duplicate_literal_removal,[],[f6289])).
fof(f6289,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,X81) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81) | ~p(s(bool,X80))) )),
  inference(superposition,[],[f3476,f247])).
fof(f7719,plain,(
  ( ! [X78,X79,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78) | ~p(s(bool,X77))) )),
  inference(duplicate_literal_removal,[],[f6288])).
fof(f6288,plain,(
  ( ! [X78,X79,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78) | ~p(s(bool,X77))) )),
  inference(superposition,[],[f3476,f249])).
fof(f7720,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(duplicate_literal_removal,[],[f6287])).
fof(f6287,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(superposition,[],[f3476,f261])).
fof(f7723,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X62) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(duplicate_literal_removal,[],[f6284])).
fof(f6284,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X62) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7748,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X43) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f6280,f4329])).
fof(f6280,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X43) | s(bool,X41) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X42)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7724,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f6279])).
fof(f6279,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7735,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f6254,f4329])).
fof(f6254,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f7727,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(duplicate_literal_removal,[],[f6253])).
fof(f6253,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f3476,f3476])).
fof(f6187,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s('type.ffi.ffi_state'(X44),'const.closSem.state_ffi_1'(s('type.closSem.state'(X44),X45))) = s('type.ffi.ffi_state'(X44),'const.closSem.state_ffi_1'(s('type.closSem.state'(X44),X46))) | ~p(s(bool,X47)) | s(bool,X48) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X44),X45),s('type.closSem.state'(X44),X46))) | s(bool,'T_0') = s(bool,X48)) ) | $spl70),
  inference(subsumption_resolution,[],[f6172,f247])).
fof(f6172,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'T_0') != s(bool,X47) | s('type.ffi.ffi_state'(X44),'const.closSem.state_ffi_1'(s('type.closSem.state'(X44),X45))) = s('type.ffi.ffi_state'(X44),'const.closSem.state_ffi_1'(s('type.closSem.state'(X44),X46))) | ~p(s(bool,X47)) | s(bool,X48) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X44),X45),s('type.closSem.state'(X44),X46))) | s(bool,'T_0') = s(bool,X48)) ) | $spl70),
  inference(superposition,[],[f3186,f3074])).
fof(f6186,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.ffi.ffi_state'(X29),'const.closSem.state_ffi_1'(s('type.closSem.state'(X29),X30))) = s('type.ffi.ffi_state'(X29),'const.closSem.state_ffi_1'(s('type.closSem.state'(X29),X31))) | ~p(s(bool,X32)) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33)) ) | $spl70),
  inference(subsumption_resolution,[],[f6169,f247])).
fof(f6169,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') != s(bool,X32) | s('type.ffi.ffi_state'(X29),'const.closSem.state_ffi_1'(s('type.closSem.state'(X29),X30))) = s('type.ffi.ffi_state'(X29),'const.closSem.state_ffi_1'(s('type.closSem.state'(X29),X31))) | ~p(s(bool,X32)) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33)) ) | $spl70),
  inference(superposition,[],[f3186,f3074])).
fof(f6145,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X61),X62),s('type.closSem.state'(X61),X63))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X61),X62))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X61),X63)))) )),
  inference(subsumption_resolution,[],[f6122,f261])).
fof(f6122,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'F_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X61),X62),s('type.closSem.state'(X61),X63))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X61),X62))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X61),X63)))) )),
  inference(superposition,[],[f4329,f2980])).
fof(f6143,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X60) | s(bool,X60) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X56),X57),s('type.closSem.state'(X56),X58))) | s(bool,'T_0') = s(bool,X59) | s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X57))) = s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X58)))) )),
  inference(subsumption_resolution,[],[f6121,f261])).
fof(f6121,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X59) | s(bool,'T_0') = s(bool,X60) | s(bool,X60) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X56),X57),s('type.closSem.state'(X56),X58))) | s(bool,'T_0') = s(bool,X59) | s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X57))) = s('type.ffi.ffi_state'(X56),'const.closSem.state_ffi_1'(s('type.closSem.state'(X56),X58)))) )),
  inference(superposition,[],[f4329,f3143])).
fof(f6083,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X44),X45))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X44),X46))) | ~p(s(bool,X47)) | s(bool,X48) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X44),X45),s('type.closSem.state'(X44),X46))) | s(bool,'T_0') = s(bool,X48)) ) | $spl64),
  inference(subsumption_resolution,[],[f6067,f247])).
fof(f6067,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'T_0') != s(bool,X47) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X44),X45))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X44),X46))) | ~p(s(bool,X47)) | s(bool,X48) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X44),X45),s('type.closSem.state'(X44),X46))) | s(bool,'T_0') = s(bool,X48)) ) | $spl64),
  inference(superposition,[],[f3029,f3074])).
fof(f6081,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X31))) | ~p(s(bool,X32)) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33)) ) | $spl64),
  inference(subsumption_resolution,[],[f6064,f247])).
fof(f6064,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') != s(bool,X32) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X30))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X29),X31))) | ~p(s(bool,X32)) | s(bool,X33) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X29),X30),s('type.closSem.state'(X29),X31))) | s(bool,'T_0') = s(bool,X33)) ) | $spl64),
  inference(superposition,[],[f3029,f3074])).
fof(f981,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1))) | ~p(s(bool,X3)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1)))) ) | $spl34),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1))) | ~p(s(bool,X3)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f6012,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X39))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X39)))) )),
  inference(resolution,[],[f544,f262])).
fof(f6005,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X29))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f544,f203])).
fof(f6003,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X21)))) )),
  inference(resolution,[],[f544,f202])).
fof(f6001,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X13)))) )),
  inference(resolution,[],[f544,f201])).
fof(f6000,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X9))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f544,f199])).
fof(f5998,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1)))) )),
  inference(resolution,[],[f544,f197])).
fof(f945,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,X3)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1)))) ) | $spl32),
  inference(cnf_transformation,[],[f945_D])).
fof(f945_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,X3)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5967,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X39))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X39),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f262])).
fof(f5960,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X29),s('type.closSem.state'(sK0),sK3))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f543,f203])).
fof(f5958,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X21),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f202])).
fof(f5956,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X13),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f201])).
fof(f5955,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X9),s('type.closSem.state'(sK0),sK3))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f543,f199])).
fof(f5953,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f197])).
fof(f5907,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X50)) ) | $spl126),
  inference(cnf_transformation,[],[f5907_D])).
fof(f5907_D,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X50)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f5908,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'T_0') != s(bool,X53)) ) | $spl58),
  inference(subsumption_resolution,[],[f5894,f246])).
fof(f5894,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'T_0') != s(bool,X53) | s(bool,'T_0') = s(bool,X51)) ) | $spl58),
  inference(superposition,[],[f4702,f261])).
fof(f5899,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,X50)) ) | $spl58),
  inference(duplicate_literal_removal,[],[f5893])).
fof(f5893,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,X50) | p(s(bool,X48))) ) | $spl58),
  inference(superposition,[],[f4702,f243])).
fof(f5904,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') != s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | $spl58),
  inference(subsumption_resolution,[],[f5886,f246])).
fof(f5886,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') != s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl58),
  inference(superposition,[],[f4702,f3074])).
fof(f5902,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X19)) ) | $spl58),
  inference(subsumption_resolution,[],[f5885,f246])).
fof(f5885,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl58),
  inference(superposition,[],[f4702,f3074])).
fof(f5900,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl122),
  inference(cnf_transformation,[],[f5900_D])).
fof(f5900_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f5866,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'T_0') != s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) ) | $spl58),
  inference(resolution,[],[f4702,f262])).
fof(f5858,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl58),
  inference(resolution,[],[f4702,f203])).
fof(f5855,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl58),
  inference(resolution,[],[f4702,f202])).
fof(f5854,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl58),
  inference(resolution,[],[f4702,f201])).
fof(f5853,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl58),
  inference(resolution,[],[f4702,f199])).
fof(f5851,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl58),
  inference(resolution,[],[f4702,f197])).
fof(f5821,plain,(
  ( ! [X380,X382,X384,X381] : (p(s(bool,X384)) | s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X381))) = s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X382))) | s(bool,X384) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X380),X381),s('type.closSem.state'(X380),X382)))) ) | $spl120),
  inference(cnf_transformation,[],[f5821_D])).
fof(f5821_D,plain,(
  ( ! [X380,X382,X384,X381] : (p(s(bool,X384)) | s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X381))) = s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X382))) | s(bool,X384) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X380),X381),s('type.closSem.state'(X380),X382)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f5820,plain,(
  ( ! [X383,X380,X382,X384,X381] : (p(s(bool,X384)) | p(s(bool,X383)) | s(bool,X384) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X380),X381),s('type.closSem.state'(X380),X382))) | s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X381))) = s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X382)))) ) | $spl62),
  inference(subsumption_resolution,[],[f5767,f243])).
fof(f5767,plain,(
  ( ! [X383,X380,X382,X384,X381] : (s(bool,'F_0') != s(bool,X383) | p(s(bool,X384)) | p(s(bool,X383)) | s(bool,X384) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X380),X381),s('type.closSem.state'(X380),X382))) | s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X381))) = s('type.ffi.ffi_state'(X380),'const.closSem.state_ffi_1'(s('type.closSem.state'(X380),X382)))) ) | $spl62),
  inference(superposition,[],[f2964,f3143])).
fof(f3186,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) | $spl70),
  inference(cnf_transformation,[],[f3186_D])).
fof(f3186_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5812,plain,(
  ( ! [X308,X310,X306,X309,X307] : (p(s(bool,X309)) | ~p(s(bool,X310)) | s(bool,X310) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X306),X307),s('type.closSem.state'(X306),X308))) | s('type.ffi.ffi_state'(X306),'const.closSem.state_ffi_1'(s('type.closSem.state'(X306),X307))) = s('type.ffi.ffi_state'(X306),'const.closSem.state_ffi_1'(s('type.closSem.state'(X306),X308)))) )),
  inference(subsumption_resolution,[],[f5749,f250])).
fof(f5749,plain,(
  ( ! [X308,X310,X306,X309,X307] : (p(s(bool,'F_0')) | p(s(bool,X309)) | ~p(s(bool,X310)) | s(bool,X310) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X306),X307),s('type.closSem.state'(X306),X308))) | s('type.ffi.ffi_state'(X306),'const.closSem.state_ffi_1'(s('type.closSem.state'(X306),X307))) = s('type.ffi.ffi_state'(X306),'const.closSem.state_ffi_1'(s('type.closSem.state'(X306),X308)))) )),
  inference(superposition,[],[f262,f3143])).
fof(f5808,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (~p(s(bool,X218)) | ~p(s(bool,X219)) | s(bool,X219) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X215),X216),s('type.closSem.state'(X215),X217))) | s(bool,X218) != s(bool,X220) | s('type.ffi.ffi_state'(X215),'const.closSem.state_ffi_1'(s('type.closSem.state'(X215),X216))) = s('type.ffi.ffi_state'(X215),'const.closSem.state_ffi_1'(s('type.closSem.state'(X215),X217)))) )),
  inference(subsumption_resolution,[],[f5725,f250])).
fof(f5725,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (p(s(bool,'F_0')) | ~p(s(bool,X218)) | ~p(s(bool,X219)) | s(bool,X219) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X215),X216),s('type.closSem.state'(X215),X217))) | s(bool,X218) != s(bool,X220) | s('type.ffi.ffi_state'(X215),'const.closSem.state_ffi_1'(s('type.closSem.state'(X215),X216))) = s('type.ffi.ffi_state'(X215),'const.closSem.state_ffi_1'(s('type.closSem.state'(X215),X217)))) )),
  inference(superposition,[],[f203,f3143])).
fof(f5805,plain,(
  ( ! [X167,X171,X169,X172,X170,X168] : (p(s(bool,X170)) | ~p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X171) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X167),X168),s('type.closSem.state'(X167),X169))) | s('type.ffi.ffi_state'(X167),'const.closSem.state_ffi_1'(s('type.closSem.state'(X167),X168))) = s('type.ffi.ffi_state'(X167),'const.closSem.state_ffi_1'(s('type.closSem.state'(X167),X169)))) )),
  inference(subsumption_resolution,[],[f5717,f250])).
fof(f5717,plain,(
  ( ! [X167,X171,X169,X172,X170,X168] : (p(s(bool,'F_0')) | p(s(bool,X170)) | ~p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X171) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X167),X168),s('type.closSem.state'(X167),X169))) | s('type.ffi.ffi_state'(X167),'const.closSem.state_ffi_1'(s('type.closSem.state'(X167),X168))) = s('type.ffi.ffi_state'(X167),'const.closSem.state_ffi_1'(s('type.closSem.state'(X167),X169)))) )),
  inference(superposition,[],[f202,f3143])).
fof(f5803,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X140) != s(bool,X142) | s(bool,X141) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X137),X138),s('type.closSem.state'(X137),X139))) | s('type.ffi.ffi_state'(X137),'const.closSem.state_ffi_1'(s('type.closSem.state'(X137),X138))) = s('type.ffi.ffi_state'(X137),'const.closSem.state_ffi_1'(s('type.closSem.state'(X137),X139)))) )),
  inference(subsumption_resolution,[],[f5712,f250])).
fof(f5712,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (p(s(bool,'F_0')) | p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X140) != s(bool,X142) | s(bool,X141) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X137),X138),s('type.closSem.state'(X137),X139))) | s('type.ffi.ffi_state'(X137),'const.closSem.state_ffi_1'(s('type.closSem.state'(X137),X138))) = s('type.ffi.ffi_state'(X137),'const.closSem.state_ffi_1'(s('type.closSem.state'(X137),X139)))) )),
  inference(superposition,[],[f201,f3143])).
fof(f5800,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X95),X96),s('type.closSem.state'(X95),X97))) | s(bool,X98) != s(bool,X100) | s('type.ffi.ffi_state'(X95),'const.closSem.state_ffi_1'(s('type.closSem.state'(X95),X96))) = s('type.ffi.ffi_state'(X95),'const.closSem.state_ffi_1'(s('type.closSem.state'(X95),X97)))) )),
  inference(subsumption_resolution,[],[f5705,f250])).
fof(f5705,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (p(s(bool,'F_0')) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X95),X96),s('type.closSem.state'(X95),X97))) | s(bool,X98) != s(bool,X100) | s('type.ffi.ffi_state'(X95),'const.closSem.state_ffi_1'(s('type.closSem.state'(X95),X96))) = s('type.ffi.ffi_state'(X95),'const.closSem.state_ffi_1'(s('type.closSem.state'(X95),X97)))) )),
  inference(superposition,[],[f199,f3143])).
fof(f5797,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X53),X54),s('type.closSem.state'(X53),X55))) | s('type.ffi.ffi_state'(X53),'const.closSem.state_ffi_1'(s('type.closSem.state'(X53),X54))) = s('type.ffi.ffi_state'(X53),'const.closSem.state_ffi_1'(s('type.closSem.state'(X53),X55)))) )),
  inference(subsumption_resolution,[],[f5698,f250])).
fof(f5698,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,'F_0')) | p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X53),X54),s('type.closSem.state'(X53),X55))) | s('type.ffi.ffi_state'(X53),'const.closSem.state_ffi_1'(s('type.closSem.state'(X53),X54))) = s('type.ffi.ffi_state'(X53),'const.closSem.state_ffi_1'(s('type.closSem.state'(X53),X55)))) )),
  inference(superposition,[],[f197,f3143])).
fof(f5616,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f3477,f262])).
fof(f5609,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X42) = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3477,f203])).
fof(f5607,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f3477,f202])).
fof(f5605,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3477,f201])).
fof(f5604,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3477,f199])).
fof(f5602,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3477,f197])).
fof(f5311,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X22)) ) | $spl108),
  inference(cnf_transformation,[],[f5311_D])).
fof(f5311_D,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X22)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5576,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f5560,f246])).
fof(f5560,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f4328,f3074])).
fof(f5575,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f5559,f246])).
fof(f5559,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f4328,f3074])).
fof(f5573,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl116),
  inference(cnf_transformation,[],[f5573_D])).
fof(f5573_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5541,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) = s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'T_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) )),
  inference(resolution,[],[f4328,f262])).
fof(f5533,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f4328,f203])).
fof(f5530,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) )),
  inference(resolution,[],[f4328,f202])).
fof(f5529,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) )),
  inference(resolution,[],[f4328,f201])).
fof(f5528,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) = s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f4328,f199])).
fof(f5526,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) )),
  inference(resolution,[],[f4328,f197])).
fof(f5519,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X22) | s(bool,X21) != s(bool,X22)) ) | $spl114),
  inference(cnf_transformation,[],[f5519_D])).
fof(f5519_D,plain,(
  ( ! [X21,X19,X22] : (s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X22) | s(bool,X21) != s(bool,X22)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f5520,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | $spl68),
  inference(subsumption_resolution,[],[f5503,f246])).
fof(f5503,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl68),
  inference(superposition,[],[f3130,f3074])).
fof(f5518,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) ) | $spl68),
  inference(subsumption_resolution,[],[f5502,f246])).
fof(f5502,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl68),
  inference(superposition,[],[f3130,f3074])).
fof(f5516,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl112),
  inference(cnf_transformation,[],[f5516_D])).
fof(f5516_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f5484,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) ) | $spl68),
  inference(resolution,[],[f3130,f262])).
fof(f5476,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl68),
  inference(resolution,[],[f3130,f203])).
fof(f5473,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl68),
  inference(resolution,[],[f3130,f202])).
fof(f5472,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl68),
  inference(resolution,[],[f3130,f201])).
fof(f5471,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl68),
  inference(resolution,[],[f3130,f199])).
fof(f5469,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl68),
  inference(resolution,[],[f3130,f197])).
fof(f5464,plain,(
  ( ! [X374,X376,X378,X375] : (p(s(bool,X378)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X375))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X376))) | s(bool,X378) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X374),X375),s('type.closSem.state'(X374),X376)))) ) | $spl110),
  inference(cnf_transformation,[],[f5464_D])).
fof(f5464_D,plain,(
  ( ! [X374,X376,X378,X375] : (p(s(bool,X378)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X375))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X376))) | s(bool,X378) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X374),X375),s('type.closSem.state'(X374),X376)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5463,plain,(
  ( ! [X377,X374,X376,X378,X375] : (p(s(bool,X378)) | p(s(bool,X377)) | s(bool,X378) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X374),X375),s('type.closSem.state'(X374),X376))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X375))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X376)))) ) | $spl62),
  inference(subsumption_resolution,[],[f5413,f243])).
fof(f5413,plain,(
  ( ! [X377,X374,X376,X378,X375] : (s(bool,'F_0') != s(bool,X377) | p(s(bool,X378)) | p(s(bool,X377)) | s(bool,X378) = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X374),X375),s('type.closSem.state'(X374),X376))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X375))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X374),X376)))) ) | $spl62),
  inference(superposition,[],[f2964,f2980])).
fof(f3029,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) | $spl64),
  inference(cnf_transformation,[],[f3029_D])).
fof(f3029_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5455,plain,(
  ( ! [X300,X302,X304,X301,X303] : (p(s(bool,X303)) | ~p(s(bool,X304)) | s(bool,X304) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X300),X301),s('type.closSem.state'(X300),X302))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X300),X301))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X300),X302)))) )),
  inference(subsumption_resolution,[],[f5395,f250])).
fof(f5395,plain,(
  ( ! [X300,X302,X304,X301,X303] : (p(s(bool,'F_0')) | p(s(bool,X303)) | ~p(s(bool,X304)) | s(bool,X304) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X300),X301),s('type.closSem.state'(X300),X302))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X300),X301))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X300),X302)))) )),
  inference(superposition,[],[f262,f2980])).
fof(f5451,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (~p(s(bool,X212)) | ~p(s(bool,X213)) | s(bool,X213) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X209),X210),s('type.closSem.state'(X209),X211))) | s(bool,X212) != s(bool,X214) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X209),X210))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X209),X211)))) )),
  inference(subsumption_resolution,[],[f5371,f250])).
fof(f5371,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (p(s(bool,'F_0')) | ~p(s(bool,X212)) | ~p(s(bool,X213)) | s(bool,X213) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X209),X210),s('type.closSem.state'(X209),X211))) | s(bool,X212) != s(bool,X214) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X209),X210))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X209),X211)))) )),
  inference(superposition,[],[f203,f2980])).
fof(f5448,plain,(
  ( ! [X165,X163,X161,X166,X164,X162] : (p(s(bool,X164)) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X165) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X161),X162),s('type.closSem.state'(X161),X163))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X161),X162))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X161),X163)))) )),
  inference(subsumption_resolution,[],[f5363,f250])).
fof(f5363,plain,(
  ( ! [X165,X163,X161,X166,X164,X162] : (p(s(bool,'F_0')) | p(s(bool,X164)) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X165) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X161),X162),s('type.closSem.state'(X161),X163))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X161),X162))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X161),X163)))) )),
  inference(superposition,[],[f202,f2980])).
fof(f5446,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (p(s(bool,X134)) | ~p(s(bool,X135)) | s(bool,X134) != s(bool,X136) | s(bool,X135) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X131),X132),s('type.closSem.state'(X131),X133))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X131),X132))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X131),X133)))) )),
  inference(subsumption_resolution,[],[f5358,f250])).
fof(f5358,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (p(s(bool,'F_0')) | p(s(bool,X134)) | ~p(s(bool,X135)) | s(bool,X134) != s(bool,X136) | s(bool,X135) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X131),X132),s('type.closSem.state'(X131),X133))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X131),X132))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X131),X133)))) )),
  inference(superposition,[],[f201,f2980])).
fof(f5443,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X89),X90),s('type.closSem.state'(X89),X91))) | s(bool,X92) != s(bool,X94) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X89),X90))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X89),X91)))) )),
  inference(subsumption_resolution,[],[f5351,f250])).
fof(f5351,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (p(s(bool,'F_0')) | ~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X89),X90),s('type.closSem.state'(X89),X91))) | s(bool,X92) != s(bool,X94) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X89),X90))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X89),X91)))) )),
  inference(superposition,[],[f199,f2980])).
fof(f5440,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X47),X48),s('type.closSem.state'(X47),X49))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X48))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X49)))) )),
  inference(subsumption_resolution,[],[f5344,f250])).
fof(f5344,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (p(s(bool,'F_0')) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X47),X48),s('type.closSem.state'(X47),X49))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X48))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X47),X49)))) )),
  inference(superposition,[],[f197,f2980])).
fof(f5312,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | $spl66),
  inference(subsumption_resolution,[],[f5296,f246])).
fof(f5296,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl66),
  inference(superposition,[],[f3095,f3074])).
fof(f5310,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) ) | $spl66),
  inference(subsumption_resolution,[],[f5295,f246])).
fof(f5295,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl66),
  inference(superposition,[],[f3095,f3074])).
fof(f5277,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) ) | $spl66),
  inference(resolution,[],[f3095,f262])).
fof(f5269,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl66),
  inference(resolution,[],[f3095,f203])).
fof(f5266,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl66),
  inference(resolution,[],[f3095,f202])).
fof(f5265,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl66),
  inference(resolution,[],[f3095,f201])).
fof(f5264,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl66),
  inference(resolution,[],[f3095,f199])).
fof(f5262,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl66),
  inference(resolution,[],[f3095,f197])).
fof(f5252,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl104),
  inference(cnf_transformation,[],[f5252_D])).
fof(f5252_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f5212,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | s(bool,X66) != s(bool,X68) | s(bool,'T_0') = s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X67) != s(bool,X70)) ) | $spl58),
  inference(resolution,[],[f3037,f262])).
fof(f5205,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl58),
  inference(resolution,[],[f3037,f203])).
fof(f5203,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl58),
  inference(resolution,[],[f3037,f202])).
fof(f5201,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl58),
  inference(resolution,[],[f3037,f201])).
fof(f5200,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl58),
  inference(resolution,[],[f3037,f199])).
fof(f5198,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl58),
  inference(resolution,[],[f3037,f197])).
fof(f5006,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X23),X25)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X23),X24))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X23),X25)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X25)))))) | s('type.ffi.ffi_state'(X23),'const.closSem.state_ffi_1'(s('type.closSem.state'(X23),X24))) != s('type.ffi.ffi_state'(X23),'const.closSem.state_ffi_1'(s('type.closSem.state'(X23),X25))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X24))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X25))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X23),X24),s('type.closSem.state'(X23),X25)))) )),
  inference(subsumption_resolution,[],[f4886,f246])).
fof(f4886,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X23),X25)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X23),X24))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X23),X25)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X25)))))) | s('type.ffi.ffi_state'(X23),'const.closSem.state_ffi_1'(s('type.closSem.state'(X23),X24))) != s('type.ffi.ffi_state'(X23),'const.closSem.state_ffi_1'(s('type.closSem.state'(X23),X25))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X24))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X23),X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X23),X24),s('type.closSem.state'(X23),X25)))) )),
  inference(superposition,[],[f258,f3074])).
fof(f5004,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) | s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X21))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X20))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21)))) )),
  inference(subsumption_resolution,[],[f4885,f246])).
fof(f4885,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),X21)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) | s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.closSem.state_ffi_1'(s('type.closSem.state'(X19),X21))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X20))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X19),X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f258,f3074])).
fof(f4992,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))))) ) | $spl100),
  inference(cnf_transformation,[],[f4992_D])).
fof(f4992_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4870,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X71)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X70))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X71)))))) | s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X70))) != s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X71))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X70))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f258,f262])).
fof(f4863,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X45),X46),s('type.closSem.state'(X45),X47)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X47)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X46))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X47)))))) | s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X46))) != s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X47))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X46))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X47))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X47))))) | s(bool,X48) != s(bool,X50)) )),
  inference(resolution,[],[f258,f203])).
fof(f4861,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X33),X34),s('type.closSem.state'(X33),X35)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X35)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X35)))))) | s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X34))) != s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X35))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X34))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X35)))))) )),
  inference(resolution,[],[f258,f202])).
fof(f4859,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X21),X22),s('type.closSem.state'(X21),X23)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X23)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X23))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X22))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X23)))))) )),
  inference(resolution,[],[f258,f201])).
fof(f4858,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17)))))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17))))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f258,f199])).
fof(f4856,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5)))))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) )),
  inference(resolution,[],[f258,f197])).
fof(f4991,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))))) ) | $spl98),
  inference(cnf_transformation,[],[f4991_D])).
fof(f4991_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4849,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X71)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X71)))))) | s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X70))) != s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X71))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X70))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X70))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f258,f262])).
fof(f4842,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X45),X46),s('type.closSem.state'(X45),X47)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X47)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X47)))))) | s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X46))) != s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X47))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X46))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X47))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X46))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X47))))) | s(bool,X48) != s(bool,X50)) )),
  inference(resolution,[],[f258,f203])).
fof(f4840,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X33),X34),s('type.closSem.state'(X33),X35)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X35)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X35)))))) | s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X34))) != s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X35))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X34))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X35)))))) )),
  inference(resolution,[],[f258,f202])).
fof(f4838,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X21),X22),s('type.closSem.state'(X21),X23)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X23)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X23))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X22))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X23)))))) )),
  inference(resolution,[],[f258,f201])).
fof(f4837,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17)))))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17))))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f258,f199])).
fof(f4835,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5)))))) )),
  inference(resolution,[],[f258,f197])).
fof(f4990,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5))))))) ) | $spl96),
  inference(cnf_transformation,[],[f4990_D])).
fof(f4990_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5))))) | ~p(s(bool,X7)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f4828,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X70))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X71)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X71)))))) | s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X70))) != s('type.ffi.ffi_state'(X69),'const.closSem.state_ffi_1'(s('type.closSem.state'(X69),X71))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X70))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X69),X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f258,f262])).
fof(f4821,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X45),X46),s('type.closSem.state'(X45),X47)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X46))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X47)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X45),X47)))))) | s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X46))) != s('type.ffi.ffi_state'(X45),'const.closSem.state_ffi_1'(s('type.closSem.state'(X45),X47))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X46))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X45),X47))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X46))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X45),X47))))) | s(bool,X48) != s(bool,X50)) )),
  inference(resolution,[],[f258,f203])).
fof(f4819,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X33),X34),s('type.closSem.state'(X33),X35)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X35)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X35)))))) | s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X34))) != s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X35))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X34))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X35)))))) )),
  inference(resolution,[],[f258,f202])).
fof(f4817,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X21),X22),s('type.closSem.state'(X21),X23)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X23)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.closSem.state_ffi_1'(s('type.closSem.state'(X21),X23))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X22))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X21),X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X22))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X23)))))) )),
  inference(resolution,[],[f258,f201])).
fof(f4816,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17)))))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17))))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f258,f199])).
fof(f4814,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X3),X4),s('type.closSem.state'(X3),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X5)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X3),X5)))))) | s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X4))) != s('type.ffi.ffi_state'(X3),'const.closSem.state_ffi_1'(s('type.closSem.state'(X3),X5))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X4))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X3),X5))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X4))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X3),X5)))))) )),
  inference(resolution,[],[f258,f197])).
fof(f4983,plain,(
  ( ! [X19,X17,X15,X16] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17)))))) | s(bool,X19) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17))) | p(s(bool,X19)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17))))))) ) | $spl94),
  inference(cnf_transformation,[],[f4983_D])).
fof(f4983_D,plain,(
  ( ! [X19,X17,X15,X16] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17)))))) | s(bool,X19) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17))) | p(s(bool,X19)) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17))))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4811,plain,(
  ( ! [X88,X87,X85,X86,X84] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X84),X85))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X84),X86)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X84),X85))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X84),X86)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X84),X85))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X84),X86)))))) | s('type.ffi.ffi_state'(X84),'const.closSem.state_ffi_1'(s('type.closSem.state'(X84),X85))) != s('type.ffi.ffi_state'(X84),'const.closSem.state_ffi_1'(s('type.closSem.state'(X84),X86))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X84),X85))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X84),X86))) | p(s(bool,X87)) | p(s(bool,X88)) | s(bool,X88) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X84),X85),s('type.closSem.state'(X84),X86)))) )),
  inference(resolution,[],[f258,f262])).
fof(f4803,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X57),X58))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X57),X59)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X57),X58))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X57),X59)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X57),X58))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X57),X59)))))) | s('type.ffi.ffi_state'(X57),'const.closSem.state_ffi_1'(s('type.closSem.state'(X57),X58))) != s('type.ffi.ffi_state'(X57),'const.closSem.state_ffi_1'(s('type.closSem.state'(X57),X59))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X57),X58))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X57),X59))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X57),X58),s('type.closSem.state'(X57),X59))) | s(bool,X61) != s(bool,X62)) )),
  inference(resolution,[],[f258,f203])).
fof(f4800,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X39),X40))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X39),X41)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X40))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X41)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X39),X40))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X39),X41)))))) | s('type.ffi.ffi_state'(X39),'const.closSem.state_ffi_1'(s('type.closSem.state'(X39),X40))) != s('type.ffi.ffi_state'(X39),'const.closSem.state_ffi_1'(s('type.closSem.state'(X39),X41))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X39),X40))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X39),X41))) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X39),X40),s('type.closSem.state'(X39),X41)))) )),
  inference(resolution,[],[f258,f202])).
fof(f4799,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X33),X35)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X35)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X34))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X33),X35)))))) | s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X34))) != s('type.ffi.ffi_state'(X33),'const.closSem.state_ffi_1'(s('type.closSem.state'(X33),X35))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X34))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X33),X35))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X33),X34),s('type.closSem.state'(X33),X35)))) )),
  inference(resolution,[],[f258,f201])).
fof(f4798,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X27),X28))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X27),X29)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X27),X28))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X27),X29)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X27),X28))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X27),X29)))))) | s('type.ffi.ffi_state'(X27),'const.closSem.state_ffi_1'(s('type.closSem.state'(X27),X28))) != s('type.ffi.ffi_state'(X27),'const.closSem.state_ffi_1'(s('type.closSem.state'(X27),X29))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X27),X28))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X27),X29))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X27),X28),s('type.closSem.state'(X27),X29))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f258,f199])).
fof(f4796,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X17)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X16))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X15),X17)))))) | s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X16))) != s('type.ffi.ffi_state'(X15),'const.closSem.state_ffi_1'(s('type.closSem.state'(X15),X17))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X16))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X15),X17))) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X15),X16),s('type.closSem.state'(X15),X17)))) )),
  inference(resolution,[],[f258,f197])).
fof(f4743,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2)) ) | $spl92),
  inference(cnf_transformation,[],[f4743_D])).
fof(f4743_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4704,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | s(bool,X66) != s(bool,X68) | p(s(bool,X68)) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X67) != s(bool,X70)) ) | $spl58),
  inference(resolution,[],[f3036,f262])).
fof(f4697,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl58),
  inference(resolution,[],[f3036,f203])).
fof(f4695,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl58),
  inference(resolution,[],[f3036,f202])).
fof(f4693,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl58),
  inference(resolution,[],[f3036,f201])).
fof(f4692,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl58),
  inference(resolution,[],[f3036,f199])).
fof(f4690,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl58),
  inference(resolution,[],[f3036,f197])).
fof(f4683,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,X22) | p(s(bool,X22))) ) | $spl90),
  inference(cnf_transformation,[],[f4683_D])).
fof(f4683_D,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,X22) | p(s(bool,X22))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4684,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,X25) | p(s(bool,X26)) | p(s(bool,X23)) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24)) ) | $spl62),
  inference(subsumption_resolution,[],[f4654,f246])).
fof(f4654,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,X25) | p(s(bool,X26)) | p(s(bool,X23)) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl62),
  inference(superposition,[],[f2964,f3074])).
fof(f4682,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X22)) | p(s(bool,X19)) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | $spl62),
  inference(subsumption_resolution,[],[f4653,f246])).
fof(f4653,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X22)) | p(s(bool,X19)) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl62),
  inference(superposition,[],[f2964,f3074])).
fof(f3095,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl66),
  inference(cnf_transformation,[],[f3095_D])).
fof(f3095_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4674,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl62),
  inference(subsumption_resolution,[],[f4639,f246])).
fof(f4639,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl62),
  inference(superposition,[],[f2964,f3074])).
fof(f4673,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X33)) | p(s(bool,X32)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl62),
  inference(subsumption_resolution,[],[f4638,f246])).
fof(f4638,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X33)) | p(s(bool,X32)) | s(bool,'T_0') = s(bool,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl62),
  inference(superposition,[],[f2964,f3074])).
fof(f4588,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X23),X24),s('type.closSem.state'(X23),X25)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X25)))))) )),
  inference(subsumption_resolution,[],[f4565,f246])).
fof(f4565,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X23),X24),s('type.closSem.state'(X23),X25)))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X24))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X23),X25)))))) )),
  inference(superposition,[],[f255,f3074])).
fof(f4586,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21)))))) )),
  inference(subsumption_resolution,[],[f4564,f246])).
fof(f4564,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X19),X20),s('type.closSem.state'(X19),X21)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X20))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X19),X21))))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f255,f3074])).
fof(f4582,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) | $spl86),
  inference(cnf_transformation,[],[f4582_D])).
fof(f4582_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4540,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X66),X67))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X66),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),X67),s('type.closSem.state'(X66),X68)))) )),
  inference(resolution,[],[f255,f262])).
fof(f4533,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X42),X43))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X42),X44)))))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f255,f203])).
fof(f4531,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X30),X31))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X30),X32)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X30),X31),s('type.closSem.state'(X30),X32)))) )),
  inference(resolution,[],[f255,f202])).
fof(f4529,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X18),X19))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X18),X20)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) )),
  inference(resolution,[],[f255,f201])).
fof(f4528,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X12),X13))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X12),X14)))))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f255,f199])).
fof(f4526,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) )),
  inference(resolution,[],[f255,f197])).
fof(f4579,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) | $spl84),
  inference(cnf_transformation,[],[f4579_D])).
fof(f4579_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4525,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X73) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f255,f262])).
fof(f4517,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X42),X43))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X42),X44))))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f255,f203])).
fof(f4514,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X24),X25),s('type.closSem.state'(X24),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X24),X25))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X24),X26)))))) )),
  inference(resolution,[],[f255,f202])).
fof(f4513,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X18),X19))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X18),X20)))))) )),
  inference(resolution,[],[f255,f201])).
fof(f4512,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X12),X13))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X12),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f255,f199])).
fof(f4510,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))))) )),
  inference(resolution,[],[f255,f197])).
fof(f4496,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f4085,f263])).
fof(f4085,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f262,f3074])).
fof(f4492,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f4062,f263])).
fof(f4062,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f203,f3074])).
fof(f4486,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f4054,f263])).
fof(f4054,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f202,f3074])).
fof(f4483,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f4049,f263])).
fof(f4049,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f201,f3074])).
fof(f4478,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f4042,f263])).
fof(f4042,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f199,f3074])).
fof(f4472,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4035,f263])).
fof(f4035,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f197,f3074])).
fof(f4460,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,X264) != s(bool,X265) | p(s(bool,X266)) | p(s(bool,X265)) | s(bool,X264) = s(bool,X267) | s(bool,'T_0') = s(bool,X267)) )),
  inference(subsumption_resolution,[],[f4010,f263])).
fof(f4010,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,X264) != s(bool,X265) | p(s(bool,X266)) | ~p(s(bool,'T_0')) | p(s(bool,X265)) | s(bool,X264) = s(bool,X267) | s(bool,'T_0') = s(bool,X267)) )),
  inference(superposition,[],[f262,f3074])).
fof(f4456,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X210) != s(bool,X211) | p(s(bool,X211)) | ~p(s(bool,X212)) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) )),
  inference(subsumption_resolution,[],[f3991,f263])).
fof(f3991,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X210) != s(bool,X211) | p(s(bool,X211)) | ~p(s(bool,X212)) | ~p(s(bool,'T_0')) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f203,f3074])).
fof(f4452,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,X165) != s(bool,X166) | p(s(bool,X167)) | ~p(s(bool,X168)) | p(s(bool,X166)) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f3982,f263])).
fof(f3982,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,X165) != s(bool,X166) | p(s(bool,X167)) | ~p(s(bool,'T_0')) | ~p(s(bool,X168)) | p(s(bool,X166)) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f202,f3074])).
fof(f4449,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X130) != s(bool,X131) | p(s(bool,X132)) | p(s(bool,X131)) | s(bool,X132) != s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f3975,f263])).
fof(f3975,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X130) != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,'T_0')) | p(s(bool,X131)) | s(bool,X132) != s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f201,f3074])).
fof(f4446,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,X108) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f3970,f263])).
fof(f3970,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | ~p(s(bool,X107)) | ~p(s(bool,'T_0')) | s(bool,X107) != s(bool,X108) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f199,f3074])).
fof(f4441,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) != s(bool,X66) | p(s(bool,X67)) | p(s(bool,X66)) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3962,f263])).
fof(f3962,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) != s(bool,X66) | p(s(bool,X67)) | ~p(s(bool,'T_0')) | p(s(bool,X66)) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f197,f3074])).
fof(f4427,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f3887,f263])).
fof(f3887,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f262,f3074])).
fof(f4423,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f3864,f263])).
fof(f3864,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f203,f3074])).
fof(f4417,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f3856,f263])).
fof(f3856,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f202,f3074])).
fof(f4414,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f3851,f263])).
fof(f3851,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f201,f3074])).
fof(f4409,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3844,f263])).
fof(f3844,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f199,f3074])).
fof(f4403,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3837,f263])).
fof(f3837,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f197,f3074])).
fof(f4391,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,X264) != s(bool,X265) | p(s(bool,X266)) | p(s(bool,X265)) | s(bool,X264) = s(bool,X267) | s(bool,'T_0') = s(bool,X267)) )),
  inference(subsumption_resolution,[],[f3812,f263])).
fof(f3812,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,X264) != s(bool,X265) | p(s(bool,X266)) | ~p(s(bool,'T_0')) | p(s(bool,X265)) | s(bool,X264) = s(bool,X267) | s(bool,'T_0') = s(bool,X267)) )),
  inference(superposition,[],[f262,f3074])).
fof(f4387,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X210) != s(bool,X211) | p(s(bool,X211)) | ~p(s(bool,X212)) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) )),
  inference(subsumption_resolution,[],[f3793,f263])).
fof(f3793,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X210) != s(bool,X211) | p(s(bool,X211)) | ~p(s(bool,X212)) | ~p(s(bool,'T_0')) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f203,f3074])).
fof(f4383,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,X165) != s(bool,X166) | p(s(bool,X167)) | ~p(s(bool,X168)) | p(s(bool,X166)) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f3784,f263])).
fof(f3784,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,X165) != s(bool,X166) | p(s(bool,X167)) | ~p(s(bool,'T_0')) | ~p(s(bool,X168)) | p(s(bool,X166)) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f202,f3074])).
fof(f4380,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X130) != s(bool,X131) | p(s(bool,X132)) | p(s(bool,X131)) | s(bool,X132) != s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f3777,f263])).
fof(f3777,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X130) != s(bool,X131) | p(s(bool,X132)) | ~p(s(bool,'T_0')) | p(s(bool,X131)) | s(bool,X132) != s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f201,f3074])).
fof(f4377,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,X108) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f3772,f263])).
fof(f3772,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,X106) | p(s(bool,X106)) | ~p(s(bool,X107)) | ~p(s(bool,'T_0')) | s(bool,X107) != s(bool,X108) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f199,f3074])).
fof(f4372,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) != s(bool,X66) | p(s(bool,X67)) | p(s(bool,X66)) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3764,f263])).
fof(f3764,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) != s(bool,X66) | p(s(bool,X67)) | ~p(s(bool,'T_0')) | p(s(bool,X66)) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f197,f3074])).
fof(f4371,plain,(
  ( ! [X372,X374,X373,X371] : (p(s(bool,X374)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373)))))) )),
  inference(subsumption_resolution,[],[f3742,f246])).
fof(f3742,plain,(
  ( ! [X372,X374,X373,X371] : (p(s(bool,X374)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,X374) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373)))))) )),
  inference(superposition,[],[f256,f3074])).
fof(f4370,plain,(
  ( ! [X368,X370,X367,X369] : (p(s(bool,X370)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369)))))) )),
  inference(subsumption_resolution,[],[f3741,f246])).
fof(f3741,plain,(
  ( ! [X368,X370,X367,X369] : (p(s(bool,X370)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,X370) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369)))))) )),
  inference(superposition,[],[f257,f3074])).
fof(f4332,plain,(
  ( ! [X372,X374,X373,X371] : (p(s(bool,X374)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373)))))) )),
  inference(subsumption_resolution,[],[f3625,f246])).
fof(f3625,plain,(
  ( ! [X372,X374,X373,X371] : (p(s(bool,X374)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X371),X372),s('type.closSem.state'(X371),X373)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X372))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X371),X373))))) | s(bool,'T_0') = s(bool,X374)) )),
  inference(superposition,[],[f256,f3074])).
fof(f4330,plain,(
  ( ! [X368,X370,X367,X369] : (p(s(bool,X370)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369)))))) )),
  inference(subsumption_resolution,[],[f3624,f246])).
fof(f3624,plain,(
  ( ! [X368,X370,X367,X369] : (p(s(bool,X370)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X367),X368),s('type.closSem.state'(X367),X369)))) | s(bool,'T_0') = s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X368))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X367),X369))))) | s(bool,'T_0') = s(bool,X370)) )),
  inference(superposition,[],[f257,f3074])).
fof(f3380,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) | $spl78),
  inference(cnf_transformation,[],[f3380_D])).
fof(f3380_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3350,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X66),X67))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X66),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),X67),s('type.closSem.state'(X66),X68)))) )),
  inference(resolution,[],[f256,f262])).
fof(f3343,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X42),X43))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X42),X44)))))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f256,f203])).
fof(f3341,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X30),X31))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X30),X32)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X30),X31),s('type.closSem.state'(X30),X32)))) )),
  inference(resolution,[],[f256,f202])).
fof(f3339,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X18),X19))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X18),X20)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) )),
  inference(resolution,[],[f256,f201])).
fof(f3338,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X12),X13))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X12),X14)))))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f256,f199])).
fof(f3336,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) )),
  inference(resolution,[],[f256,f197])).
fof(f3377,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) | $spl76),
  inference(cnf_transformation,[],[f3377_D])).
fof(f3377_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3335,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X73) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X70))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f256,f262])).
fof(f3327,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X42),X43))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X42),X44))))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f256,f203])).
fof(f3324,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X24),X25),s('type.closSem.state'(X24),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X24),X25))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X24),X26)))))) )),
  inference(resolution,[],[f256,f202])).
fof(f3323,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X18),X19))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X18),X20)))))) )),
  inference(resolution,[],[f256,f201])).
fof(f3322,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X12),X13))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X12),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f256,f199])).
fof(f3320,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))))) )),
  inference(resolution,[],[f256,f197])).
fof(f3312,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) | $spl74),
  inference(cnf_transformation,[],[f3312_D])).
fof(f3312_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3278,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X66),X67))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X66),X68)))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),X67),s('type.closSem.state'(X66),X68)))) )),
  inference(resolution,[],[f257,f262])).
fof(f3271,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X42),X43))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X42),X44)))))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f257,f203])).
fof(f3269,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X30),X31))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X30),X32)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X30),X31),s('type.closSem.state'(X30),X32)))) )),
  inference(resolution,[],[f257,f202])).
fof(f3267,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X18),X19))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X18),X20)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) )),
  inference(resolution,[],[f257,f201])).
fof(f3266,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X12),X13))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X12),X14)))))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f257,f199])).
fof(f3264,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) )),
  inference(resolution,[],[f257,f197])).
fof(f3309,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) | $spl72),
  inference(cnf_transformation,[],[f3309_D])).
fof(f3309_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))) | p(s(bool,X4))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3263,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X69),X70),s('type.closSem.state'(X69),X71)))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X73) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X70))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X69),X71)))))) )),
  inference(resolution,[],[f257,f262])).
fof(f3255,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X42),X43))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X42),X44))))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f257,f203])).
fof(f3252,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X24),X25),s('type.closSem.state'(X24),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X24),X25))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X24),X26)))))) )),
  inference(resolution,[],[f257,f202])).
fof(f3251,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X18),X19))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X18),X20)))))) )),
  inference(resolution,[],[f257,f201])).
fof(f3250,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X12),X13))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X12),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f257,f199])).
fof(f3248,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))) )),
  inference(resolution,[],[f257,f197])).
fof(f3146,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X67))) = s('type.ffi.ffi_state'(X66),'const.closSem.state_ffi_1'(s('type.closSem.state'(X66),X68))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),X67),s('type.closSem.state'(X66),X68)))) )),
  inference(resolution,[],[f254,f262])).
fof(f3139,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.ffi.ffi_state'(X42),'const.closSem.state_ffi_1'(s('type.closSem.state'(X42),X43))) = s('type.ffi.ffi_state'(X42),'const.closSem.state_ffi_1'(s('type.closSem.state'(X42),X44))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f254,f203])).
fof(f3137,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s('type.ffi.ffi_state'(X30),'const.closSem.state_ffi_1'(s('type.closSem.state'(X30),X31))) = s('type.ffi.ffi_state'(X30),'const.closSem.state_ffi_1'(s('type.closSem.state'(X30),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X30),X31),s('type.closSem.state'(X30),X32)))) )),
  inference(resolution,[],[f254,f202])).
fof(f3135,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.ffi.ffi_state'(X18),'const.closSem.state_ffi_1'(s('type.closSem.state'(X18),X19))) = s('type.ffi.ffi_state'(X18),'const.closSem.state_ffi_1'(s('type.closSem.state'(X18),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) )),
  inference(resolution,[],[f254,f201])).
fof(f3134,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.ffi.ffi_state'(X12),'const.closSem.state_ffi_1'(s('type.closSem.state'(X12),X13))) = s('type.ffi.ffi_state'(X12),'const.closSem.state_ffi_1'(s('type.closSem.state'(X12),X14))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f254,f199])).
fof(f3132,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) )),
  inference(resolution,[],[f254,f197])).
fof(f3130,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl68),
  inference(cnf_transformation,[],[f3130_D])).
fof(f3130_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3114,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) ) | $spl60),
  inference(resolution,[],[f2856,f262])).
fof(f3106,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl60),
  inference(resolution,[],[f2856,f203])).
fof(f3103,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl60),
  inference(resolution,[],[f2856,f202])).
fof(f3102,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl60),
  inference(resolution,[],[f2856,f201])).
fof(f3101,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl60),
  inference(resolution,[],[f2856,f199])).
fof(f3099,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl60),
  inference(resolution,[],[f2856,f197])).
fof(f3076,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X54) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f2660,f262])).
fof(f3068,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f2660,f203])).
fof(f3065,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f2660,f202])).
fof(f3064,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f2660,f201])).
fof(f3063,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f2660,f199])).
fof(f3061,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f2660,f197])).
fof(f2983,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X66),X67))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X66),X68))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X66),X67),s('type.closSem.state'(X66),X68)))) )),
  inference(resolution,[],[f253,f262])).
fof(f2976,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X42),X43))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X42),X44))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X42),X43),s('type.closSem.state'(X42),X44))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f253,f203])).
fof(f2974,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X30),X31))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X30),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X30),X31),s('type.closSem.state'(X30),X32)))) )),
  inference(resolution,[],[f253,f202])).
fof(f2972,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X18),X19))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X18),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X18),X19),s('type.closSem.state'(X18),X20)))) )),
  inference(resolution,[],[f253,f201])).
fof(f2971,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X12),X13))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X12),X14))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X12),X13),s('type.closSem.state'(X12),X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f253,f199])).
fof(f2969,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(X0),X1),s('type.closSem.state'(X0),X2)))) )),
  inference(resolution,[],[f253,f197])).
fof(f2964,plain,(
  ( ! [X0,X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) = s(bool,X1)) ) | $spl62),
  inference(cnf_transformation,[],[f2964_D])).
fof(f2964_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2935,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | s(bool,X54) = s(bool,X55) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f2659,f262])).
fof(f2927,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f2659,f203])).
fof(f2924,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f2659,f202])).
fof(f2923,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f2659,f201])).
fof(f2922,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f2659,f199])).
fof(f2920,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f2659,f197])).
fof(f2918,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | s(bool,X54) = s(bool,X55) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f2659,f262])).
fof(f2910,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f2659,f203])).
fof(f2907,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f2659,f202])).
fof(f2906,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f2659,f201])).
fof(f2905,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f2659,f199])).
fof(f2903,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f2659,f197])).
fof(f2856,plain,(
  ( ! [X2,X0] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl60),
  inference(cnf_transformation,[],[f2856_D])).
fof(f2856_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2657,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') = s(bool,X39) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f243,f262])).
fof(f2649,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f243,f203])).
fof(f2646,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f243,f202])).
fof(f2645,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f243,f201])).
fof(f2644,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f243,f199])).
fof(f2642,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f243,f197])).
fof(f2634,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f2634_D])).
fof(f2634_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2592,plain,(
  ( ! [X39,X37,X38] : (s(bool,'F_0') != s(bool,X37) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f242,f262])).
fof(f2586,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f242,f203])).
fof(f2584,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f242,f202])).
fof(f2582,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f242,f201])).
fof(f2581,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f242,f199])).
fof(f2579,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f242,f197])).
fof(f2577,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) | $spl56),
  inference(cnf_transformation,[],[f2577_D])).
fof(f2577_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2570,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f250,f262])).
fof(f2564,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f250,f203])).
fof(f2562,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f250,f202])).
fof(f2560,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f250,f201])).
fof(f2559,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f250,f199])).
fof(f2557,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f250,f197])).
fof(f2555,plain,(
  ~p(s(bool,'F_0')) | $spl55),
  inference(cnf_transformation,[],[f2555_D])).
fof(f2555_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f213,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.IMP_CLAUSES|split|4')).
fof(f2556,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f217,f263])).
fof(f217,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f148])).
fof(f231,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : ((p(s(bool,X9)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.AND_CLAUSES|split|3')).
fof(f235,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9] : ((p(s(bool,'F_0')) & p(s(bool,X9))) <=> p(s(bool,'F_0')))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f2485,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f2266,f252])).
fof(f252,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb',aHL_TRUTH)).
fof(f2266,plain,(
  ~p(s(bool,'T_0')) | $spl44),
  inference(backward_demodulation,[],[f2191,f250])).
fof(f2465,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f2451])).
fof(f2451,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl44),
  inference(superposition,[],[f560,f2191])).
fof(f2462,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f2454])).
fof(f2454,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl44),
  inference(superposition,[],[f321,f2191])).
fof(f2252,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6)))) ) | $spl52),
  inference(cnf_transformation,[],[f2252_D])).
fof(f2252_D,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2251,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2184,f261])).
fof(f2184,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6)))) ) | $spl6),
  inference(superposition,[],[f890,f261])).
fof(f890,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X28) != s(bool,'const.combin.K_2'(s(bool,X26),s(X25,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),'const.combin.K_1'(s(bool,X26))),s(X25,X27)))) ) | $spl6),
  inference(superposition,[],[f828,f179])).
fof(f179,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X19] : ! [X20,X21] : s(X6,'const.combin.K_2'(s(X6,X20),s(X19,X21))) = s(X6,happ(s(fun(X19,X6),'const.combin.K_1'(s(X6,X20))),s(X19,X21)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','ARITY_1_2_const.combin.K')).
fof(f2192,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X5),s(X6,X7))) | s(bool,'F_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X5))),s(X6,X7)))) ) | $spl46),
  inference(cnf_transformation,[],[f2192_D])).
fof(f2192_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X5),s(X6,X7))) | s(bool,'F_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X5))),s(X6,X7)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2202,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X64),s(X65,X66))) | s(bool,'F_0') = s(bool,happ(s(fun(X65,bool),'const.combin.K_1'(s(bool,X64))),s(X65,X66))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X63)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2113,f968])).
fof(f2113,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X64),s(X65,X66))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X63))) | s(bool,'F_0') = s(bool,happ(s(fun(X65,bool),'const.combin.K_1'(s(bool,X64))),s(X65,X66))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X63)))) ) | $spl6),
  inference(superposition,[],[f890,f966])).
fof(f966,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X39))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X39)))) )),
  inference(resolution,[],[f544,f243])).
fof(f2201,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X60),s(X61,X62))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),'const.combin.K_1'(s(bool,X60))),s(X61,X62))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X59)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2112,f1072])).
fof(f2112,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X60),s(X61,X62))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X59))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),'const.combin.K_1'(s(bool,X60))),s(X61,X62))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X59)))) ) | $spl6),
  inference(superposition,[],[f890,f1070])).
fof(f1070,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X39))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X39)))) )),
  inference(resolution,[],[f783,f243])).
fof(f2199,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X56),s(X57,X58))) | s(bool,'F_0') = s(bool,happ(s(fun(X57,bool),'const.combin.K_1'(s(bool,X56))),s(X57,X58))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X55)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2111,f931])).
fof(f2111,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X56),s(X57,X58))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X55),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,happ(s(fun(X57,bool),'const.combin.K_1'(s(bool,X56))),s(X57,X58))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X55)))) ) | $spl6),
  inference(superposition,[],[f890,f929])).
fof(f929,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X39),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X39)))) )),
  inference(resolution,[],[f543,f243])).
fof(f2197,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X52),s(X53,X54))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),'const.combin.K_1'(s(bool,X52))),s(X53,X54))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X51)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2110,f1036])).
fof(f2110,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X52),s(X53,X54))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X51),s('type.closSem.state'(sK0),sK3))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),'const.combin.K_1'(s(bool,X52))),s(X53,X54))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X51)))) ) | $spl6),
  inference(superposition,[],[f890,f1034])).
fof(f1034,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X39),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X39)))) )),
  inference(resolution,[],[f782,f243])).
fof(f2190,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X5),s(X6,X7))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X5))),s(X6,X7)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2104])).
fof(f2104,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X5),s(X6,X7))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X5))),s(X6,X7))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f890,f261])).
fof(f1084,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1))) | ~p(s(bool,X3)) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1)))) ) | $spl40),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1))) | ~p(s(bool,X3)) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1910,plain,(
  ( ! [X182,X183,X181] : (p(s(bool,X182)) | ~p(s(bool,X183)) | s(bool,X183) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X181))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X181)))) )),
  inference(subsumption_resolution,[],[f1853,f231])).
fof(f1853,plain,(
  ( ! [X182,X183,X181] : (p(s(bool,'F_0')) | p(s(bool,X182)) | ~p(s(bool,X183)) | s(bool,X183) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X181))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X181)))) )),
  inference(superposition,[],[f262,f1070])).
fof(f1899,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X136) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X134))) | s(bool,X135) != s(bool,X137) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X134)))) )),
  inference(subsumption_resolution,[],[f1822,f213])).
fof(f1822,plain,(
  ( ! [X134,X136,X135,X137] : (p(s(bool,'F_0')) | ~p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X136) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X134))) | s(bool,X135) != s(bool,X137) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X134)))) )),
  inference(superposition,[],[f203,f1070])).
fof(f1894,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X102))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X102)))) )),
  inference(subsumption_resolution,[],[f1814,f231])).
fof(f1814,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,'F_0')) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X102))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X102)))) )),
  inference(superposition,[],[f202,f1070])).
fof(f1889,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X82))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X82)))) )),
  inference(subsumption_resolution,[],[f1809,f231])).
fof(f1809,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'F_0')) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X82))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X82)))) )),
  inference(superposition,[],[f201,f1070])).
fof(f1886,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X54))) | s(bool,X55) != s(bool,X57) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X54)))) )),
  inference(subsumption_resolution,[],[f1802,f213])).
fof(f1802,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X54))) | s(bool,X55) != s(bool,X57) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X54)))) )),
  inference(superposition,[],[f199,f1070])).
fof(f1880,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X26))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X26)))) )),
  inference(subsumption_resolution,[],[f1795,f231])).
fof(f1795,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X26))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X26)))) )),
  inference(superposition,[],[f197,f1070])).
fof(f1049,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,X3)) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1)))) ) | $spl38),
  inference(cnf_transformation,[],[f1049_D])).
fof(f1049_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,X3)) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1576,plain,(
  ( ! [X182,X183,X181] : (p(s(bool,X182)) | ~p(s(bool,X183)) | s(bool,X183) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X181),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X181)))) )),
  inference(subsumption_resolution,[],[f1520,f231])).
fof(f1520,plain,(
  ( ! [X182,X183,X181] : (p(s(bool,'F_0')) | p(s(bool,X182)) | ~p(s(bool,X183)) | s(bool,X183) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X181),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X181)))) )),
  inference(superposition,[],[f262,f1034])).
fof(f1565,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X136) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X134),s('type.closSem.state'(sK0),sK3))) | s(bool,X135) != s(bool,X137) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X134)))) )),
  inference(subsumption_resolution,[],[f1489,f213])).
fof(f1489,plain,(
  ( ! [X134,X136,X135,X137] : (p(s(bool,'F_0')) | ~p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X136) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X134),s('type.closSem.state'(sK0),sK3))) | s(bool,X135) != s(bool,X137) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X134)))) )),
  inference(superposition,[],[f203,f1034])).
fof(f1560,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X102),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X102)))) )),
  inference(subsumption_resolution,[],[f1481,f231])).
fof(f1481,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,'F_0')) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X102),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X102)))) )),
  inference(superposition,[],[f202,f1034])).
fof(f1555,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X82),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X82)))) )),
  inference(subsumption_resolution,[],[f1476,f231])).
fof(f1476,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'F_0')) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X82),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X82)))) )),
  inference(superposition,[],[f201,f1034])).
fof(f1552,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X54),s('type.closSem.state'(sK0),sK3))) | s(bool,X55) != s(bool,X57) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X54)))) )),
  inference(subsumption_resolution,[],[f1469,f213])).
fof(f1469,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X54),s('type.closSem.state'(sK0),sK3))) | s(bool,X55) != s(bool,X57) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X54)))) )),
  inference(superposition,[],[f199,f1034])).
fof(f1546,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X26),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X26)))) )),
  inference(subsumption_resolution,[],[f1462,f231])).
fof(f1462,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X26),s('type.closSem.state'(sK0),sK3))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X26)))) )),
  inference(superposition,[],[f197,f1034])).
fof(f1439,plain,(
  $false | $spl17),
  inference(subsumption_resolution,[],[f346,f268])).
fof(f1416,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X179))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X179)))) )),
  inference(subsumption_resolution,[],[f1358,f231])).
fof(f1358,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,'F_0')) | p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X179))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X179)))) )),
  inference(superposition,[],[f262,f966])).
fof(f1405,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X132))) | s(bool,X133) != s(bool,X135) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X132)))) )),
  inference(subsumption_resolution,[],[f1327,f213])).
fof(f1327,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,'F_0')) | ~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X132))) | s(bool,X133) != s(bool,X135) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X132)))) )),
  inference(superposition,[],[f203,f966])).
fof(f1400,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X100))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X100)))) )),
  inference(subsumption_resolution,[],[f1319,f231])).
fof(f1319,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,'F_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X100))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X100)))) )),
  inference(superposition,[],[f202,f966])).
fof(f1395,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X80))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X80)))) )),
  inference(subsumption_resolution,[],[f1314,f231])).
fof(f1314,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,'F_0')) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X80))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X80)))) )),
  inference(superposition,[],[f201,f966])).
fof(f1392,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X52))) | s(bool,X53) != s(bool,X55) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X52)))) )),
  inference(subsumption_resolution,[],[f1307,f213])).
fof(f1307,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'F_0')) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X52))) | s(bool,X53) != s(bool,X55) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X52)))) )),
  inference(superposition,[],[f199,f966])).
fof(f1386,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X24))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X24)))) )),
  inference(subsumption_resolution,[],[f1300,f231])).
fof(f1300,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'F_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X24))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X24)))) )),
  inference(superposition,[],[f197,f966])).
fof(f1266,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f262])).
fof(f1256,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f268,f203])).
fof(f1253,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f202])).
fof(f1252,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f201])).
fof(f1251,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f268,f199])).
fof(f1249,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f268,f197])).
fof(f1226,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X179),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X179)))) )),
  inference(subsumption_resolution,[],[f1169,f231])).
fof(f1169,plain,(
  ( ! [X180,X181,X179] : (p(s(bool,'F_0')) | p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X179),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X179)))) )),
  inference(superposition,[],[f262,f929])).
fof(f1215,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X132),s('type.closSem.state'(sK0),sK3))) | s(bool,X133) != s(bool,X135) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X132)))) )),
  inference(subsumption_resolution,[],[f1138,f213])).
fof(f1138,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,'F_0')) | ~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X132),s('type.closSem.state'(sK0),sK3))) | s(bool,X133) != s(bool,X135) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X132)))) )),
  inference(superposition,[],[f203,f929])).
fof(f1210,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X100),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X100)))) )),
  inference(subsumption_resolution,[],[f1130,f231])).
fof(f1130,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,'F_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X102) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X100),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X100)))) )),
  inference(superposition,[],[f202,f929])).
fof(f1205,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X80),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X80)))) )),
  inference(subsumption_resolution,[],[f1125,f231])).
fof(f1125,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,'F_0')) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X80),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X80)))) )),
  inference(superposition,[],[f201,f929])).
fof(f1202,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X52),s('type.closSem.state'(sK0),sK3))) | s(bool,X53) != s(bool,X55) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X52)))) )),
  inference(subsumption_resolution,[],[f1118,f213])).
fof(f1118,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'F_0')) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X52),s('type.closSem.state'(sK0),sK3))) | s(bool,X53) != s(bool,X55) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X52)))) )),
  inference(superposition,[],[f199,f929])).
fof(f1196,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X24),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X24)))) )),
  inference(subsumption_resolution,[],[f1111,f231])).
fof(f1111,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'F_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X24),s('type.closSem.state'(sK0),sK3))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X24)))) )),
  inference(superposition,[],[f197,f929])).
fof(f1090,plain,(
  $false | $spl15),
  inference(subsumption_resolution,[],[f344,f269])).
fof(f1074,plain,(
  ( ! [X45,X43,X44] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X43))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X43)))) )),
  inference(resolution,[],[f783,f262])).
fof(f1063,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X29))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f783,f203])).
fof(f1061,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X21)))) )),
  inference(resolution,[],[f783,f202])).
fof(f1059,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X13)))) )),
  inference(resolution,[],[f783,f201])).
fof(f1058,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X9))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f783,f199])).
fof(f1056,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1)))) )),
  inference(resolution,[],[f783,f197])).
fof(f1038,plain,(
  ( ! [X45,X43,X44] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X43))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X43),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f782,f262])).
fof(f1027,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X29),s('type.closSem.state'(sK0),sK3))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f782,f203])).
fof(f1025,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X21),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f782,f202])).
fof(f1023,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X13),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f782,f201])).
fof(f1022,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X9),s('type.closSem.state'(sK0),sK3))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f782,f199])).
fof(f1020,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f782,f197])).
fof(f1005,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f262])).
fof(f995,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f269,f203])).
fof(f992,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f202])).
fof(f991,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f201])).
fof(f990,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f269,f199])).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f269,f197])).
fof(f970,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X43))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X43)))) )),
  inference(resolution,[],[f544,f262])).
fof(f959,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X29))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f544,f203])).
fof(f957,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X21)))) )),
  inference(resolution,[],[f544,f202])).
fof(f955,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X13)))) )),
  inference(resolution,[],[f544,f201])).
fof(f954,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X9))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f544,f199])).
fof(f952,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK3),s('type.closSem.state'(sK0),X1)))) )),
  inference(resolution,[],[f544,f197])).
fof(f933,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X43))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X43),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f262])).
fof(f922,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X29))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X29),s('type.closSem.state'(sK0),sK3))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f543,f203])).
fof(f920,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X21),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f202])).
fof(f918,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X13),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f201])).
fof(f917,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X9))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X9),s('type.closSem.state'(sK0),sK3))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f543,f199])).
fof(f915,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f543,f197])).
fof(f876,plain,(
  $false | $spl23),
  inference(trivial_inequality_removal,[],[f875])).
fof(f875,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) | $spl23),
  inference(forward_demodulation,[],[f352,f265])).
fof(f723,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f713])).
fof(f713,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f296,f285])).
fof(f722,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f296])).
fof(f297,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f297_D])).
fof(f297_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f690,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) )),
  inference(forward_demodulation,[],[f689,f285])).
fof(f689,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f617,f262])).
fof(f617,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f203,f285])).
fof(f686,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(subsumption_resolution,[],[f612,f262])).
fof(f612,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(superposition,[],[f203,f285])).
fof(f679,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X52) != s(bool,X53)) )),
  inference(subsumption_resolution,[],[f602,f203])).
fof(f602,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,X51) | p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X52) != s(bool,X53)) )),
  inference(superposition,[],[f201,f285])).
fof(f674,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) )),
  inference(forward_demodulation,[],[f673,f285])).
fof(f673,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(subsumption_resolution,[],[f597,f262])).
fof(f597,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(superposition,[],[f199,f285])).
fof(f670,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(subsumption_resolution,[],[f593,f262])).
fof(f593,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(superposition,[],[f199,f285])).
fof(f667,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X8)) )),
  inference(subsumption_resolution,[],[f587,f203])).
fof(f587,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X6)) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X8)) )),
  inference(superposition,[],[f197,f285])).
fof(f540,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f509,f299])).
fof(f509,plain,(
  ( ! [X43] : (p(s(bool,X43))) ) | $spl10),
  inference(resolution,[],[f299,f224])).
fof(f539,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f508,f299])).
fof(f508,plain,(
  ( ! [X42] : (p(s(bool,X42))) ) | $spl10),
  inference(resolution,[],[f299,f221])).
fof(f538,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f507,f299])).
fof(f507,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl10),
  inference(resolution,[],[f299,f218])).
fof(f537,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f505,f299])).
fof(f505,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f299,f215])).
fof(f215,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f148])).
fof(f513,plain,(
  $false | $spl10),
  inference(resolution,[],[f299,f251])).
fof(f251,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.TRUTH')).
fof(f514,plain,(
  $false | $spl10),
  inference(resolution,[],[f299,f252])).
fof(f515,plain,(
  $false | $spl10),
  inference(resolution,[],[f299,f263])).
fof(f518,plain,(
  $false | $spl10),
  inference(resolution,[],[f299,f172])).
fof(f475,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X12))) ) | $spl8),
  inference(superposition,[],[f298,f247])).
fof(f476,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f453])).
fof(f453,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl8),
  inference(superposition,[],[f298,f249])).
fof(f480,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f452,f298])).
fof(f452,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f298,f261])).
fof(f479,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(subsumption_resolution,[],[f451,f298])).
fof(f451,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f298,f261])).
fof(f436,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f417])).
fof(f417,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | p(s(bool,X16))) ) | $spl2),
  inference(superposition,[],[f295,f243])).
fof(f437,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f416])).
fof(f416,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f295,f245])).
fof(f245,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f443,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f413,f295])).
fof(f413,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f295,f261])).
fof(f441,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl2),
  inference(subsumption_resolution,[],[f412,f295])).
fof(f412,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f295,f261])).
fof(f408,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f382,f300])).
fof(f382,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f300,f213])).
fof(f378,plain,(
  $false | $spl12),
  inference(resolution,[],[f300,f250])).
fof(f388,plain,(
  $false | $spl12),
  inference(resolution,[],[f300,f173])).
fof(f353,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) ) | $spl24),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f323,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f173,f262])).
fof(f312,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f173,f203])).
fof(f310,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f173,f202])).
fof(f354,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) ) | $spl26),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f309,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f173,f201])).
fof(f308,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f173,f201])).
fof(f307,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f173,f199])).
fof(f306,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f173,f197])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))) )),
  inference(resolution,[],[f173,f197])).
fof(f344,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | $spl15),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f346,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | $spl17),
  inference(cnf_transformation,[],[f346_D])).
fof(f346_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f348,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | $spl19),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f350,plain,(
  s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | $spl21),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f352,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) | $spl23),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f342,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(forward_demodulation,[],[f341,f179])).
fof(f341,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK1),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(forward_demodulation,[],[f340,f179])).
fof(f340,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(forward_demodulation,[],[f339,f174])).
fof(f174,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X16] : ! [X8,X17] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X17),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X8))))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_86_prove_ltb')).
fof(f339,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK3))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3))))))),
  inference(forward_demodulation,[],[f338,f174])).
fof(f338,plain,(
  s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f337,f178])).
fof(f337,plain,(
  s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),sK3))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f336,f178])).
fof(f336,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f335,f176])).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X16] : ! [X8,X17] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X16),X8))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X16),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_69_prove_ltb')).
fof(f335,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f334,f176])).
fof(f334,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f333,f177])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16] : ! [X8,X17] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X16),X8))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X16),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_65_prove_ltb')).
fof(f333,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f332,f177])).
fof(f332,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK2))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f331,f175])).
fof(f175,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X16] : ! [X8,X17] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X16),X8))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X16),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X17),s('type.closSem.state'(X16),X8)))))),
  file('thm.clos_numberProof.clos_numberProof_|49|4_prove_ltb','thm.closSem.closSem_82_prove_ltb')).
fof(f331,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(forward_demodulation,[],[f304,f175])).
fof(f304,plain,(
  ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.pair.prod'('type.num.num','type.closLang.exp'),fun('type.pair.prod'('type.num.num','type.closLang.exp'),bool)),'const.clos_numberProof.code_rel_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.closSem.v'),fun('type.option.option'('type.closSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | ~p(s(bool,'const.finite_map.fmap_rel_3'(s(fun('type.closSem.ref'('type.closSem.v'),fun('type.closSem.ref'('type.closSem.v'),bool)),'const.closProps.ref_rel_1'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3))))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK2))))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK1))),s('type.closSem.state'(sK0),sK3)))))),
  inference(resolution,[],[f173,f258])).
fof(f294,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) ) | $spl0),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f287,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f262])).
fof(f277,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f172,f203])).
fof(f276,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f203])).
fof(f274,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f202])).
fof(f273,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f201])).
fof(f272,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f172,f199])).
fof(f271,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f199])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_numberProof.state_rel_2'(s('type.closSem.state'(sK0),sK2),s('type.closSem.state'(sK0),sK3)))) )),
  inference(resolution,[],[f172,f197])).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 151
% Passive clauses: 1264
% Generated clauses: 11990
% Final active clauses: 90
% Final passive clauses: 958
% Input formulas: 55
% Initial clauses: 90
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 539
% Trivial inequalities: 277
% Fw subsumption resolutions: 2458
% Fw demodulations: 546
% Bw demodulations: 51
% 
% Simple tautologies: 418
% Equational tautologies: 40
% Forward subsumptions: 5495
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 1012
% Factoring: 1
% Forward superposition: 2052
% Backward superposition: 4615
% Self superposition: 121
% Equality factoring: 102
% Equality resolution: 32
% 
% Split clauses: 903
% Split components: 1990
% SAT solver clauses: 960
% SAT solver unit clauses: 25
% SAT solver binary clauses: 483
% 
% Sat splits: 903
% Sat splitting refutations: 25
% 
% Memory used [KB]: 5245
% Time elapsed: 0.338 s
% ------------------------------
% ------------------------------
% Success in time 12.081 s
