Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvpSem
% (54747)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 3517
% Generated clauses: 14128
% Final active clauses: 88
% Final passive clauses: 3303
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 653
% Trivial inequalities: 161
% Fw subsumption resolutions: 1704
% Fw demodulations: 917
% Bw demodulations: 2
% 
% Simple tautologies: 227
% Equational tautologies: 79
% Forward subsumptions: 5805
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 182
% Factoring: 3
% Forward superposition: 4914
% Backward superposition: 4956
% Self superposition: 277
% Equality factoring: 199
% Equality resolution: 15
% 
% Split clauses: 330
% Split components: 683
% SAT solver clauses: 1029
% SAT solver unit clauses: 18
% SAT solver binary clauses: 182
% 
% Sat splits: 330
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1028
% TWLsolver calls for satisfiability: 138
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvpSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvpSem
% SZS output start Proof for thm.bvpSem
fof(f3368,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3367,f2083,f2084,f76,f106,f65,f93,f91,f2210,f3313,f3231,f3232,f78,f108,f72,f102,f69,f99,f64,f92,f90])).
fof(f90,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) | $spl0),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f92,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) = s('type.sptree.spt'('type.bvlSem.v'),sK1) | $spl2),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) = s('type.sptree.spt'('type.bvlSem.v'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f64,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) = s('type.sptree.spt'('type.bvlSem.v'),sK1) | s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1)))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) | s('type.sptree.spt'('type.bvlSem.v'),sK0) = s('type.sptree.spt'('type.bvlSem.v'),sK1)) & (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) | s('type.sptree.spt'('type.bvlSem.v'),sK0) != s('type.sptree.spt'('type.bvlSem.v'),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f59])).
fof(f59,plain,(
  ? [X0,X1] : ((s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))) | s('type.sptree.spt'('type.bvlSem.v'),X0) = s('type.sptree.spt'('type.bvlSem.v'),X1)) & (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))) | s('type.sptree.spt'('type.bvlSem.v'),X0) != s('type.sptree.spt'('type.bvlSem.v'),X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ? [X0,X1] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))) <~> s('type.sptree.spt'('type.bvlSem.v'),X0) = s('type.sptree.spt'('type.bvlSem.v'),X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))) <=> s('type.sptree.spt'('type.bvlSem.v'),X0) = s('type.sptree.spt'('type.bvlSem.v'),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X6))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X7))) <=> s('type.sptree.spt'('type.bvlSem.v'),X6) = s('type.sptree.spt'('type.bvlSem.v'),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X6))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X7))) <=> s('type.sptree.spt'('type.bvlSem.v'),X6) = s('type.sptree.spt'('type.bvlSem.v'),X7))),
  file('thm.bvpSem.stack_|49||49||split|0','thm.bvpSem.stack_|49||49||split|0')).
fof(f99,plain,(
  ( ! [X0] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),X0)))) ) | $spl10),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  ( ! [X0] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),X0)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f69,plain,(
  ( ! [X0] : (s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X0))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X8))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),X8)))),
  file('thm.bvpSem.stack_|49||49||split|0','ARITY_0_1_const.bvpSem.Env')).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f72,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvpSem.stack_|49||49||split|0','thm.ind_type.FCONS|split|0')).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) | $spl28),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ! [X8,X12,X20] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X12),s('type.num.num',X20))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X12))),s('type.num.num',X20)))),
  file('thm.bvpSem.stack_|49||49||split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f3232,plain,(
  ( ! [X4,X5] : (s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK0),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5))) = s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5)))) ) | $spl596),
  inference(cnf_transformation,[],[f3232_D])).
fof(f3232_D,plain,(
  ( ! [X4,X5] : (s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK0),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5))) = s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5)))) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f3231,plain,(
  ( ! [X4,X5] : (s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK0),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5))) = s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5)))) ) | ($spl28 | $spl370)),
  inference(forward_demodulation,[],[f3123,f108])).
fof(f3123,plain,(
  ( ! [X4,X5] : (s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4),s('type.num.num',X5))) = s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),happ(s(fun(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),fun('type.num.num','type.sptree.spt'('type.bvlSem.v'))),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),fun('type.num.num','type.sptree.spt'('type.bvlSem.v')))),'const.ind_type.FCONS_0'),s('type.sptree.spt'('type.bvlSem.v'),sK0))),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X4))),s('type.num.num',X5)))) ) | ($spl28 | $spl370)),
  inference(superposition,[],[f108,f2084])).
fof(f3313,plain,(
  $false | ($spl3 | $spl16 | $spl596)),
  inference(subsumption_resolution,[],[f3312,f93])).
fof(f3312,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) = s('type.sptree.spt'('type.bvlSem.v'),sK1) | ($spl16 | $spl596)),
  inference(forward_demodulation,[],[f3293,f102])).
fof(f3293,plain,(
  ( ! [X2] : (s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),'const.ind_type.FCONS_3'(s('type.sptree.spt'('type.bvlSem.v'),sK0),s(fun('type.num.num','type.sptree.spt'('type.bvlSem.v')),X2),s('type.num.num','const.num.|48|_0')))) ) | ($spl16 | $spl596)),
  inference(superposition,[],[f3232,f102])).
fof(f2210,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) | ($spl10 | $spl370)),
  inference(forward_demodulation,[],[f2172,f99])).
fof(f2172,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) = s('type.bvpSem.stack',happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.bvpSem.stack'),'const.bvpSem.Env_0'),s('type.sptree.spt'('type.bvlSem.v'),sK0))) | ($spl10 | $spl370)),
  inference(superposition,[],[f99,f2084])).
fof(f91,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) | $spl1),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f93,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) != s('type.sptree.spt'('type.bvlSem.v'),sK1) | $spl3),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) != s('type.sptree.spt'('type.bvlSem.v'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f65,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK0) != s('type.sptree.spt'('type.bvlSem.v'),sK1) | s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK1)))),
  inference(cnf_transformation,[],[f60])).
fof(f106,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s('type.sptree.spt'('type.bvlSem.v'),X1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X3)))) ) | $spl24),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s('type.sptree.spt'('type.bvlSem.v'),X1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X3)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s('type.sptree.spt'('type.bvlSem.v'),X1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X3)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s('type.sptree.spt'('type.bvlSem.v'),X1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X3)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s('type.sptree.spt'('type.bvlSem.v'),X1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X1))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X3)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11] : s(X9,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X9),X10),s('type.sptree.spt'('type.bvlSem.v'),X6))) = s(X9,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),X6))),s(fun('type.sptree.spt'('type.bvlSem.v'),X9),X10),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X9)),X11)))),
  file('thm.bvpSem.stack_|49||49||split|0','thm.bvpSem.stack_case_def|split|0')).
fof(f2084,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK1)))) ) | $spl370),
  inference(cnf_transformation,[],[f2084_D])).
fof(f2084_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK1)))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f2083,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK0))) = s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK1)))) ) | ($spl0 | $spl24)),
  inference(forward_demodulation,[],[f2012,f106])).
fof(f2012,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s('type.sptree.spt'('type.bvlSem.v'),sK1))) = s(X0,'const.bvpSem.stack_CASE_3'(s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))),s(fun('type.sptree.spt'('type.bvlSem.v'),X0),X1),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.num.num',X0)),X2)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f106,f90])).
fof(f3367,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f3366])).
fof(f3366,plain,(
  s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) != s('type.bvpSem.stack','const.bvpSem.Env_1'(s('type.sptree.spt'('type.bvlSem.v'),sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f91,f92])).
% SZS output end Proof for thm.bvpSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 48
% Passive clauses: 311
% Generated clauses: 3252
% Final active clauses: 39
% Final passive clauses: 225
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 175
% Trivial inequalities: 15
% Fw subsumption resolutions: 88
% Fw demodulations: 207
% Bw demodulations: 41
% Condensations: 304
% 
% Simple tautologies: 48
% Equational tautologies: 29
% Forward subsumptions: 1534
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 69
% Factoring: 1
% Forward superposition: 935
% Backward superposition: 942
% Self superposition: 80
% Equality factoring: 52
% Equality resolution: 1
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 567
% SAT solver unit clauses: 39
% SAT solver binary clauses: 25
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 8
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 565
% TWLsolver calls for satisfiability: 110
% 
% Memory used [KB]: 5245
% Time elapsed: 0.170 s
% ------------------------------
% ------------------------------
% Success in time 0.79 s
