Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.infer_eComplete
% (43478)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 578
% Passive clauses: 3691
% Generated clauses: 14312
% Final active clauses: 572
% Final passive clauses: 2865
% Input formulas: 36
% Initial clauses: 153
% 
% Duplicate literals: 640
% Trivial inequalities: 39
% Fw subsumption resolutions: 1898
% Fw demodulations: 273
% Bw demodulations: 6
% 
% Simple tautologies: 232
% Equational tautologies: 63
% Forward subsumptions: 5900
% 
% Binary resolution: 4467
% Factoring: 1024
% Forward superposition: 1757
% Backward superposition: 3453
% Self superposition: 109
% Equality factoring: 108
% Equality resolution: 6
% 
% Split clauses: 1411
% Split components: 2988
% SAT solver clauses: 1731
% SAT solver unit clauses: 34
% SAT solver binary clauses: 176
% 
% Sat splits: 1411
% 
% TWLsolver clauses: 1724
% TWLsolver calls for satisfiability: 323
% 
% Memory used [KB]: 6652
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.infer_eComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.infer_eComplete
% SZS output start Proof for thm.infer_eComplete
fof(f1755,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1754,f209,f368,f1716,f1718,f210,f376,f486,f492,f208,f370,f225,f384,f1675,f1676,f358,f474,f484,f488,f206,f374,f1751,f1752,f283,f405,f485,f490,f207,f372])).
fof(f372,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) | $spl4),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f207,plain,(
  p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))))),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15)))))) & s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',sK15))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f91])).
fof(f91,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) & s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))))))) & s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))))))) => s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X8),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t',X9))))))) => s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.infer_t.infer_t',X9))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t',X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X8),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t',X9))))))) => s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.infer_t.infer_t',X9))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t',X9))))),
  file('thm.infer_eComplete.infer_eComplete_35_prove_ltb','thm.infer_eComplete.infer_eComplete_35_prove_ltb')).
fof(f490,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) | $spl186),
  inference(cnf_transformation,[],[f490_D])).
fof(f490_D,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f485,plain,(
  s(bool,'T_0') = s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) | ($spl4 | $spl14)),
  inference(resolution,[],[f384,f372])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2)))))) ) | $spl38),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) => s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20))))) => s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.infer_t.infer_t',X11))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19),s('type.infer_t.infer_t',X11))))))),
  file('thm.infer_eComplete.infer_eComplete_35_prove_ltb','thm.unify.t_walkstar_SUBMAP')).
fof(f1752,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) | $spl644),
  inference(cnf_transformation,[],[f1752_D])).
fof(f1752_D,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f1751,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) | ($spl8 | $spl38 | $spl184 | $spl186)),
  inference(subsumption_resolution,[],[f1750,f376])).
fof(f1750,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) | ($spl8 | $spl38 | $spl184 | $spl186)),
  inference(forward_demodulation,[],[f1749,f488])).
fof(f1749,plain,(
  ( ! [X0] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) | ($spl8 | $spl38 | $spl186)),
  inference(subsumption_resolution,[],[f1734,f376])).
fof(f1734,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',X0)))))) ) | ($spl38 | $spl186)),
  inference(superposition,[],[f405,f490])).
fof(f374,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) | $spl6),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f206,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))))),
  inference(cnf_transformation,[],[f138])).
fof(f488,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) | $spl184),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f484,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12))) | ($spl6 | $spl14)),
  inference(resolution,[],[f384,f374])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))) ) | $spl176),
  inference(cnf_transformation,[],[f474_D])).
fof(f474_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1))))),
  inference(rectify,[],[f119])).
fof(f119,plain,(
  ! [X0,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2)))) | s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X2,X3] : ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2))))) | s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2))))) => s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13,X14,X11,X7] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X11))))) => s('type.infer_t.infer_t',X11) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t',X11))))),
  file('thm.infer_eComplete.infer_eComplete_35_prove_ltb','thm.inferProps.t_walkstar_no_vars')).
fof(f1676,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X1)))) ) | $spl588),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X1)))) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f1675,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X1)))) ) | ($spl8 | $spl176 | $spl184)),
  inference(subsumption_resolution,[],[f1645,f376])).
fof(f1645,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',X1)))) ) | ($spl176 | $spl184)),
  inference(superposition,[],[f474,f488])).
fof(f384,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f384_D])).
fof(f384_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f225,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.infer_eComplete.infer_eComplete_35_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f370,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15)))))) | $spl2),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f208,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))))),
  inference(cnf_transformation,[],[f138])).
fof(f492,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) | $spl188),
  inference(cnf_transformation,[],[f492_D])).
fof(f492_D,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f486,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK14),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) | ($spl2 | $spl14)),
  inference(resolution,[],[f384,f370])).
fof(f376,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f376_D])).
fof(f376_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f210,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_35_prove_ltb','thm.bool.TRUTH')).
fof(f1718,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) | $spl616),
  inference(cnf_transformation,[],[f1718_D])).
fof(f1718_D,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f1716,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) | ($spl8 | $spl188 | $spl588)),
  inference(subsumption_resolution,[],[f1701,f376])).
fof(f1701,plain,(
  ~p(s(bool,'T_0')) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))))) | ($spl188 | $spl588)),
  inference(superposition,[],[f1676,f492])).
fof(f368,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',sK15))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) | $spl1),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',sK15))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f209,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',sK15))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15)))),
  inference(cnf_transformation,[],[f138])).
fof(f1754,plain,(
  $false | ($spl1 | $spl616 | $spl644)),
  inference(subsumption_resolution,[],[f1753,f368])).
fof(f1753,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK12),s('type.infer_t.infer_t',sK15))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK13),s('type.infer_t.infer_t',sK15))) | ($spl616 | $spl644)),
  inference(backward_demodulation,[],[f1752,f1718])).
% SZS output end Proof for thm.infer_eComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 176
% Passive clauses: 300
% Generated clauses: 1530
% Final active clauses: 173
% Final passive clauses: 123
% Input formulas: 36
% Initial clauses: 153
% 
% Duplicate literals: 73
% Trivial inequalities: 2
% Fw subsumption resolutions: 139
% Fw demodulations: 19
% Bw demodulations: 4
% Condensations: 23
% Global subsumptions: 13
% 
% Simple tautologies: 152
% Equational tautologies: 3
% Forward subsumptions: 437
% 
% Binary resolution: 588
% Factoring: 44
% Forward superposition: 103
% Backward superposition: 48
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 549
% SAT solver unit clauses: 170
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 1
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 496
% TWLsolver calls for satisfiability: 1507
% 
% Memory used [KB]: 1918
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
% Success in time 0.65 s
