Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modProps
% (134050)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 793
% Passive clauses: 7653
% Generated clauses: 23884
% Final active clauses: 781
% Final passive clauses: 6571
% Input formulas: 78
% Initial clauses: 201
% 
% Duplicate literals: 309
% Trivial inequalities: 119
% Fw subsumption resolutions: 1750
% Fw demodulations: 1786
% Bw demodulations: 11
% 
% Simple tautologies: 227
% Equational tautologies: 57
% Forward subsumptions: 11735
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 7792
% Factoring: 1308
% Forward superposition: 5292
% Backward superposition: 5141
% Self superposition: 46
% Equality factoring: 26
% Equality resolution: 21
% 
% Split clauses: 254
% Split components: 509
% SAT solver clauses: 481
% SAT solver unit clauses: 11
% SAT solver binary clauses: 70
% 
% Sat splits: 254
% 
% TWLsolver clauses: 481
% TWLsolver calls for satisfiability: 326
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modProps
% SZS output start Proof for thm.modProps
fof(f11616,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f11586,f298,f506,f669,f673,f303,f516,f374,f560,f11469,f11472,f6546,f9385,f297,f508,f310,f525,f377,f562,f378,f563])).
fof(f563,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) != s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3)))) ) | $spl84),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) != s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f378,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3))) | s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) != s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5)))))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) != s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) = s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X1),s('type.modSem.state'(X0),X2),s('type.list.list'('type.modLang.exp'),X3))) = s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X5))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X3))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X19] : ! [X10,X9,X20,X8,X7] : (s('type.pair.prod'('type.modSem.state'(X19),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X10),s('type.modSem.state'(X19),X9),s('type.list.list'('type.modLang.exp'),X20))) = s('type.pair.prod'('type.modSem.state'(X19),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X19),X8),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X7))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X7))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),X20))))),
  file('thm.modProps.evaluate_sing|split|0','thm.modProps.evaluate_length|split|0')).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) ) | $spl82),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X17,X15] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X15))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))),
  file('thm.modProps.evaluate_sing|split|0','thm.list.LENGTH|split|1')).
fof(f525,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f525_D])).
fof(f525_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f310,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.modProps.evaluate_sing|split|0','thm.list.LENGTH|split|0')).
fof(f508,plain,(
  s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) = s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',sK16),s('type.modSem.state'(sK12),sK15),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',sK17),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) | $spl2),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) = s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',sK16),s('type.modSem.state'(sK12),sK15),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',sK17),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f297,plain,(
  s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) = s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',sK16),s('type.modSem.state'(sK12),sK15),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',sK17),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0')))))),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) = s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',sK16),s('type.modSem.state'(sK12),sK15),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',sK17),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) & ! [X6] : s('type.list.list'('type.modSem.v'),sK13) != s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17])],[f179])).
fof(f179,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) = s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X4),s('type.modSem.state'(X0),X3),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',X5),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) & ! [X6] : s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) = s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X4),s('type.modSem.state'(X0),X3),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',X5),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) => ? [X6] : s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) = s('type.pair.prod'('type.modSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X4),s('type.modSem.state'(X0),X3),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',X5),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) => ? [X6] : s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11] : (s('type.pair.prod'('type.modSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X6),X8),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X7))))) = s('type.pair.prod'('type.modSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X10),s('type.modSem.state'(X6),X9),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',X11),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) => ? [X12] : s('type.list.list'('type.modSem.v'),X7) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X12),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11] : (s('type.pair.prod'('type.modSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(X6),X8),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X7))))) = s('type.pair.prod'('type.modSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.evaluate_3'(s('type.modSem.environment',X10),s('type.modSem.state'(X6),X9),s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',X11),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) => ? [X12] : s('type.list.list'('type.modSem.v'),X7) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X12),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))),
  file('thm.modProps.evaluate_sing|split|0','thm.modProps.evaluate_sing|split|0')).
fof(f9385,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) != s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),X0),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X1)))) ) | $spl2746),
  inference(cnf_transformation,[],[f9385_D])).
fof(f9385_D,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) != s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),X0),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X1)))) ) <=> ~$spl2746),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2746])])).
fof(f6546,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X1))) | s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) != s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),X0),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1)))))) ) | ($spl2 | $spl16 | $spl82 | $spl84)),
  inference(forward_demodulation,[],[f6545,f525])).
fof(f6545,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),'const.list.NIL_0'))))) | s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) != s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),X0),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1)))))) ) | ($spl2 | $spl82 | $spl84)),
  inference(forward_demodulation,[],[f6541,f562])).
fof(f6541,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),sK13))))) != s('type.pair.prod'('type.modSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.pair.,_2'(s('type.modSem.state'(sK12),X0),s('type.semanticPrimitives.result'('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.modSem.v'),X1))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modLang.exp'),'const.list.CONS_2'(s('type.modLang.exp',sK17),s('type.list.list'('type.modLang.exp'),'const.list.NIL_0')))))) ) | ($spl2 | $spl84)),
  inference(superposition,[],[f563,f508])).
fof(f11472,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK13))) | $spl3106),
  inference(cnf_transformation,[],[f11472_D])).
fof(f11472_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK13))) <=> ~$spl3106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3106])])).
fof(f11469,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.modSem.v'),sK13))) | $spl2746),
  inference(equality_resolution,[],[f9385])).
fof(f560,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl78),
  inference(cnf_transformation,[],[f560_D])).
fof(f560_D,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f374,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X18] : (s('type.list.list'(X6),X18) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,'const.list.HD_1'(s('type.list.list'(X6),X18))),s('type.list.list'(X6),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X18))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  file('thm.modProps.evaluate_sing|split|0','thm.misc.SING_HD|split|0')).
fof(f516,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl10),
  inference(cnf_transformation,[],[f516_D])).
fof(f516_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f303,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.modProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|13')).
fof(f673,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl272),
  inference(cnf_transformation,[],[f673_D])).
fof(f673_D,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f669,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) ) | ($spl10 | $spl78)),
  inference(backward_demodulation,[],[f516,f560])).
fof(f506,plain,(
  ( ! [X6] : (s('type.list.list'('type.modSem.v'),sK13) != s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f506_D])).
fof(f506_D,plain,(
  ( ! [X6] : (s('type.list.list'('type.modSem.v'),sK13) != s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f298,plain,(
  ( ! [X6] : (s('type.list.list'('type.modSem.v'),sK13) != s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X6),s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f226])).
fof(f11586,plain,(
  $false | ($spl0 | $spl272 | $spl3106)),
  inference(subsumption_resolution,[],[f11585,f506])).
fof(f11585,plain,(
  s('type.list.list'('type.modSem.v'),sK13) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v','const.list.HD_1'(s('type.list.list'('type.modSem.v'),sK13))),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))) | ($spl272 | $spl3106)),
  inference(trivial_inequality_removal,[],[f11473])).
fof(f11473,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.list.list'('type.modSem.v'),sK13) = s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v','const.list.HD_1'(s('type.list.list'('type.modSem.v'),sK13))),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))) | ($spl272 | $spl3106)),
  inference(superposition,[],[f673,f11472])).
% SZS output end Proof for thm.modProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 498
% Passive clauses: 1576
% Generated clauses: 11313
% Final active clauses: 473
% Final passive clauses: 1016
% Input formulas: 78
% Initial clauses: 201
% 
% Duplicate literals: 114
% Trivial inequalities: 57
% Fw subsumption resolutions: 787
% Fw demodulations: 969
% Bw demodulations: 73
% Condensations: 122
% Global subsumptions: 13
% 
% Simple tautologies: 201
% Equational tautologies: 8
% Forward subsumptions: 5752
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 2121
% Factoring: 306
% Forward superposition: 2600
% Backward superposition: 2291
% Self superposition: 27
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 1930
% SAT solver unit clauses: 252
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 11
% 
% Sat splits: 8
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 1876
% TWLsolver calls for satisfiability: 1727
% 
% Memory used [KB]: 9466
% Time elapsed: 0.218 s
% ------------------------------
% ------------------------------
% Success in time 0.86 s
