Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (1970)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 3585
% Generated clauses: 21678
% Final active clauses: 110
% Final passive clauses: 3277
% Input formulas: 44
% Initial clauses: 81
% 
% Duplicate literals: 513
% Trivial inequalities: 158
% Fw subsumption resolutions: 2199
% Fw demodulations: 3181
% Bw demodulations: 5
% 
% Simple tautologies: 211
% Equational tautologies: 74
% Forward subsumptions: 10761
% 
% Binary resolution: 18
% Forward superposition: 7816
% Backward superposition: 7107
% Self superposition: 318
% Equality factoring: 202
% Equality resolution: 27
% 
% Split clauses: 189
% Split components: 385
% SAT solver clauses: 412
% SAT solver unit clauses: 1
% SAT solver binary clauses: 42
% 
% Sat splits: 189
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 412
% TWLsolver calls for satisfiability: 90
% 
% Memory used [KB]: 20340
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (2004)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 1015
% Generated clauses: 10256
% Final active clauses: 106
% Final passive clauses: 844
% Input formulas: 44
% Initial clauses: 81
% 
% Duplicate literals: 278
% Trivial inequalities: 20
% Fw subsumption resolutions: 246
% Fw demodulations: 797
% Bw demodulations: 64
% Condensations: 525
% Global subsumptions: 7
% 
% Simple tautologies: 63
% Equational tautologies: 18
% Forward subsumptions: 5723
% 
% Binary resolution: 11
% Forward superposition: 4583
% Backward superposition: 2403
% Self superposition: 116
% Equality factoring: 34
% Equality resolution: 6
% 
% Split clauses: 61
% Split components: 134
% SAT solver clauses: 1619
% SAT solver unit clauses: 123
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 10
% 
% Sat splits: 61
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1591
% TWLsolver calls for satisfiability: 659
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f9840,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f162,f247,f245,f243,f241,f161,f246,f160,f244,f159,f242,f240,f277,f283,f281,f1007,f1008,f1024,f1009,f1010,f1025,f1011,f1026,f1023,f3472,f3473,f3474,f3475,f3476,f3494,f3495,f3496,f3497,f3498,f3502,f3504,f3506,f3507,f3508,f3509,f3510,f3511,f3551,f3600,f3553,f3601,f3555,f3556,f3557,f3558,f3559,f3603,f3599,f3560,f3602,f3598,f3604,f3618,f3605,f3606,f3619,f3607,f3622,f3620,f9180,f9181,f9182,f9183,f9184,f9185,f9186,f9187,f9188,f9189,f9190,f9191,f9192,f9193,f9194,f9195,f9196,f9197,f9198,f9199,f9200,f9201,f9202,f9203,f9204,f9205,f9206,f9207,f9208,f9209,f9210,f9211,f9212,f9213,f9214,f9215,f9216,f9217,f9218,f9219,f9220,f9221,f9222,f9223,f9224,f9225,f9226,f9227,f9228,f9229,f9233,f9238,f9234,f9274,f9232,f9237,f9235,f9278,f9288,f9291,f9329,f9338,f9802,f9804,f9805,f9806,f9807,f9808,f9839])).
fof(f9839,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f247,f9336])).
fof(f9336,plain,(
  s(sK1,sK4) = s(sK1,sK6) | $spl0),
  inference(unit_resulting_resolution,[],[f9327,f197])).
fof(f197,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | s(X1,X4) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & ((~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : ! [X8,X9,X16,X15,X18] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.infer.guard_3'(s(bool,X8),s(X6,X9),s(X7,X16))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X6),'const.infer.Success_1'(s('type.one.one',X15))),s(X7,X18))) <=> (p(s(bool,X8)) & s('type.one.one','const.one.one_0') = s('type.one.one',X15) & s(X7,X16) = s(X7,X18)))),
  file('thm.inferProps.success_eqns|split|13','thm.inferProps.inferProps_38_prove_ltb')).
fof(f9327,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK6))) | $spl0),
  inference(forward_demodulation,[],[f240,f165])).
fof(f165,plain,(
  ( ! [X0] : (s('type.one.one','const.one.one_0') = s('type.one.one',X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.one.one','const.one.one_0') = s('type.one.one',X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : s('type.one.one','const.one.one_0') = s('type.one.one',X15)),
  file('thm.inferProps.success_eqns|split|13','thm.one.one')).
fof(f9808,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9512,f9279])).
fof(f9279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,sK2),s(X0,X2),s(X1,X3))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,X5)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f243,f195])).
fof(f195,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f9512,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK6))) | $spl0),
  inference(superposition,[],[f9327,f165])).
fof(f9807,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9511,f9279])).
fof(f9511,plain,(
  ( ! [X2] : (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK1,sK6)))) ) | $spl0),
  inference(superposition,[],[f9327,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (s('type.one.one',X0) = s('type.one.one',X1)) )),
  inference(superposition,[],[f165,f165])).
fof(f9806,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9510,f9279])).
fof(f9510,plain,(
  ( ! [X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',X1))),s(sK1,sK6)))) ) | $spl0),
  inference(superposition,[],[f9327,f257])).
fof(f9805,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9509,f9279])).
fof(f9509,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',X0))),s(sK1,sK6)))) ) | $spl0),
  inference(superposition,[],[f9327,f165])).
fof(f9804,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9803,f9279])).
fof(f9803,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK4))) | $spl0),
  inference(backward_demodulation,[],[f9336,f9327])).
fof(f9802,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f9337,f243])).
fof(f9337,plain,(
  p(s(bool,sK2)) | $spl0),
  inference(unit_resulting_resolution,[],[f9327,f195])).
fof(f9338,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f243,f9327,f195])).
fof(f9329,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f9328,f1172])).
fof(f1172,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'F_0'),s(X0,X2),s(X1,X3))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f164,f195])).
fof(f164,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.inferProps.success_eqns|split|13',aHL_FALSITY)).
fof(f9328,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,'F_0'),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK6))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f9327,f281])).
fof(f9291,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f9276,f9279])).
fof(f9276,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK4))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9275,f165])).
fof(f9275,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK4))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f240,f246])).
fof(f9288,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(subsumption_resolution,[],[f9287,f1172])).
fof(f9287,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,'F_0'),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK4))) | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(backward_demodulation,[],[f9280,f9276])).
fof(f9280,plain,(
  s(bool,'F_0') = s(bool,sK2) | ($spl3 | $spl11)),
  inference(unit_resulting_resolution,[],[f283,f243,f378])).
fof(f378,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f163,f271])).
fof(f271,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f167,f167])).
fof(f167,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.inferProps.success_eqns|split|13',aHL_BOOL_CASES)).
fof(f163,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.inferProps.success_eqns|split|13',aHL_TRUTH)).
fof(f9278,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) | $spl44),
  inference(cnf_transformation,[],[f9278_D])).
fof(f9278_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f9235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(cnf_transformation,[],[f9235_D])).
fof(f9235_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9237,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f9236])).
fof(f9236,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f9173,f378])).
fof(f9173,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,X47)) | s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f250,f271])).
fof(f250,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(subsumption_resolution,[],[f198,f165])).
fof(f198,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | s(X1,X4) != s(X1,X6) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f136])).
fof(f9232,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f9231])).
fof(f9231,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f9172,f378])).
fof(f9172,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f250,f271])).
fof(f9274,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f245,f165])).
fof(f9234,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) | $spl40),
  inference(cnf_transformation,[],[f9234_D])).
fof(f9234_D,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f9238,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9237,f3633])).
fof(f3633,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,'T_0'),s(X29,X31),s(X30,X32))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34)))) ) | ($spl22 | $spl26)),
  inference(backward_demodulation,[],[f3599,f3601])).
fof(f9233,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9232,f3633])).
fof(f9229,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9117,f3633])).
fof(f9117,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f168,f250])).
fof(f168,plain,(
  ( ! [X0,X1] : (s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X38] : s(X6,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X6),X38))) = s(X6,happ(s(fun('type.infer.infer_st'(X6),X6),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X6),X38)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_1_const.infer.infer_st_subst')).
fof(f9228,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9118,f3633])).
fof(f9118,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f163,f206,f250])).
fof(f206,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X40,X41,X42] : ! [X44,X45,X47] : s(X40,happ(s(fun(X41,X40),X44),s(X41,happ(s(fun(X42,X41),X45),s(X42,X47))))) = s(X40,happ(s(fun(X42,X40),'_dst_x0x1_2'(s(fun(X41,X40),X44),s(fun(X42,X41),X45))),s(X42,X47)))),
  file('thm.inferProps.success_eqns|split|13','_dst_x0x1_ax')).
fof(f9227,plain,(
  $false | ($spl2 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9119,f3633])).
fof(f9119,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(bool,sK2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f163,f278,f250])).
fof(f278,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl2),
  inference(subsumption_resolution,[],[f275,f164])).
fof(f275,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK2) | $spl2),
  inference(superposition,[],[f242,f167])).
fof(f9226,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9120,f3633])).
fof(f9120,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f163,f589,f250])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,'type.one.one'),X1) = s(fun(X0,'type.one.one'),X2)) )),
  inference(unit_resulting_resolution,[],[f257,f182])).
fof(f182,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.inferProps.success_eqns|split|13',aHL_EXT)).
fof(f9225,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9121,f3633])).
fof(f9121,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f163,f3687,f250])).
fof(f3687,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun('type.one.one',bool)),X1) = s(fun(X0,fun('type.one.one',bool)),X2)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f1023,f182])).
fof(f9224,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9122,f3633])).
fof(f9122,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f163,f607,f250])).
fof(f607,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,'type.one.one')),X2) = s(fun(X0,fun(X1,'type.one.one')),X3)) )),
  inference(unit_resulting_resolution,[],[f589,f182])).
fof(f9223,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9123,f3633])).
fof(f9123,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f163,f732,f250])).
fof(f732,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X3) = s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X4)) )),
  inference(unit_resulting_resolution,[],[f607,f182])).
fof(f9222,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9124,f3633])).
fof(f9124,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f163,f760,f250])).
fof(f760,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X4) = s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X5)) )),
  inference(unit_resulting_resolution,[],[f732,f182])).
fof(f9221,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9125,f3633])).
fof(f9125,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f163,f795,f250])).
fof(f795,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X5) = s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X6)) )),
  inference(unit_resulting_resolution,[],[f760,f182])).
fof(f9220,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9126,f3633])).
fof(f9126,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f163,f1023,f250])).
fof(f9219,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9127,f3633])).
fof(f9127,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one','const.one.one_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.one.one',X3)))) )),
  inference(unit_resulting_resolution,[],[f163,f165,f250])).
fof(f9218,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9128,f3633])).
fof(f9128,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f163,f257,f250])).
fof(f9217,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9129,f3633])).
fof(f9129,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f183,f250])).
fof(f183,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X38,X39,X48] : s('type.pair.prod'('type.infer.exc'('type.one.one',X7),X6),'const.infer.guard_3'(s(bool,X38),s(X7,X39),s(X6,X48))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X7),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'('type.one.one',X7),X6)),'const.infer.guard_2'(s(bool,X38),s(X7,X39))),s(X6,X48)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_2_3_const.infer.guard')).
fof(f9216,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9130,f3633])).
fof(f9130,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f177,f250])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : ! [X5] : s('type.pair.prod'('type.infer.exc'(X6,X7),X6),'const.infer.read_1'(s(X6,X5))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X6),'const.pair.,_2'(s('type.infer.exc'(X6,X7),'const.infer.Success_1'(s(X6,X5))),s(X6,X5)))),
  file('thm.inferProps.success_eqns|split|13','thm.infer.read_def')).
fof(f9215,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9131,f3633])).
fof(f9131,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f178,f250])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : ! [X38] : s('type.pair.prod'('type.infer.exc'(X6,X7),X6),'const.infer.read_1'(s(X6,X38))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X6,X7),X6)),'const.infer.read_0'),s(X6,X38)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_1_const.infer.read')).
fof(f9214,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9132,f3633])).
fof(f9132,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.read_1'(s(sK1,sK4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.st_ex_return_2'(s(sK1,sK4),s(sK1,sK4)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f163,f5903,f250])).
fof(f5903,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.read_1'(s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK4),s(sK1,sK4)))) ) | $spl6),
  inference(forward_demodulation,[],[f5902,f246])).
fof(f5902,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.read_1'(s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK6),s(sK1,sK6)))) ) | $spl6),
  inference(forward_demodulation,[],[f5818,f177])).
fof(f5818,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK6),s(sK1,sK6))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.pair.,_2'(s('type.infer.exc'(sK1,X0),'const.infer.Success_1'(s(sK1,sK4))),s(sK1,sK4)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f246,f246,f212])).
fof(f212,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & ((s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6)) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X6,X7] : ! [X15,X16,X17,X18] : (s('type.pair.prod'('type.infer.exc'(X6,X14),X7),'const.infer.st_ex_return_2'(s(X6,X15),s(X7,X16))) = s('type.pair.prod'('type.infer.exc'(X6,X14),X7),'const.pair.,_2'(s('type.infer.exc'(X6,X14),'const.infer.Success_1'(s(X6,X17))),s(X7,X18))) <=> (s(X6,X15) = s(X6,X17) & s(X7,X16) = s(X7,X18)))),
  file('thm.inferProps.success_eqns|split|13','thm.inferProps.inferProps_24_prove_ltb')).
fof(f9213,plain,(
  $false | ($spl2 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9133,f3633])).
fof(f9133,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(bool,X1),bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.pair.prod'('type.infer.exc'(bool,X1),bool),'const.infer.read_1'(s(bool,'T_0'))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(bool,X1),bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(bool,X1),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f163,f6828,f250])).
fof(f6828,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f6827,f278])).
fof(f6827,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,sK2),s(bool,sK2)))) ) | $spl2),
  inference(forward_demodulation,[],[f4963,f177])).
fof(f4963,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,sK2),s(bool,sK2))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f278,f278,f212])).
fof(f9212,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9134,f3633])).
fof(f9134,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f179,f250])).
fof(f179,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7] : ! [X38] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X7)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X7),X38))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X7)),happ(s(fun('type.infer.infer_st'(X7),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X7))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X7),X38)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_1_const.infer.fresh_uvar')).
fof(f9211,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9135,f3633])).
fof(f9135,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f204,f250])).
fof(f204,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X14,X6,X7] : ! [X38,X39] : s('type.pair.prod'('type.infer.exc'(X14,X6),X7),'const.infer.failwith_2'(s(X6,X38),s(X7,X39))) = s('type.pair.prod'('type.infer.exc'(X14,X6),X7),happ(s(fun(X7,'type.pair.prod'('type.infer.exc'(X14,X6),X7)),'const.infer.failwith_1'(s(X6,X38))),s(X7,X39)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_1_2_const.infer.failwith')).
fof(f9210,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9136,f3633])).
fof(f9136,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f175,f250])).
fof(f175,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X38,X39] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X39))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X38))),s('type.num.num',X39)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_1_2_const.arithmetic.+')).
fof(f9209,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9137,f3633])).
fof(f9137,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f163,f166,f250])).
fof(f166,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X38] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X38))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X38)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_1_const.infer_t.Infer_Tuvar')).
fof(f9208,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9138,f3633])).
fof(f9138,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f176,f250])).
fof(f176,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X38,X39] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38),s('type.infer_t.infer_t',X39))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38))),s('type.infer_t.infer_t',X39)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_2_const.unify.t_walkstar')).
fof(f9207,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9139,f3633])).
fof(f9139,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f163,f180,f250])).
fof(f180,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7,X6] : ! [X38,X39] : s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X38),s('type.list.list'(X6),X39))) = s('type.list.list'(X7),happ(s(fun('type.list.list'(X6),'type.list.list'(X7)),happ(s(fun(fun(X6,X7),fun('type.list.list'(X6),'type.list.list'(X7))),'const.list.MAP_0'),s(fun(X6,X7),X38))),s('type.list.list'(X6),X39)))),
  file('thm.inferProps.success_eqns|split|13','ARITY_0_2_const.list.MAP')).
fof(f9206,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9140,f3633])).
fof(f9140,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK1,sK6)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f163,f246,f250])).
fof(f9205,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9141,f3633])).
fof(f9141,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f168,f250])).
fof(f9204,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9142,f3633])).
fof(f9142,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X8),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f163,f206,f250])).
fof(f9203,plain,(
  $false | ($spl2 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9143,f3633])).
fof(f9143,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(bool,sK2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(bool,'T_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f163,f278,f250])).
fof(f9202,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9144,f3633])).
fof(f9144,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f163,f589,f250])).
fof(f9201,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9145,f3633])).
fof(f9145,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f163,f3687,f250])).
fof(f9200,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9146,f3633])).
fof(f9146,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f163,f607,f250])).
fof(f9199,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9147,f3633])).
fof(f9147,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f163,f732,f250])).
fof(f9198,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9148,f3633])).
fof(f9148,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f163,f760,f250])).
fof(f9197,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9149,f3633])).
fof(f9149,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f163,f795,f250])).
fof(f9196,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9150,f3633])).
fof(f9150,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f163,f1023,f250])).
fof(f9195,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9151,f3633])).
fof(f9151,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one','const.one.one_0')))) )),
  inference(unit_resulting_resolution,[],[f163,f165,f250])).
fof(f9194,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9152,f3633])).
fof(f9152,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f163,f257,f250])).
fof(f9193,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9153,f3633])).
fof(f9153,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f183,f250])).
fof(f9192,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9154,f3633])).
fof(f9154,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f177,f250])).
fof(f9191,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9155,f3633])).
fof(f9155,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f178,f250])).
fof(f9190,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9156,f3633])).
fof(f9156,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.read_1'(s(sK1,sK4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.st_ex_return_2'(s(sK1,sK4),s(sK1,sK4)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f163,f5903,f250])).
fof(f9189,plain,(
  $false | ($spl2 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9157,f3633])).
fof(f9157,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(bool,X1),bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(bool,X1),bool),'const.infer.read_1'(s(bool,'T_0'))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(bool,X1),bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(bool,X1),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f163,f6828,f250])).
fof(f9188,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9158,f3633])).
fof(f9158,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f179,f250])).
fof(f9187,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9159,f3633])).
fof(f9159,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f204,f250])).
fof(f9186,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9160,f3633])).
fof(f9160,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f175,f250])).
fof(f9185,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9161,f3633])).
fof(f9161,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f163,f166,f250])).
fof(f9184,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9162,f3633])).
fof(f9162,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f176,f250])).
fof(f9183,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9163,f3633])).
fof(f9163,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f163,f180,f250])).
fof(f9182,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f9164,f3633])).
fof(f9164,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK1,sK6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK1,sK4)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f163,f246,f250])).
fof(f9181,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f9166])).
fof(f9166,plain,(
  s(sK1,sK4) != s(sK1,sK4) | ($spl1 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f163,f284,f250])).
fof(f284,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',X0))),s(sK1,sK4)))) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f279,f165])).
fof(f279,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK4))) | ($spl1 | $spl2 | $spl6)),
  inference(backward_demodulation,[],[f278,f256])).
fof(f256,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one','const.one.one_0'))),s(sK1,sK4))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f255,f165])).
fof(f255,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK4))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f241,f246])).
fof(f9180,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f9165])).
fof(f9165,plain,(
  s(sK1,sK4) != s(sK1,sK4) | ($spl1 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f163,f279,f250])).
fof(f3620,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl36),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3622,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl39),
  inference(cnf_transformation,[],[f3622_D])).
fof(f3622_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f3607,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl18),
  inference(superposition,[],[f1026,f165])).
fof(f3619,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl34),
  inference(cnf_transformation,[],[f3619_D])).
fof(f3619_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3606,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X3)))) ) | $spl18),
  inference(superposition,[],[f1026,f257])).
fof(f3605,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X2))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl18),
  inference(superposition,[],[f1026,f257])).
fof(f3618,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) | $spl32),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3604,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl18),
  inference(superposition,[],[f1026,f165])).
fof(f3598,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl20),
  inference(cnf_transformation,[],[f3598_D])).
fof(f3598_D,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3602,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) | $spl28),
  inference(cnf_transformation,[],[f3602_D])).
fof(f3602_D,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3560,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl12),
  inference(superposition,[],[f378,f3330])).
fof(f3330,plain,(
  ( ! [X255,X253,X254,X252] : (s(bool,happ(s(fun('type.one.one',bool),X254),s('type.one.one',X252))) = s(bool,happ(s(fun('type.one.one',bool),X255),s('type.one.one',X253)))) ) | $spl12),
  inference(superposition,[],[f2043,f257])).
fof(f2043,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,happ(s(fun('type.one.one',bool),X143),s('type.one.one',X144))) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun('type.one.one',X140),X142),s('type.one.one',X144)))))) ) | $spl12),
  inference(superposition,[],[f206,f1023])).
fof(f3599,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) | $spl22),
  inference(cnf_transformation,[],[f3599_D])).
fof(f3599_D,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3603,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) | $spl30),
  inference(cnf_transformation,[],[f3603_D])).
fof(f3603_D,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3559,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl12),
  inference(superposition,[],[f273,f3330])).
fof(f273,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f164,f167])).
fof(f3558,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3557,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3556,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3555,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3601,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl26),
  inference(cnf_transformation,[],[f3601_D])).
fof(f3601_D,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3553,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl12),
  inference(superposition,[],[f195,f3330])).
fof(f3600,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) | $spl24),
  inference(cnf_transformation,[],[f3600_D])).
fof(f3600_D,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3551,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl12),
  inference(superposition,[],[f167,f3330])).
fof(f3511,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl12),
  inference(superposition,[],[f378,f3330])).
fof(f3510,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl12),
  inference(superposition,[],[f273,f3330])).
fof(f3509,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3508,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3507,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3506,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl12),
  inference(superposition,[],[f271,f3330])).
fof(f3504,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl12),
  inference(superposition,[],[f195,f3330])).
fof(f3502,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl12),
  inference(superposition,[],[f167,f3330])).
fof(f3498,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl12),
  inference(superposition,[],[f3330,f167])).
fof(f3497,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3496,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3495,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3494,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3476,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl12),
  inference(superposition,[],[f3330,f167])).
fof(f3475,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3474,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3473,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f3472,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl12),
  inference(superposition,[],[f3330,f271])).
fof(f1023,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) | $spl12),
  inference(cnf_transformation,[],[f1023_D])).
fof(f1023_D,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1026,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) | $spl18),
  inference(cnf_transformation,[],[f1026_D])).
fof(f1026_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1011,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(fun('type.one.one',bool),X17) = s(fun('type.one.one',bool),X18) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) )),
  inference(superposition,[],[f595,f167])).
fof(f595,plain,(
  ( ! [X21,X19,X22,X20] : (s(X21,f(s('type.one.one',X22))) != s(X21,g(s('type.one.one',X22))) | s(fun('type.one.one',X21),X19) = s(fun('type.one.one',X21),X20)) )),
  inference(superposition,[],[f182,f257])).
fof(f1025,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl16),
  inference(cnf_transformation,[],[f1025_D])).
fof(f1025_D,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1010,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,f(s('type.one.one',X12))) | s(fun('type.one.one',bool),X14) = s(fun('type.one.one',bool),X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,g(s('type.one.one',X12)))) )),
  inference(superposition,[],[f595,f271])).
fof(f1009,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(fun('type.one.one',bool),X10) = s(fun('type.one.one',bool),X11) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f595,f271])).
fof(f1024,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) | $spl14),
  inference(cnf_transformation,[],[f1024_D])).
fof(f1024_D,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1008,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X4))) | s(fun('type.one.one',bool),X5) = s(fun('type.one.one',bool),X6) | s(bool,X7) = s(bool,g(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f595,f271])).
fof(f1007,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2) | s(bool,X3) = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f595,f271])).
fof(f281,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f281_D])).
fof(f281_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f283,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f277,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f167])).
fof(f240,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) | $spl0),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f242,plain,(
  p(s(bool,sK2)) | $spl2),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  p(s(bool,sK2)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f159,plain,(
  p(s(bool,sK2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) | (p(s(bool,sK2)) & s('type.one.one','const.one.one_0') = s('type.one.one',sK5) & s(sK1,sK4) = s(sK1,sK6))) & (s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) | ~p(s(bool,sK2)) | s('type.one.one','const.one.one_0') != s('type.one.one',sK5) | s(sK1,sK4) != s(sK1,sK6))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f124])).
fof(f124,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | ~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <~> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.infer.guard_3'(s(bool,X8),s(X6,X9),s(X7,X10))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X6),'const.infer.Success_1'(s('type.one.one',X11))),s(X7,X12))) <=> (p(s(bool,X8)) & s('type.one.one','const.one.one_0') = s('type.one.one',X11) & s(X7,X10) = s(X7,X12)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.infer.guard_3'(s(bool,X8),s(X6,X9),s(X7,X10))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X7),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X6),'const.infer.Success_1'(s('type.one.one',X11))),s(X7,X12))) <=> (p(s(bool,X8)) & s('type.one.one','const.one.one_0') = s('type.one.one',X11) & s(X7,X10) = s(X7,X12)))),
  file('thm.inferProps.success_eqns|split|13','thm.inferProps.success_eqns|split|13')).
fof(f244,plain,(
  s('type.one.one','const.one.one_0') = s('type.one.one',sK5) | $spl4),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s('type.one.one','const.one.one_0') = s('type.one.one',sK5) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f160,plain,(
  s('type.one.one','const.one.one_0') = s('type.one.one',sK5) | s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f246,plain,(
  s(sK1,sK4) = s(sK1,sK6) | $spl6),
  inference(cnf_transformation,[],[f246_D])).
fof(f246_D,plain,(
  s(sK1,sK4) = s(sK1,sK6) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f161,plain,(
  s(sK1,sK4) = s(sK1,sK6) | s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f241,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) | $spl1),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f243,plain,(
  ~p(s(bool,sK2)) | $spl3),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ~p(s(bool,sK2)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f245,plain,(
  s('type.one.one','const.one.one_0') != s('type.one.one',sK5) | $spl5),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  s('type.one.one','const.one.one_0') != s('type.one.one',sK5) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f247,plain,(
  s(sK1,sK4) != s(sK1,sK6) | $spl7),
  inference(cnf_transformation,[],[f247_D])).
fof(f247_D,plain,(
  s(sK1,sK4) != s(sK1,sK6) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f162,plain,(
  s(sK1,sK4) != s(sK1,sK6) | s('type.one.one','const.one.one_0') != s('type.one.one',sK5) | ~p(s(bool,sK2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.infer.guard_3'(s(bool,sK2),s(sK0,sK3),s(sK1,sK4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',sK0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',sK0),'const.infer.Success_1'(s('type.one.one',sK5))),s(sK1,sK6)))),
  inference(cnf_transformation,[],[f125])).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 68
% Passive clauses: 1017
% Generated clauses: 9757
% Final active clauses: 40
% Final passive clauses: 675
% Input formulas: 44
% Initial clauses: 81
% 
% Duplicate literals: 53
% Trivial inequalities: 6
% Fw subsumption resolutions: 93
% Fw demodulations: 3146
% Bw demodulations: 40
% Condensations: 155
% 
% Simple tautologies: 16
% Equational tautologies: 25
% Forward subsumptions: 4837
% Fw demodulations to eq. taut.: 303
% Bw demodulations to eq. taut.: 11
% 
% Unit resulting resolution: 2711
% Forward superposition: 1988
% Backward superposition: 1210
% Self superposition: 123
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 44
% Split components: 90
% SAT solver clauses: 1169
% SAT solver unit clauses: 380
% SAT solver binary clauses: 460
% 
% Sat splits: 44
% Sat splitting refutations: 62
% 
% Memory used [KB]: 13048
% Time elapsed: 0.445 s
% ------------------------------
% ------------------------------
% Success in time 1.788 s
