Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (122461)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 102
% Passive clauses: 2688
% Generated clauses: 18199
% Final active clauses: 99
% Final passive clauses: 2486
% Input formulas: 37
% Initial clauses: 56
% 
% Duplicate literals: 355
% Trivial inequalities: 187
% Fw subsumption resolutions: 2234
% Fw demodulations: 2805
% Bw demodulations: 2
% 
% Simple tautologies: 184
% Equational tautologies: 51
% Forward subsumptions: 9335
% 
% Binary resolution: 186
% Factoring: 1
% Forward superposition: 6264
% Backward superposition: 5606
% Self superposition: 286
% Equality factoring: 137
% Equality resolution: 28
% 
% Split clauses: 200
% Split components: 402
% SAT solver clauses: 379
% SAT solver unit clauses: 15
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 200
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 378
% TWLsolver calls for satisfiability: 90
% 
% Memory used [KB]: 13432
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (122505)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 84
% Passive clauses: 797
% Generated clauses: 9418
% Final active clauses: 78
% Final passive clauses: 672
% Input formulas: 37
% Initial clauses: 56
% 
% Duplicate literals: 341
% Trivial inequalities: 106
% Fw subsumption resolutions: 413
% Fw demodulations: 724
% Bw demodulations: 41
% Condensations: 886
% Global subsumptions: 3
% 
% Simple tautologies: 120
% Equational tautologies: 37
% Forward subsumptions: 4670
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 147
% Factoring: 1
% Forward superposition: 3055
% Backward superposition: 2642
% Self superposition: 149
% Equality factoring: 72
% Equality resolution: 18
% 
% Split clauses: 15
% Split components: 31
% SAT solver clauses: 1390
% SAT solver unit clauses: 78
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 7
% SAT solver learnt literals: 1
% 
% Sat splits: 15
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1377
% TWLsolver calls for satisfiability: 242
% 
% Memory used [KB]: 14967
% Time elapsed: 0.708 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f16408,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f126,f186,f184,f182,f125,f185,f124,f183,f181,f191,f995,f998,f997,f999,f1003,f1007,f1009,f1013,f1015,f1017,f1019,f1076,f1077,f1730,f1714,f1702,f1697,f1688,f1739,f1687,f1667,f1661,f1833,f1660,f1834,f1651,f1764,f1914,f1916,f1917,f6529,f6528,f6506,f6613,f6503,f6493,f6620,f6405,f6402,f6868,f6865,f6870,f6867,f6871,f6400,f1738,f6876,f6873,f6869,f6887,f6888,f6872,f7071,f7116,f7296,f7298,f7506,f7508,f7513,f7516,f7520,f7534,f7544,f7545,f7547,f7551,f7556,f7560,f7568,f7579,f15645,f15691,f16354,f16401])).
fof(f16401,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f16400,f15701])).
fof(f15701,plain,(
  ( ! [X0,X1] : (s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),X0))) != s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f184,f162])).
fof(f162,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4))) & (s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X3) | s('type.wordSem.word_loc'(X0),X2) != s('type.wordSem.word_loc'(X0),X4) | s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4))) & ((s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X3) | s('type.wordSem.word_loc'(X0),X2) != s('type.wordSem.word_loc'(X0),X4)) | s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) <=> (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) <=> (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))) = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X9),s('type.wordSem.word_loc'(X6),X10))) <=> (s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X9) & s('type.wordSem.word_loc'(X6),X8) = s('type.wordSem.word_loc'(X6),X10)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_|49||49||split|0')).
fof(f16400,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | $spl0),
  inference(forward_demodulation,[],[f16199,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X11,X12] : s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X11),s('type.wordSem.word_loc'(X6),X12))) = s('type.wordSem.result'(X6),happ(s(fun('type.wordSem.word_loc'(X6),'type.wordSem.result'(X6)),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),'type.wordSem.result'(X6))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(X6),X11))),s('type.wordSem.word_loc'(X6),X12)))),
  file('thm.wordSem.result_|49||49||split|1','ARITY_0_2_const.wordSem.Result')).
fof(f16199,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) = s('type.wordSem.result'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0)),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) | $spl0),
  inference(superposition,[],[f154,f15740])).
fof(f15740,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) = s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK3))),s('type.wordSem.word_loc'(sK0),sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f15702,f177])).
fof(f177,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X6] : ! [X7,X8,X13,X15,X16,X17,X18] : s(X14,happ(s(fun('type.wordSem.word_loc'(X6),X14),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s('type.wordSem.word_loc'(X6),X7))),s('type.wordSem.word_loc'(X6),X8))) = s(X14,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s(X14,X16),s(X14,X17),s(X14,X18)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_case_def|split|1')).
fof(f15702,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK3))),s('type.wordSem.word_loc'(sK0),sK4))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(X0,X3),s(X0,X4),s(X0,X5)))) ) | $spl0),
  inference(superposition,[],[f177,f181])).
fof(f16354,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f16353,f15701])).
fof(f16353,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | $spl0),
  inference(forward_demodulation,[],[f16076,f154])).
fof(f16076,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) = s('type.wordSem.result'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0)),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) | $spl0),
  inference(superposition,[],[f15740,f154])).
fof(f15691,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f15690,f2426])).
fof(f2426,plain,(
  ( ! [X0,X1] : (s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),X0),s('type.wordSem.word_loc'(sK0),sK4))) != s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),X1),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f186,f163])).
fof(f163,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)) )),
  inference(cnf_transformation,[],[f121])).
fof(f15690,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f15486,f154])).
fof(f15486,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK4))) = s('type.wordSem.result'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0)),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f154,f12856])).
fof(f12856,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) = s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK4)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f12388,f177])).
fof(f12388,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK4))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(X0,X3),s(X0,X4),s(X0,X5)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f177,f192])).
fof(f192,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f181,f183])).
fof(f15645,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f15644,f2426])).
fof(f15644,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f15364,f154])).
fof(f15364,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK4))) = s('type.wordSem.result'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0)),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),'type.wordSem.result'(sK0))),'const.wordSem.Result_0'),s('type.wordSem.word_loc'(sK0),sK1))),s('type.wordSem.word_loc'(sK0),sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f12856,f154])).
fof(f7579,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267,f7546,f7555,f7559,f7567,f7578])).
fof(f7578,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,X82) = s(bool,'const.bool.DATATYPE_1'(s(X83,X84))) | s(bool,X85) = s(bool,'const.bool.DATATYPE_1'(s(bool,X82)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7240,f6893])).
fof(f7240,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,X82) = s(bool,'const.bool.DATATYPE_1'(s(X83,X84))) | s(bool,X85) = s(bool,'const.bool.DATATYPE_1'(s(bool,X82))) | s(bool,'T_0') = s(bool,X85)) ) | $spl28),
  inference(superposition,[],[f3843,f6869])).
fof(f7567,plain,(
  ( ! [X52,X51] : (s(bool,X52) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7566,f7115])).
fof(f7566,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X52) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51)))) ) | $spl28),
  inference(forward_demodulation,[],[f7565,f7115])).
fof(f7565,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,X52) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7225,f6893])).
fof(f7225,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,X52) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51))) | s(bool,'T_0') = s(bool,X52)) ) | $spl28),
  inference(superposition,[],[f1774,f6869])).
fof(f7559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f7558,f131])).
fof(f7558,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f7557])).
fof(f7557,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl28),
  inference(forward_demodulation,[],[f7209,f148])).
fof(f7209,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) = s(bool,'const.bool.DATATYPE_1'(s(bool,X19)))) ) | $spl28),
  inference(superposition,[],[f165,f6869])).
fof(f7555,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f7554,f131])).
fof(f7554,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f7553])).
fof(f7553,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X18))) ) | $spl28),
  inference(forward_demodulation,[],[f7208,f148])).
fof(f7208,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) = s(bool,'const.bool.DATATYPE_1'(s(bool,X17)))) ) | $spl28),
  inference(superposition,[],[f165,f6869])).
fof(f7546,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl28),
  inference(subsumption_resolution,[],[f7204,f7115])).
fof(f7204,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'T_0') != s(bool,'const.bool.DATATYPE_1'(s(bool,X13)))) ) | $spl28),
  inference(superposition,[],[f131,f6869])).
fof(f7267,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f7199])).
fof(f7199,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl28),
  inference(superposition,[],[f3606,f6869])).
fof(f7543,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X37) = s(bool,'const.bool.DATATYPE_1'(s(X38,X39))) | s(bool,X40) = s(bool,'const.bool.DATATYPE_1'(s(bool,X37)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7194,f6893])).
fof(f7194,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X37) = s(bool,'const.bool.DATATYPE_1'(s(X38,X39))) | s(bool,X40) = s(bool,'const.bool.DATATYPE_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X40)) ) | $spl28),
  inference(superposition,[],[f6869,f3843])).
fof(f7271,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f7178])).
fof(f7178,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl28),
  inference(superposition,[],[f6869,f3606])).
fof(f7503,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (s(fun(X442,bool),X443) = s(fun(X442,bool),X444) | s(bool,X445) = s(bool,'const.bool.DATATYPE_1'(s(X446,X447))) | s(bool,X445) = s(bool,g(s(X442,sK7(X443,X444,bool,X442))))) ) | $spl28),
  inference(trivial_inequality_removal,[],[f7456])).
fof(f7456,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X442,bool),X443) = s(fun(X442,bool),X444) | s(bool,X445) = s(bool,'const.bool.DATATYPE_1'(s(X446,X447))) | s(bool,X445) = s(bool,g(s(X442,sK7(X443,X444,bool,X442))))) ) | $spl28),
  inference(backward_demodulation,[],[f7115,f6306])).
fof(f7507,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f7301,f7115])).
fof(f7301,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl28),
  inference(backward_demodulation,[],[f7115,f135])).
fof(f7115,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(unit_resulting_resolution,[],[f6869,f381])).
fof(f1328,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f173,f211])).
fof(f1330,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK7(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK7(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f173,f211])).
fof(f1332,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK7(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK7(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f173,f130])).
fof(f3889,plain,(
  ( ! [X313,X315,X314] : (s(bool,'T_0') != s(bool,f(s(X313,sK7(X314,X315,bool,X313)))) | s(fun(X313,bool),X314) = s(fun(X313,bool),X315) | s(bool,'F_0') = s(bool,g(s(X313,sK7(X314,X315,bool,X313))))) )),
  inference(forward_demodulation,[],[f3748,f148])).
fof(f3748,plain,(
  ( ! [X313,X315,X316,X314,X317] : (s(bool,'const.bool.DATATYPE_1'(s(X316,X317))) != s(bool,f(s(X313,sK7(X314,X315,bool,X313)))) | s(fun(X313,bool),X314) = s(fun(X313,bool),X315) | s(bool,'F_0') = s(bool,g(s(X313,sK7(X314,X315,bool,X313))))) )),
  inference(superposition,[],[f173,f3606])).
fof(f4429,plain,(
  ( ! [X269,X267,X268,X270] : (s(bool,X270) != s(bool,f(s(X267,sK7(X268,X269,bool,X267)))) | s(fun(X267,bool),X268) = s(fun(X267,bool),X269) | s(bool,'F_0') = s(bool,X270) | s(bool,'F_0') = s(bool,g(s(X267,sK7(X268,X269,bool,X267))))) )),
  inference(superposition,[],[f173,f3673])).
fof(f4653,plain,(
  ( ! [X255,X253,X254,X252] : (s(bool,'F_0') != s(bool,f(s(X252,sK7(X253,X254,bool,X252)))) | s(fun(X252,bool),X253) = s(fun(X252,bool),X254) | s(bool,X255) = s(bool,g(s(X252,sK7(X253,X254,bool,X252)))) | s(bool,'F_0') = s(bool,X255)) )),
  inference(superposition,[],[f173,f3673])).
fof(f6163,plain,(
  ( ! [X472,X474,X469,X471,X473,X470] : (s(bool,X472) != s(bool,f(s(X469,sK7(X470,X471,bool,X469)))) | s(fun(X469,bool),X470) = s(fun(X469,bool),X471) | s(bool,X472) = s(bool,'const.bool.DATATYPE_1'(s(X473,X474))) | s(bool,'T_0') = s(bool,g(s(X469,sK7(X470,X471,bool,X469))))) )),
  inference(superposition,[],[f173,f3843])).
fof(f6306,plain,(
  ( ! [X445,X447,X443,X444,X446,X442] : (s(bool,'T_0') != s(bool,f(s(X442,sK7(X443,X444,bool,X442)))) | s(fun(X442,bool),X443) = s(fun(X442,bool),X444) | s(bool,X445) = s(bool,'const.bool.DATATYPE_1'(s(X446,X447))) | s(bool,X445) = s(bool,g(s(X442,sK7(X443,X444,bool,X442))))) )),
  inference(superposition,[],[f173,f3843])).
fof(f1190,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,X7) = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f157,f211])).
fof(f1192,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) = s(bool,happ(s(fun('type.num.num',bool),X13),s('type.num.num',X14))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X12),s(fun('type.num.num',bool),X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f157,f211])).
fof(f1194,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X20),s(fun('type.num.num',bool),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f157,f130])).
fof(f1234,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X90),s('type.num.num',X91))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X89),s(fun('type.num.num',bool),X90),s('type.num.num','const.num.SUC_1'(s('type.num.num',X91)))))) )),
  inference(forward_demodulation,[],[f1216,f148])).
fof(f1216,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,'const.bool.DATATYPE_1'(s(X92,X93))) != s(bool,happ(s(fun('type.num.num',bool),X90),s('type.num.num',X91))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X89),s(fun('type.num.num',bool),X90),s('type.num.num','const.num.SUC_1'(s('type.num.num',X91)))))) )),
  inference(superposition,[],[f381,f157])).
fof(f2099,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X36),s('type.num.num',X37))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X35),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))) )),
  inference(forward_demodulation,[],[f2043,f148])).
fof(f2043,plain,(
  ( ! [X37,X35,X36] : (s(bool,happ(s(fun('type.num.num',bool),X36),s('type.num.num',X37))) = s(bool,'const.bool.DATATYPE_1'(s(bool,happ(s(fun('type.num.num',bool),X36),s('type.num.num',X37))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X35),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f1650,f157])).
fof(f2501,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,X51) = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f1774,f157])).
fof(f2757,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X64) != s(bool,happ(s(fun('type.num.num',bool),X62),s('type.num.num',X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X61),s(fun('type.num.num',bool),X62),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63)))))) )),
  inference(superposition,[],[f2056,f157])).
fof(f2781,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X61) != s(bool,happ(s(fun('type.num.num',bool),X59),s('type.num.num',X60))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X58),s(fun('type.num.num',bool),X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f2056,f157])).
fof(f3218,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun('type.num.num',bool),X72),s('type.num.num',X73))),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.SUC_1'(s('type.num.num',X73)))))) )),
  inference(superposition,[],[f267,f157])).
fof(f4226,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,X117) = s(bool,happ(s(fun('type.num.num',bool),X115),s('type.num.num',X116))) | s(bool,'F_0') = s(bool,X117) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X114),s(fun('type.num.num',bool),X115),s('type.num.num','const.num.SUC_1'(s('type.num.num',X116)))))) )),
  inference(superposition,[],[f3673,f157])).
fof(f4306,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X101),s('type.num.num',X102))) | s(bool,X103) = s(bool,'const.ind_type.FCONS_3'(s(bool,X100),s(fun('type.num.num',bool),X101),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102))))) | s(bool,'F_0') = s(bool,X103)) )),
  inference(superposition,[],[f3673,f157])).
fof(f5172,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X117),s('type.num.num',X118))) | s(bool,X119) = s(bool,'const.ind_type.FCONS_3'(s(bool,X116),s(fun('type.num.num',bool),X117),s('type.num.num','const.num.SUC_1'(s('type.num.num',X118))))) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f4534,f157])).
fof(f5774,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (s(bool,X231) = s(bool,happ(s(fun('type.num.num',bool),X229),s('type.num.num',X230))) | s(bool,X231) = s(bool,'const.bool.DATATYPE_1'(s(X232,X233))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X228),s(fun('type.num.num',bool),X229),s('type.num.num','const.num.SUC_1'(s('type.num.num',X230)))))) )),
  inference(superposition,[],[f3843,f157])).
fof(f5823,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X186),s('type.num.num',X187))) | s(bool,X188) = s(bool,'const.bool.DATATYPE_1'(s(X189,X190))) | s(bool,X188) = s(bool,'const.ind_type.FCONS_3'(s(bool,X185),s(fun('type.num.num',bool),X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X187)))))) )),
  inference(superposition,[],[f3843,f157])).
fof(f3887,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,'|dollar|exists_1'(s(fun(X141,bool),X142)))) | s(bool,'F_0') = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143)))) )),
  inference(subsumption_resolution,[],[f3886,f127])).
fof(f3886,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X141,bool),X142)))) | s(bool,'F_0') = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143)))) )),
  inference(forward_demodulation,[],[f3717,f148])).
fof(f3717,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X144,X145)))) | p(s(bool,'|dollar|exists_1'(s(fun(X141,bool),X142)))) | s(bool,'F_0') = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143)))) )),
  inference(superposition,[],[f897,f3606])).
fof(f903,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X17)))) | p(s(bool,'|dollar|exists_1'(s(fun(X14,bool),sK5(X15,X16,fun(X14,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X14,bool),bool),X16)))) )),
  inference(superposition,[],[f897,f419])).
fof(f932,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f920,f127])).
fof(f920,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f897,f211])).
fof(f897,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f841,f131])).
fof(f841,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f156,f419])).
fof(f4019,plain,(
  ( ! [X189,X187,X185,X188,X186] : (~p(s(bool,happ(s(fun(X187,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X187,X185),X188))),s(X187,X189)))) | p(s(bool,'|dollar|exists_1'(s(fun(X185,bool),X186))))) )),
  inference(superposition,[],[f897,f179])).
fof(f4021,plain,(
  ( ! [X198,X196,X199,X197,X195] : (p(s(bool,happ(s(fun(X197,bool),'_dst_x0x1_2'(s(fun(X195,bool),X196),s(fun(X197,X195),X198))),s(X197,X199)))) | s(bool,'T_0') != s(bool,happ(s(fun(X195,bool),X196),s(X195,happ(s(fun(X197,X195),X198),s(X197,X199)))))) )),
  inference(superposition,[],[f131,f179])).
fof(f4022,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,happ(s(fun(X202,bool),'_dst_x0x1_2'(s(fun(X200,bool),X201),s(fun(X202,X200),X203))),s(X202,X204)))) | s(bool,'T_0') = s(bool,happ(s(fun(X200,bool),X201),s(X200,happ(s(fun(X202,X200),X203),s(X202,X204)))))) )),
  inference(superposition,[],[f132,f179])).
fof(f4023,plain,(
  ( ! [X208,X206,X209,X207,X205] : (s(bool,'F_0') != s(bool,happ(s(fun(X207,bool),'_dst_x0x1_2'(s(fun(X205,bool),X206),s(fun(X207,X205),X208))),s(X207,X209))) | ~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,happ(s(fun(X207,X205),X208),s(X207,X209))))))) )),
  inference(superposition,[],[f135,f179])).
fof(f4024,plain,(
  ( ! [X212,X210,X213,X211,X214] : (p(s(bool,happ(s(fun(X212,bool),'_dst_x0x1_2'(s(fun(X210,bool),X211),s(fun(X212,X210),X213))),s(X212,X214)))) | s(bool,'F_0') = s(bool,happ(s(fun(X210,bool),X211),s(X210,happ(s(fun(X212,X210),X213),s(X212,X214)))))) )),
  inference(superposition,[],[f136,f179])).
fof(f4025,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (p(s(bool,happ(s(fun(X217,bool),'_dst_x0x1_2'(s(fun(X215,bool),X216),s(fun(X217,X215),X218))),s(X217,X219)))) | p(s(bool,X220)) | s(bool,X220) = s(bool,happ(s(fun(X215,bool),X216),s(X215,happ(s(fun(X217,X215),X218),s(X217,X219)))))) )),
  inference(superposition,[],[f165,f179])).
fof(f4027,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (~p(s(bool,happ(s(fun(X229,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X229,X227),X230))),s(X229,X231)))) | ~p(s(bool,X232)) | s(bool,X232) = s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X229,X227),X230),s(X229,X231)))))) )),
  inference(superposition,[],[f168,f179])).
fof(f4032,plain,(
  ( ! [X261,X257,X259,X260,X262,X258] : (~p(s(bool,happ(s(fun(X259,bool),'_dst_x0x1_2'(s(fun(X257,bool),X258),s(fun(X259,X257),X260))),s(X259,X261)))) | s(bool,X262) != s(bool,happ(s(fun(X257,bool),X258),s(X257,happ(s(fun(X259,X257),X260),s(X259,X261))))) | s(bool,'T_0') = s(bool,X262)) )),
  inference(superposition,[],[f214,f179])).
fof(f4033,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (p(s(bool,happ(s(fun(X265,bool),'_dst_x0x1_2'(s(fun(X263,bool),X264),s(fun(X265,X263),X266))),s(X265,X267)))) | s(bool,X268) = s(bool,happ(s(fun(X263,bool),X264),s(X263,happ(s(fun(X265,X263),X266),s(X265,X267))))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f247,f179])).
fof(f4034,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun(X271,bool),'_dst_x0x1_2'(s(fun(X269,bool),X270),s(fun(X271,X269),X272))),s(X271,X273))),s(fun('type.num.num',bool),X274),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X269,bool),X270),s(X269,happ(s(fun(X271,X269),X272),s(X271,X273)))))) )),
  inference(superposition,[],[f267,f179])).
fof(f4050,plain,(
  ( ! [X366,X368,X370,X365,X367,X369] : (~p(s(bool,happ(s(fun(X367,bool),'_dst_x0x1_2'(s(fun(X365,bool),X366),s(fun(X367,X365),X368))),s(X367,X369)))) | s(bool,X370) = s(bool,happ(s(fun(X365,bool),X366),s(X365,happ(s(fun(X367,X365),X368),s(X367,X369))))) | s(bool,'F_0') = s(bool,X370)) )),
  inference(superposition,[],[f3302,f179])).
fof(f4100,plain,(
  ( ! [X191,X194,X192,X190,X193] : (p(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f131,f179])).
fof(f4101,plain,(
  ( ! [X198,X196,X199,X197,X195] : (~p(s(bool,happ(s(fun(X196,bool),X197),s(X196,happ(s(fun(X195,X196),X198),s(X195,X199)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X196,bool),X197),s(fun(X195,X196),X198))),s(X195,X199)))) )),
  inference(superposition,[],[f132,f179])).
fof(f4102,plain,(
  ( ! [X204,X202,X200,X203,X201] : (s(bool,'F_0') != s(bool,happ(s(fun(X201,bool),X202),s(X201,happ(s(fun(X200,X201),X203),s(X200,X204))))) | ~p(s(bool,happ(s(fun(X200,bool),'_dst_x0x1_2'(s(fun(X201,bool),X202),s(fun(X200,X201),X203))),s(X200,X204))))) )),
  inference(superposition,[],[f135,f179])).
fof(f4103,plain,(
  ( ! [X208,X206,X209,X207,X205] : (p(s(bool,happ(s(fun(X206,bool),X207),s(X206,happ(s(fun(X205,X206),X208),s(X205,X209)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X205,bool),'_dst_x0x1_2'(s(fun(X206,bool),X207),s(fun(X205,X206),X208))),s(X205,X209)))) )),
  inference(superposition,[],[f136,f179])).
fof(f4104,plain,(
  ( ! [X212,X210,X215,X213,X211,X214] : (p(s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X210,X211),X213),s(X210,X214)))))) | p(s(bool,X215)) | s(bool,X215) = s(bool,happ(s(fun(X210,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X210,X211),X213))),s(X210,X214)))) )),
  inference(superposition,[],[f165,f179])).
fof(f4106,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (~p(s(bool,happ(s(fun(X223,bool),X224),s(X223,happ(s(fun(X222,X223),X225),s(X222,X226)))))) | ~p(s(bool,X227)) | s(bool,X227) = s(bool,happ(s(fun(X222,bool),'_dst_x0x1_2'(s(fun(X223,bool),X224),s(fun(X222,X223),X225))),s(X222,X226)))) )),
  inference(superposition,[],[f168,f179])).
fof(f4111,plain,(
  ( ! [X255,X253,X257,X254,X252,X256] : (~p(s(bool,happ(s(fun(X253,bool),X254),s(X253,happ(s(fun(X252,X253),X255),s(X252,X256)))))) | s(bool,X257) != s(bool,happ(s(fun(X252,bool),'_dst_x0x1_2'(s(fun(X253,bool),X254),s(fun(X252,X253),X255))),s(X252,X256))) | s(bool,'T_0') = s(bool,X257)) )),
  inference(superposition,[],[f214,f179])).
fof(f4112,plain,(
  ( ! [X261,X263,X259,X260,X262,X258] : (p(s(bool,happ(s(fun(X259,bool),X260),s(X259,happ(s(fun(X258,X259),X261),s(X258,X262)))))) | s(bool,X263) = s(bool,happ(s(fun(X258,bool),'_dst_x0x1_2'(s(fun(X259,bool),X260),s(fun(X258,X259),X261))),s(X258,X262))) | s(bool,'T_0') = s(bool,X263)) )),
  inference(superposition,[],[f247,f179])).
fof(f4129,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (~p(s(bool,happ(s(fun(X361,bool),X362),s(X361,happ(s(fun(X360,X361),X363),s(X360,X364)))))) | s(bool,X365) = s(bool,happ(s(fun(X360,bool),'_dst_x0x1_2'(s(fun(X361,bool),X362),s(fun(X360,X361),X363))),s(X360,X364))) | s(bool,'F_0') = s(bool,X365)) )),
  inference(superposition,[],[f3302,f179])).
fof(f4055,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (~p(s(bool,happ(s(fun(X399,bool),happ(s(fun(X402,fun(X399,bool)),'_dst_x0x1_2'(s(fun(X400,fun(X399,bool)),X401),s(fun(X402,X400),X403))),s(X402,X404))),s(X399,X405)))) | p(s(bool,'|dollar|exists_1'(s(fun(X399,bool),happ(s(fun(X400,fun(X399,bool)),X401),s(X400,happ(s(fun(X402,X400),X403),s(X402,X404))))))))) )),
  inference(superposition,[],[f897,f179])).
fof(f4134,plain,(
  ( ! [X397,X399,X395,X396,X398,X394,X400] : (~p(s(bool,happ(s(fun(X394,bool),happ(s(fun(X396,fun(X394,bool)),X397),s(X396,happ(s(fun(X395,X396),X398),s(X395,X399))))),s(X394,X400)))) | p(s(bool,'|dollar|exists_1'(s(fun(X394,bool),happ(s(fun(X395,fun(X394,bool)),'_dst_x0x1_2'(s(fun(X396,fun(X394,bool)),X397),s(fun(X395,X396),X398))),s(X395,X399))))))) )),
  inference(superposition,[],[f897,f179])).
fof(f4098,plain,(
  ( ! [X182,X180,X184,X183,X181] : (~p(s(bool,happ(s(fun(X181,bool),X182),s(X181,happ(s(fun(X180,X181),X183),s(X180,X184)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X180,bool),'_dst_x0x1_2'(s(fun(X181,bool),X182),s(fun(X180,X181),X183))))))) )),
  inference(superposition,[],[f897,f179])).
fof(f4113,plain,(
  ( ! [X269,X265,X267,X268,X264,X266] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun(X265,bool),X266),s(X265,happ(s(fun(X264,X265),X267),s(X264,X268))))),s(fun('type.num.num',bool),X269),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X264,bool),'_dst_x0x1_2'(s(fun(X265,bool),X266),s(fun(X264,X265),X267))),s(X264,X268)))) )),
  inference(superposition,[],[f267,f179])).
fof(f938,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f169,f211])).
fof(f940,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f169,f211])).
fof(f942,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f169,f130])).
fof(f982,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),X77))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X76),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X76),X77)))) )),
  inference(forward_demodulation,[],[f967,f148])).
fof(f967,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'const.bool.DATATYPE_1'(s(X78,X79))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),X77))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X76),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X76),X77)))) )),
  inference(superposition,[],[f381,f169])).
fof(f2097,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X31),X32)))) )),
  inference(forward_demodulation,[],[f2041,f148])).
fof(f2041,plain,(
  ( ! [X31,X32] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X31),X32)))) )),
  inference(superposition,[],[f1650,f169])).
fof(f2499,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X42),X43))) | s(bool,X44) = s(bool,happ(s(fun('type.pair.prod'(bool,X42),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X42),X43))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f1774,f169])).
fof(f2754,plain,(
  ( ! [X50,X48,X49] : (s(bool,X50) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),X49))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X48),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X48),X49)))) )),
  inference(superposition,[],[f2056,f169])).
fof(f2778,plain,(
  ( ! [X47,X45,X46] : (s(bool,X47) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X45),X46))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X45),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X45),X46))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f2056,f169])).
fof(f3214,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X51),X52))),s(fun('type.num.num',bool),X53),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X51),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X51),X52)))) )),
  inference(superposition,[],[f267,f169])).
fof(f4219,plain,(
  ( ! [X76,X77,X75] : (s(bool,X77) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),X76))) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X75),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X75),X76)))) )),
  inference(superposition,[],[f3673,f169])).
fof(f4299,plain,(
  ( ! [X61,X62,X63] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),X62))) | s(bool,X63) = s(bool,happ(s(fun('type.pair.prod'(bool,X61),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X61),X62))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f3673,f169])).
fof(f5165,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X77),X78))) | s(bool,X79) = s(bool,happ(s(fun('type.pair.prod'(bool,X77),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X77),X78))) | s(bool,'F_0') = s(bool,X79)) )),
  inference(superposition,[],[f4534,f169])).
fof(f5767,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s(bool,X177) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X175),X176))) | s(bool,X177) = s(bool,'const.bool.DATATYPE_1'(s(X178,X179))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X175),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X175),X176)))) )),
  inference(superposition,[],[f3843,f169])).
fof(f5816,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),X133))) | s(bool,X134) = s(bool,'const.bool.DATATYPE_1'(s(X135,X136))) | s(bool,X134) = s(bool,happ(s(fun('type.pair.prod'(bool,X132),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X132),X133)))) )),
  inference(superposition,[],[f3843,f169])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14] : ! [X31,X32] : (! [X21] : (s(X14,X21) = s(X14,X32) => p(s(bool,happ(s(fun(X14,bool),X31),s(X14,X21))))) => p(s(bool,'|dollar|exists_1'(s(fun(X14,bool),X31)))))),
  file('thm.wordSem.result_|49||49||split|1','thm.bool.JRH_INDUCT_UTIL')).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X3) = s(X0,sK5(X3,'const.ind_type.FCONS_3'(s(fun(X0,bool),X1),s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0')),X0))) )),
  inference(superposition,[],[f155,f152])).
fof(f416,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X7)))) | s(X4,X8) = s(X4,sK5(X8,'const.pair.SND_1'(s('type.pair.prod'(X5,fun(X4,bool)),'const.pair.,_2'(s(X5,X6),s(fun(X4,bool),X7)))),X4))) )),
  inference(superposition,[],[f155,f171])).
fof(f417,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X11)))) | s(X9,X13) = s(X9,sK5(X13,'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9))) )),
  inference(superposition,[],[f155,f172])).
fof(f834,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X3) = s(X0,sK5(X3,'const.ind_type.FCONS_3'(s(fun(X0,bool),X1),s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0')),X0))) )),
  inference(superposition,[],[f419,f152])).
fof(f835,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X7))) | s(X4,X8) = s(X4,sK5(X8,'const.pair.SND_1'(s('type.pair.prod'(X5,fun(X4,bool)),'const.pair.,_2'(s(X5,X6),s(fun(X4,bool),X7)))),X4))) )),
  inference(superposition,[],[f419,f171])).
fof(f836,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X9,bool),X11))) | s(X9,X13) = s(X9,sK5(X13,'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9))) )),
  inference(superposition,[],[f419,f172])).
fof(f837,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(X14,X17) = s(X14,sK5(X17,sK5(X15,X16,fun(X14,bool)),X14)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X14,bool),bool),X16)))) )),
  inference(superposition,[],[f419,f419])).
fof(f866,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) | s(X83,X86) = s(X83,sK5(X86,sK5(X84,X85,fun(X83,bool)),X83)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X83,bool),bool),X85)))) )),
  inference(superposition,[],[f155,f419])).
fof(f898,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK5(X30,X28,X27))) )),
  inference(subsumption_resolution,[],[f875,f127])).
fof(f875,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK5(X30,X28,X27))) )),
  inference(superposition,[],[f168,f419])).
fof(f2669,plain,(
  ( ! [X277,X279,X273,X275,X276,X278,X272,X274] : (p(s(bool,'|dollar|exists_1'(s(fun(X272,bool),X276)))) | s(X272,X279) = s(X272,sK5(X279,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X273),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X273),fun('type.wordSem.word_loc'(X273),fun(X272,bool))),X274),s(fun('type.wordSem.word_loc'(X273),fun('type.wordSem.word_loc'(X273),fun(X272,bool))),X275),s(fun(X272,bool),X276),s(fun(X272,bool),X277),s(fun(X272,bool),X278)),X272))) )),
  inference(superposition,[],[f155,f174])).
fof(f2670,plain,(
  ( ! [X280,X282,X285,X287,X281,X283,X284,X286] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X280,bool),X284))) | s(X280,X287) = s(X280,sK5(X287,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X281),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X281),fun('type.wordSem.word_loc'(X281),fun(X280,bool))),X282),s(fun('type.wordSem.word_loc'(X281),fun('type.wordSem.word_loc'(X281),fun(X280,bool))),X283),s(fun(X280,bool),X284),s(fun(X280,bool),X285),s(fun(X280,bool),X286)),X280))) )),
  inference(superposition,[],[f419,f174])).
fof(f3051,plain,(
  ( ! [X341,X343,X337,X339,X340,X342,X336,X338] : (p(s(bool,'|dollar|exists_1'(s(fun(X336,bool),X342)))) | s(X336,X343) = s(X336,sK5(X343,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X337),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X337),fun('type.wordSem.word_loc'(X337),fun(X336,bool))),X338),s(fun('type.wordSem.word_loc'(X337),fun('type.wordSem.word_loc'(X337),fun(X336,bool))),X339),s(fun(X336,bool),X340),s(fun(X336,bool),X341),s(fun(X336,bool),X342)),X336))) )),
  inference(superposition,[],[f155,f175])).
fof(f3052,plain,(
  ( ! [X350,X344,X346,X349,X351,X345,X347,X348] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X344,bool),X350))) | s(X344,X351) = s(X344,sK5(X351,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X345),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X345),fun('type.wordSem.word_loc'(X345),fun(X344,bool))),X346),s(fun('type.wordSem.word_loc'(X345),fun('type.wordSem.word_loc'(X345),fun(X344,bool))),X347),s(fun(X344,bool),X348),s(fun(X344,bool),X349),s(fun(X344,bool),X350)),X344))) )),
  inference(superposition,[],[f419,f175])).
fof(f3495,plain,(
  ( ! [X397,X399,X393,X395,X396,X398,X394,X400] : (p(s(bool,'|dollar|exists_1'(s(fun(X393,bool),X398)))) | s(X393,X400) = s(X393,sK5(X400,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X394),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X394),fun('type.wordSem.word_loc'(X394),fun(X393,bool))),X395),s(fun('type.wordSem.word_loc'(X394),fun('type.wordSem.word_loc'(X394),fun(X393,bool))),X396),s(fun(X393,bool),X397),s(fun(X393,bool),X398),s(fun(X393,bool),X399)),X393))) )),
  inference(superposition,[],[f155,f176])).
fof(f3496,plain,(
  ( ! [X408,X405,X407,X401,X403,X404,X406,X402] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X401,bool),X406))) | s(X401,X408) = s(X401,sK5(X408,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X402),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X402),fun('type.wordSem.word_loc'(X402),fun(X401,bool))),X403),s(fun('type.wordSem.word_loc'(X402),fun('type.wordSem.word_loc'(X402),fun(X401,bool))),X404),s(fun(X401,bool),X405),s(fun(X401,bool),X406),s(fun(X401,bool),X407)),X401))) )),
  inference(superposition,[],[f419,f176])).
fof(f1219,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X104,bool),happ(s(fun('type.num.num',fun(X104,bool)),X106),s('type.num.num',X107))))) | s(X104,X108) = s(X104,sK5(X108,'const.ind_type.FCONS_3'(s(fun(X104,bool),X105),s(fun('type.num.num',fun(X104,bool)),X106),s('type.num.num','const.num.SUC_1'(s('type.num.num',X107)))),X104))) )),
  inference(superposition,[],[f419,f157])).
fof(f1218,plain,(
  ( ! [X103,X101,X99,X102,X100] : (p(s(bool,'|dollar|exists_1'(s(fun(X99,bool),happ(s(fun('type.num.num',fun(X99,bool)),X101),s('type.num.num',X102)))))) | s(X99,X103) = s(X99,sK5(X103,'const.ind_type.FCONS_3'(s(fun(X99,bool),X100),s(fun('type.num.num',fun(X99,bool)),X101),s('type.num.num','const.num.SUC_1'(s('type.num.num',X102)))),X99))) )),
  inference(superposition,[],[f155,f157])).
fof(f4135,plain,(
  ( ! [X405,X407,X401,X403,X404,X406,X402] : (p(s(bool,'|dollar|exists_1'(s(fun(X401,bool),happ(s(fun(X403,fun(X401,bool)),X404),s(X403,happ(s(fun(X402,X403),X405),s(X402,X406)))))))) | s(X401,X407) = s(X401,sK5(X407,happ(s(fun(X402,fun(X401,bool)),'_dst_x0x1_2'(s(fun(X403,fun(X401,bool)),X404),s(fun(X402,X403),X405))),s(X402,X406)),X401))) )),
  inference(superposition,[],[f155,f179])).
fof(f4056,plain,(
  ( ! [X412,X408,X410,X407,X409,X411,X406] : (p(s(bool,'|dollar|exists_1'(s(fun(X406,bool),happ(s(fun(X409,fun(X406,bool)),'_dst_x0x1_2'(s(fun(X407,fun(X406,bool)),X408),s(fun(X409,X407),X410))),s(X409,X411)))))) | s(X406,X412) = s(X406,sK5(X412,happ(s(fun(X407,fun(X406,bool)),X408),s(X407,happ(s(fun(X409,X407),X410),s(X409,X411)))),X406))) )),
  inference(superposition,[],[f155,f179])).
fof(f972,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,'|dollar|exists_1'(s(fun(X96,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X96,bool),X97),X98)))))) | s(X96,X99) = s(X96,sK5(X99,happ(s(fun('type.pair.prod'(fun(X96,bool),X97),fun(X96,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X96,bool),X97),X98)),X96))) )),
  inference(superposition,[],[f155,f169])).
fof(f970,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X88,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X88,bool),X89),X90))))) | s(X88,X91) = s(X88,sK5(X91,happ(s(fun('type.pair.prod'(fun(X88,bool),X89),fun(X88,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X88,bool),X89),X90)),X88))) )),
  inference(superposition,[],[f419,f169])).
fof(f850,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,sK5(X36,X37,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37)))) )),
  inference(superposition,[],[f131,f419])).
fof(f851,plain,(
  ( ! [X39,X38] : (~p(s(bool,X38)) | s(bool,'T_0') = s(bool,sK5(X38,X39,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X39)))) )),
  inference(superposition,[],[f132,f419])).
fof(f852,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X40) | ~p(s(bool,sK5(X40,X41,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X41)))) )),
  inference(superposition,[],[f135,f419])).
fof(f853,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | s(bool,'F_0') = s(bool,sK5(X42,X43,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f136,f419])).
fof(f854,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X46)) | s(bool,X46) = s(bool,sK5(X44,X45,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f165,f419])).
fof(f856,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | s(bool,X52) = s(bool,sK5(X50,X51,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X51)))) )),
  inference(superposition,[],[f168,f419])).
fof(f861,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | s(bool,X67) != s(bool,sK5(X65,X66,bool)) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66)))) )),
  inference(superposition,[],[f214,f419])).
fof(f862,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | s(bool,X70) = s(bool,sK5(X68,X69,bool)) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X69)))) )),
  inference(superposition,[],[f247,f419])).
fof(f3590,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,X90)) | s(bool,X92) = s(bool,sK5(X90,X91,bool)) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f3302,f419])).
fof(f6318,plain,(
  ( ! [X532,X534,X530,X533,X529,X531] : (s(X531,f(s(bool,'T_0'))) != s(X531,g(s(bool,'T_0'))) | s(fun(bool,X531),X529) = s(fun(bool,X531),X530) | s(bool,X532) = s(bool,'const.bool.DATATYPE_1'(s(X533,X534))) | s(bool,X532) = s(bool,sK7(X529,X530,X531,bool))) )),
  inference(superposition,[],[f173,f3843])).
fof(f6929,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,X49) = s(bool,'const.bool.DATATYPE_1'(s(X50,X51))) | s(bool,X47) = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f6237,f127])).
fof(f6237,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | s(bool,X47) != s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,X49) = s(bool,'const.bool.DATATYPE_1'(s(X50,X51))) | s(bool,X47) = s(bool,X49)) )),
  inference(superposition,[],[f214,f3843])).
fof(f6181,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.bool.DATATYPE_1'(s(X32,X33))) | s(bool,X28) = s(bool,X31)) )),
  inference(superposition,[],[f4534,f3843])).
fof(f6180,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,X23) != s(bool,X24) | s(bool,X24) = s(bool,'const.bool.DATATYPE_1'(s(bool,X24))) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(X26,X27))) | s(bool,X23) = s(bool,X25)) )),
  inference(superposition,[],[f1837,f3843])).
fof(f6175,plain,(
  ( ! [X561,X557,X559,X560,X562,X558] : (s(X559,f(s(bool,X560))) != s(X559,g(s(bool,X560))) | s(fun(bool,X559),X557) = s(fun(bool,X559),X558) | s(bool,X560) = s(bool,'const.bool.DATATYPE_1'(s(X561,X562))) | s(bool,'T_0') = s(bool,sK7(X557,X558,X559,bool))) )),
  inference(superposition,[],[f173,f3843])).
fof(f6894,plain,(
  ( ! [X231,X229,X230,X228,X232] : (s(bool,X228) != s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,X228) = s(bool,'const.bool.DATATYPE_1'(s(X231,X232)))) )),
  inference(subsumption_resolution,[],[f6125,f195])).
fof(f6125,plain,(
  ( ! [X231,X229,X230,X228,X232] : (s(bool,X228) != s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,X228) = s(bool,'const.bool.DATATYPE_1'(s(X231,X232))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1774,f3843])).
fof(f6893,plain,(
  ( ! [X227,X225,X226] : (s(bool,'T_0') != s(bool,X225) | s(bool,X225) = s(bool,'const.bool.DATATYPE_1'(s(X226,X227)))) )),
  inference(subsumption_resolution,[],[f6124,f195])).
fof(f6124,plain,(
  ( ! [X227,X225,X226] : (s(bool,'T_0') != s(bool,X225) | s(bool,X225) = s(bool,'const.bool.DATATYPE_1'(s(X226,X227))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f195,f3843])).
fof(f5826,plain,(
  ( ! [X206,X204,X207,X205,X203] : (s(bool,'T_0') = s(bool,X203) | s(bool,X205) = s(bool,'const.bool.DATATYPE_1'(s(X206,X207))) | s(bool,X205) = s(bool,sK5(X203,X204,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X204)))) )),
  inference(superposition,[],[f3843,f419])).
fof(f5825,plain,(
  ( ! [X198,X202,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,X198) | s(bool,X200) = s(bool,'const.bool.DATATYPE_1'(s(X201,X202))) | s(bool,X200) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X197),'const.pair.,_2'(s(bool,X198),s(X197,X199)))))) )),
  inference(superposition,[],[f3843,f172])).
fof(f5824,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (s(bool,'T_0') = s(bool,X193) | s(bool,X194) = s(bool,'const.bool.DATATYPE_1'(s(X195,X196))) | s(bool,X194) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X191,bool),'const.pair.,_2'(s(X191,X192),s(bool,X193)))))) )),
  inference(superposition,[],[f3843,f171])).
fof(f5822,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,'T_0') = s(bool,X180) | s(bool,X182) = s(bool,'const.bool.DATATYPE_1'(s(X183,X184))) | s(bool,X182) = s(bool,'const.ind_type.FCONS_3'(s(bool,X180),s(fun('type.num.num',bool),X181),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3843,f152])).
fof(f5821,plain,(
  ( ! [X177,X175,X173,X171,X178,X176,X174,X172,X179] : (s(bool,'T_0') = s(bool,X175) | s(bool,X177) = s(bool,'const.bool.DATATYPE_1'(s(X178,X179))) | s(bool,X177) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X171),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X171),fun('type.wordSem.word_loc'(X171),bool)),X172),s(fun('type.wordSem.word_loc'(X171),fun('type.wordSem.word_loc'(X171),bool)),X173),s(bool,X174),s(bool,X175),s(bool,X176)))) )),
  inference(superposition,[],[f3843,f176])).
fof(f5820,plain,(
  ( ! [X167,X165,X163,X169,X166,X164,X162,X170,X168] : (s(bool,'T_0') = s(bool,X167) | s(bool,X168) = s(bool,'const.bool.DATATYPE_1'(s(X169,X170))) | s(bool,X168) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X162),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X162),fun('type.wordSem.word_loc'(X162),bool)),X163),s(fun('type.wordSem.word_loc'(X162),fun('type.wordSem.word_loc'(X162),bool)),X164),s(bool,X165),s(bool,X166),s(bool,X167)))) )),
  inference(superposition,[],[f3843,f175])).
fof(f5819,plain,(
  ( ! [X158,X156,X154,X161,X159,X157,X155,X153,X160] : (s(bool,'T_0') = s(bool,X156) | s(bool,X159) = s(bool,'const.bool.DATATYPE_1'(s(X160,X161))) | s(bool,X159) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X153),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X153),fun('type.wordSem.word_loc'(X153),bool)),X154),s(fun('type.wordSem.word_loc'(X153),fun('type.wordSem.word_loc'(X153),bool)),X155),s(bool,X156),s(bool,X157),s(bool,X158)))) )),
  inference(superposition,[],[f3843,f174])).
fof(f6494,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,X21) = s(bool,X23) | s(bool,X21) = s(bool,'const.bool.DATATYPE_1'(s(bool,X21)))) )),
  inference(duplicate_literal_removal,[],[f5798])).
fof(f5798,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,X21) = s(bool,X23) | s(bool,X21) = s(bool,'const.bool.DATATYPE_1'(s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f3843,f553])).
fof(f5780,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,X14) = s(bool,X15) | s(bool,X16) = s(bool,'const.bool.DATATYPE_1'(s(X17,X18))) | s(bool,X15) = s(bool,X16) | s(bool,X14) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f3843,f211])).
fof(f5778,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.DATATYPE_1'(s(X3,X4))) | s(bool,X1) = s(bool,X2) | s(bool,X5) = s(bool,'const.bool.DATATYPE_1'(s(X6,X7))) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f3843,f3843])).
fof(f5777,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s(bool,X246) = s(bool,X248) | s(bool,X248) = s(bool,'const.bool.DATATYPE_1'(s(X249,X250))) | s(bool,'T_0') = s(bool,sK5(X246,X247,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X247)))) )),
  inference(superposition,[],[f3843,f419])).
fof(f5776,plain,(
  ( ! [X243,X241,X244,X242,X240,X245] : (s(bool,X241) = s(bool,X243) | s(bool,X243) = s(bool,'const.bool.DATATYPE_1'(s(X244,X245))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X240),'const.pair.,_2'(s(bool,X241),s(X240,X242)))))) )),
  inference(superposition,[],[f3843,f172])).
fof(f5775,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (s(bool,X236) = s(bool,X237) | s(bool,X237) = s(bool,'const.bool.DATATYPE_1'(s(X238,X239))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X234,bool),'const.pair.,_2'(s(X234,X235),s(bool,X236)))))) )),
  inference(superposition,[],[f3843,f171])).
fof(f5773,plain,(
  ( ! [X227,X225,X223,X226,X224] : (s(bool,X223) = s(bool,X225) | s(bool,X225) = s(bool,'const.bool.DATATYPE_1'(s(X226,X227))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X223),s(fun('type.num.num',bool),X224),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3843,f152])).
fof(f5772,plain,(
  ( ! [X218,X216,X222,X220,X215,X219,X217,X221,X214] : (s(bool,X218) = s(bool,X220) | s(bool,X220) = s(bool,'const.bool.DATATYPE_1'(s(X221,X222))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X214),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X214),fun('type.wordSem.word_loc'(X214),bool)),X215),s(fun('type.wordSem.word_loc'(X214),fun('type.wordSem.word_loc'(X214),bool)),X216),s(bool,X217),s(bool,X218),s(bool,X219)))) )),
  inference(superposition,[],[f3843,f176])).
fof(f5771,plain,(
  ( ! [X212,X210,X208,X206,X213,X211,X209,X207,X205] : (s(bool,X210) = s(bool,X211) | s(bool,X211) = s(bool,'const.bool.DATATYPE_1'(s(X212,X213))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X205),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X205),fun('type.wordSem.word_loc'(X205),bool)),X206),s(fun('type.wordSem.word_loc'(X205),fun('type.wordSem.word_loc'(X205),bool)),X207),s(bool,X208),s(bool,X209),s(bool,X210)))) )),
  inference(superposition,[],[f3843,f175])).
fof(f5770,plain,(
  ( ! [X198,X196,X204,X202,X200,X199,X197,X203,X201] : (s(bool,X199) = s(bool,X202) | s(bool,X202) = s(bool,'const.bool.DATATYPE_1'(s(X203,X204))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X196),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X196),fun('type.wordSem.word_loc'(X196),bool)),X197),s(fun('type.wordSem.word_loc'(X196),fun('type.wordSem.word_loc'(X196),bool)),X198),s(bool,X199),s(bool,X200),s(bool,X201)))) )),
  inference(superposition,[],[f3843,f174])).
fof(f5769,plain,(
  ( ! [X191,X189,X194,X192,X190,X188,X195,X193] : (s(bool,X193) = s(bool,happ(s(fun(X189,bool),X190),s(X189,happ(s(fun(X188,X189),X191),s(X188,X192))))) | s(bool,X193) = s(bool,'const.bool.DATATYPE_1'(s(X194,X195))) | s(bool,'T_0') = s(bool,happ(s(fun(X188,bool),'_dst_x0x1_2'(s(fun(X189,bool),X190),s(fun(X188,X189),X191))),s(X188,X192)))) )),
  inference(superposition,[],[f3843,f179])).
fof(f5768,plain,(
  ( ! [X187,X185,X182,X180,X186,X184,X183,X181] : (s(bool,X185) = s(bool,happ(s(fun(X182,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X182,X180),X183))),s(X182,X184))) | s(bool,X185) = s(bool,'const.bool.DATATYPE_1'(s(X186,X187))) | s(bool,'T_0') = s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X182,X180),X183),s(X182,X184)))))) )),
  inference(superposition,[],[f3843,f179])).
fof(f6660,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X114) = s(bool,X115) | s(bool,X115) = s(bool,'const.bool.DATATYPE_1'(s(X116,X117))) | s(bool,X114) = s(bool,X118) | s(bool,'F_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f5755,f195])).
fof(f5755,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X114) = s(bool,X115) | s(bool,X115) = s(bool,'const.bool.DATATYPE_1'(s(X116,X117))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X114) = s(bool,X118) | s(bool,'F_0') = s(bool,X118)) )),
  inference(superposition,[],[f3843,f3673])).
fof(f5616,plain,(
  ( ! [X241,X239,X237,X240,X238] : (s(bool,X237) = s(bool,'const.bool.DATATYPE_1'(s(X239,X240))) | s(bool,X241) = s(bool,sK5(X237,X238,bool)) | s(bool,'T_0') = s(bool,X241) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X238)))) )),
  inference(superposition,[],[f3843,f419])).
fof(f5615,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (s(bool,X232) = s(bool,'const.bool.DATATYPE_1'(s(X234,X235))) | s(bool,X236) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X231),'const.pair.,_2'(s(bool,X232),s(X231,X233))))) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f3843,f172])).
fof(f5614,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,X227) = s(bool,'const.bool.DATATYPE_1'(s(X228,X229))) | s(bool,X230) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X225,bool),'const.pair.,_2'(s(X225,X226),s(bool,X227))))) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f3843,f171])).
fof(f5612,plain,(
  ( ! [X218,X216,X215,X217,X214] : (s(bool,X214) = s(bool,'const.bool.DATATYPE_1'(s(X216,X217))) | s(bool,X218) = s(bool,'const.ind_type.FCONS_3'(s(bool,X214),s(fun('type.num.num',bool),X215),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f3843,f152])).
fof(f5611,plain,(
  ( ! [X212,X210,X208,X206,X213,X211,X209,X207,X205] : (s(bool,X209) = s(bool,'const.bool.DATATYPE_1'(s(X211,X212))) | s(bool,X213) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X205),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X205),fun('type.wordSem.word_loc'(X205),bool)),X206),s(fun('type.wordSem.word_loc'(X205),fun('type.wordSem.word_loc'(X205),bool)),X207),s(bool,X208),s(bool,X209),s(bool,X210))) | s(bool,'T_0') = s(bool,X213)) )),
  inference(superposition,[],[f3843,f176])).
fof(f5610,plain,(
  ( ! [X198,X196,X204,X202,X200,X199,X197,X203,X201] : (s(bool,X201) = s(bool,'const.bool.DATATYPE_1'(s(X202,X203))) | s(bool,X204) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X196),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X196),fun('type.wordSem.word_loc'(X196),bool)),X197),s(fun('type.wordSem.word_loc'(X196),fun('type.wordSem.word_loc'(X196),bool)),X198),s(bool,X199),s(bool,X200),s(bool,X201))) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f3843,f175])).
fof(f5609,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X195,X193] : (s(bool,X190) = s(bool,'const.bool.DATATYPE_1'(s(X193,X194))) | s(bool,X195) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X187),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X187),fun('type.wordSem.word_loc'(X187),bool)),X188),s(fun('type.wordSem.word_loc'(X187),fun('type.wordSem.word_loc'(X187),bool)),X189),s(bool,X190),s(bool,X191),s(bool,X192))) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f3843,f174])).
fof(f5175,plain,(
  ( ! [X130,X128,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,X130) = s(bool,sK5(X128,X129,bool)) | s(bool,'F_0') = s(bool,X130) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X129)))) )),
  inference(superposition,[],[f4534,f419])).
fof(f5174,plain,(
  ( ! [X127,X125,X126,X124] : (s(bool,'T_0') != s(bool,X125) | s(bool,X127) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X125),s(X124,X126))))) | s(bool,'F_0') = s(bool,X127)) )),
  inference(superposition,[],[f4534,f172])).
fof(f5173,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') != s(bool,X122) | s(bool,X123) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X120,bool),'const.pair.,_2'(s(X120,X121),s(bool,X122))))) | s(bool,'F_0') = s(bool,X123)) )),
  inference(superposition,[],[f4534,f171])).
fof(f5171,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,X113) | s(bool,X115) = s(bool,'const.ind_type.FCONS_3'(s(bool,X113),s(fun('type.num.num',bool),X114),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f4534,f152])).
fof(f5170,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s(bool,'T_0') != s(bool,X110) | s(bool,X112) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X106),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X106),fun('type.wordSem.word_loc'(X106),bool)),X107),s(fun('type.wordSem.word_loc'(X106),fun('type.wordSem.word_loc'(X106),bool)),X108),s(bool,X109),s(bool,X110),s(bool,X111))) | s(bool,'F_0') = s(bool,X112)) )),
  inference(superposition,[],[f4534,f176])).
fof(f5169,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,X105) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X99),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X99),fun('type.wordSem.word_loc'(X99),bool)),X100),s(fun('type.wordSem.word_loc'(X99),fun('type.wordSem.word_loc'(X99),bool)),X101),s(bool,X102),s(bool,X103),s(bool,X104))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(superposition,[],[f4534,f175])).
fof(f5168,plain,(
  ( ! [X94,X92,X97,X95,X93,X98,X96] : (s(bool,'T_0') != s(bool,X95) | s(bool,X98) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X92),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X92),fun('type.wordSem.word_loc'(X92),bool)),X93),s(fun('type.wordSem.word_loc'(X92),fun('type.wordSem.word_loc'(X92),bool)),X94),s(bool,X95),s(bool,X96),s(bool,X97))) | s(bool,'F_0') = s(bool,X98)) )),
  inference(superposition,[],[f4534,f174])).
fof(f5167,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'T_0') != s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90))))) | s(bool,X91) = s(bool,happ(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))),s(X86,X90))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f4534,f179])).
fof(f5166,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,happ(s(fun(X82,bool),'_dst_x0x1_2'(s(fun(X80,bool),X81),s(fun(X82,X80),X83))),s(X82,X84))) | s(bool,X85) = s(bool,happ(s(fun(X80,bool),X81),s(X80,happ(s(fun(X82,X80),X83),s(X82,X84))))) | s(bool,'F_0') = s(bool,X85)) )),
  inference(superposition,[],[f4534,f179])).
fof(f5127,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.DATATYPE_1'(s(bool,X3))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f4534,f553])).
fof(f4665,plain,(
  ( ! [X315,X316,X318,X317] : (s(X317,f(s(bool,'F_0'))) != s(X317,g(s(bool,'F_0'))) | s(fun(bool,X317),X315) = s(fun(bool,X317),X316) | s(bool,X318) = s(bool,sK7(X315,X316,X317,bool)) | s(bool,'F_0') = s(bool,X318)) )),
  inference(superposition,[],[f173,f3673])).
fof(f5060,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f4587,f129])).
fof(f4587,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f165,f3673])).
fof(f4535,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X13) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f1774,f3673])).
fof(f4534,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f195,f3673])).
fof(f4533,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X5) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f135,f3673])).
fof(f4441,plain,(
  ( ! [X333,X331,X332,X334] : (s(X333,f(s(bool,X334))) != s(X333,g(s(bool,X334))) | s(fun(bool,X333),X331) = s(fun(bool,X333),X332) | s(bool,'F_0') = s(bool,X334) | s(bool,'F_0') = s(bool,sK7(X331,X332,X333,bool))) )),
  inference(superposition,[],[f173,f3673])).
fof(f4309,plain,(
  ( ! [X114,X112,X113] : (s(bool,'F_0') = s(bool,X112) | s(bool,X114) = s(bool,sK5(X112,X113,bool)) | s(bool,'F_0') = s(bool,X114) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) )),
  inference(superposition,[],[f3673,f419])).
fof(f4308,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,X109) | s(bool,X111) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X108),'const.pair.,_2'(s(bool,X109),s(X108,X110))))) | s(bool,'F_0') = s(bool,X111)) )),
  inference(superposition,[],[f3673,f172])).
fof(f4307,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'F_0') = s(bool,X106) | s(bool,X107) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X104,bool),'const.pair.,_2'(s(X104,X105),s(bool,X106))))) | s(bool,'F_0') = s(bool,X107)) )),
  inference(superposition,[],[f3673,f171])).
fof(f4304,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X96] : (s(bool,'F_0') = s(bool,X94) | s(bool,X96) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X90),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X90),fun('type.wordSem.word_loc'(X90),bool)),X91),s(fun('type.wordSem.word_loc'(X90),fun('type.wordSem.word_loc'(X90),bool)),X92),s(bool,X93),s(bool,X94),s(bool,X95))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f3673,f176])).
fof(f4303,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84] : (s(bool,'F_0') = s(bool,X88) | s(bool,X89) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X83),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X83),fun('type.wordSem.word_loc'(X83),bool)),X84),s(fun('type.wordSem.word_loc'(X83),fun('type.wordSem.word_loc'(X83),bool)),X85),s(bool,X86),s(bool,X87),s(bool,X88))) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f3673,f175])).
fof(f4302,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (s(bool,'F_0') = s(bool,X79) | s(bool,X82) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X76),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X76),fun('type.wordSem.word_loc'(X76),bool)),X77),s(fun('type.wordSem.word_loc'(X76),fun('type.wordSem.word_loc'(X76),bool)),X78),s(bool,X79),s(bool,X80),s(bool,X81))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f3673,f174])).
fof(f4267,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f3673,f3673])).
fof(f4229,plain,(
  ( ! [X127,X128,X126] : (s(bool,X126) = s(bool,X128) | s(bool,'F_0') = s(bool,X128) | s(bool,'F_0') = s(bool,sK5(X126,X127,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X127)))) )),
  inference(superposition,[],[f3673,f419])).
fof(f4228,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X123) = s(bool,X125) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X122),'const.pair.,_2'(s(bool,X123),s(X122,X124)))))) )),
  inference(superposition,[],[f3673,f172])).
fof(f4227,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X120) = s(bool,X121) | s(bool,'F_0') = s(bool,X121) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X118,bool),'const.pair.,_2'(s(X118,X119),s(bool,X120)))))) )),
  inference(superposition,[],[f3673,f171])).
fof(f4225,plain,(
  ( ! [X111,X112,X113] : (s(bool,X111) = s(bool,X113) | s(bool,'F_0') = s(bool,X113) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X111),s(fun('type.num.num',bool),X112),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3673,f152])).
fof(f4224,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s(bool,X108) = s(bool,X110) | s(bool,'F_0') = s(bool,X110) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X104),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X104),fun('type.wordSem.word_loc'(X104),bool)),X105),s(fun('type.wordSem.word_loc'(X104),fun('type.wordSem.word_loc'(X104),bool)),X106),s(bool,X107),s(bool,X108),s(bool,X109)))) )),
  inference(superposition,[],[f3673,f176])).
fof(f4223,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98] : (s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X97),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X97),fun('type.wordSem.word_loc'(X97),bool)),X98),s(fun('type.wordSem.word_loc'(X97),fun('type.wordSem.word_loc'(X97),bool)),X99),s(bool,X100),s(bool,X101),s(bool,X102)))) )),
  inference(superposition,[],[f3673,f175])).
fof(f4222,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X96] : (s(bool,X93) = s(bool,X96) | s(bool,'F_0') = s(bool,X96) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X90),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X90),fun('type.wordSem.word_loc'(X90),bool)),X91),s(fun('type.wordSem.word_loc'(X90),fun('type.wordSem.word_loc'(X90),bool)),X92),s(bool,X93),s(bool,X94),s(bool,X95)))) )),
  inference(superposition,[],[f3673,f174])).
fof(f4221,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X89) = s(bool,happ(s(fun(X85,bool),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88))))) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,happ(s(fun(X84,bool),'_dst_x0x1_2'(s(fun(X85,bool),X86),s(fun(X84,X85),X87))),s(X84,X88)))) )),
  inference(superposition,[],[f3673,f179])).
fof(f4220,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (s(bool,X83) = s(bool,happ(s(fun(X80,bool),'_dst_x0x1_2'(s(fun(X78,bool),X79),s(fun(X80,X78),X81))),s(X80,X82))) | s(bool,'F_0') = s(bool,X83) | s(bool,'F_0') = s(bool,happ(s(fun(X78,bool),X79),s(X78,happ(s(fun(X80,X78),X81),s(X80,X82)))))) )),
  inference(superposition,[],[f3673,f179])).
fof(f4152,plain,(
  ( ! [X557,X553,X555,X551,X556,X558,X552,X554,X550] : (s('type.wordSem.result'(X550),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X550),X557),s('type.wordSem.word_loc'(X550),X558))) != s('type.wordSem.result'(X550),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X550),X556),s('type.wordSem.word_loc'(X550),happ(s(fun(X552,'type.wordSem.word_loc'(X550)),X553),s(X552,happ(s(fun(X551,X552),X554),s(X551,X555))))))) | s('type.wordSem.word_loc'(X550),X558) = s('type.wordSem.word_loc'(X550),happ(s(fun(X551,'type.wordSem.word_loc'(X550)),'_dst_x0x1_2'(s(fun(X552,'type.wordSem.word_loc'(X550)),X553),s(fun(X551,X552),X554))),s(X551,X555)))) )),
  inference(superposition,[],[f163,f179])).
fof(f4147,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X506,X513] : (s('type.wordSem.result'(X505),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X505),X512),s('type.wordSem.word_loc'(X505),X513))) != s('type.wordSem.result'(X505),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X505),happ(s(fun(X507,'type.wordSem.word_loc'(X505)),X508),s(X507,happ(s(fun(X506,X507),X509),s(X506,X510))))),s('type.wordSem.word_loc'(X505),X511))) | s('type.wordSem.word_loc'(X505),X512) = s('type.wordSem.word_loc'(X505),happ(s(fun(X506,'type.wordSem.word_loc'(X505)),'_dst_x0x1_2'(s(fun(X507,'type.wordSem.word_loc'(X505)),X508),s(fun(X506,X507),X509))),s(X506,X510)))) )),
  inference(superposition,[],[f162,f179])).
fof(f4136,plain,(
  ( ! [X412,X414,X408,X410,X413,X409,X411] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X408,bool),happ(s(fun(X410,fun(X408,bool)),X411),s(X410,happ(s(fun(X409,X410),X412),s(X409,X413))))))) | s(X408,X414) = s(X408,sK5(X414,happ(s(fun(X409,fun(X408,bool)),'_dst_x0x1_2'(s(fun(X410,fun(X408,bool)),X411),s(fun(X409,X410),X412))),s(X409,X413)),X408))) )),
  inference(superposition,[],[f419,f179])).
fof(f4126,plain,(
  ( ! [X344,X346,X343,X345,X347,X342] : (s(bool,X347) != s(bool,happ(s(fun(X343,bool),X344),s(X343,happ(s(fun(X342,X343),X345),s(X342,X346))))) | s(bool,'T_0') = s(bool,happ(s(fun(X342,bool),'_dst_x0x1_2'(s(fun(X343,bool),X344),s(fun(X342,X343),X345))),s(X342,X346))) | s(bool,'F_0') = s(bool,X347)) )),
  inference(superposition,[],[f2056,f179])).
fof(f4125,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (s(bool,X341) != s(bool,happ(s(fun(X337,bool),X338),s(X337,happ(s(fun(X336,X337),X339),s(X336,X340))))) | s(bool,'T_0') = s(bool,X341) | s(bool,'F_0') = s(bool,happ(s(fun(X336,bool),'_dst_x0x1_2'(s(fun(X337,bool),X338),s(fun(X336,X337),X339))),s(X336,X340)))) )),
  inference(superposition,[],[f2056,f179])).
fof(f4121,plain,(
  ( ! [X313,X315,X316,X312,X314,X317] : (s(bool,'F_0') != s(bool,happ(s(fun(X313,bool),X314),s(X313,happ(s(fun(X312,X313),X315),s(X312,X316))))) | s(bool,X317) = s(bool,happ(s(fun(X312,bool),'_dst_x0x1_2'(s(fun(X313,bool),X314),s(fun(X312,X313),X315))),s(X312,X316))) | s(bool,'T_0') = s(bool,X317)) )),
  inference(superposition,[],[f1774,f179])).
fof(f4178,plain,(
  ( ! [X271,X273,X270,X272,X274] : (s(bool,'T_0') != s(bool,happ(s(fun(X271,bool),X272),s(X271,happ(s(fun(X270,X271),X273),s(X270,X274))))) | s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),'_dst_x0x1_2'(s(fun(X271,bool),X272),s(fun(X270,X271),X273))),s(X270,X274)))) )),
  inference(forward_demodulation,[],[f4114,f148])).
fof(f4114,plain,(
  ( ! [X271,X273,X275,X270,X276,X272,X274] : (s(bool,'const.bool.DATATYPE_1'(s(X275,X276))) != s(bool,happ(s(fun(X271,bool),X272),s(X271,happ(s(fun(X270,X271),X273),s(X270,X274))))) | s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),'_dst_x0x1_2'(s(fun(X271,bool),X272),s(fun(X270,X271),X273))),s(X270,X274)))) )),
  inference(superposition,[],[f381,f179])).
fof(f4073,plain,(
  ( ! [X561,X563,X557,X559,X555,X560,X562,X556,X558] : (s('type.wordSem.result'(X555),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X555),X562),s('type.wordSem.word_loc'(X555),X563))) != s('type.wordSem.result'(X555),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X555),X561),s('type.wordSem.word_loc'(X555),happ(s(fun(X558,'type.wordSem.word_loc'(X555)),'_dst_x0x1_2'(s(fun(X556,'type.wordSem.word_loc'(X555)),X557),s(fun(X558,X556),X559))),s(X558,X560))))) | s('type.wordSem.word_loc'(X555),X563) = s('type.wordSem.word_loc'(X555),happ(s(fun(X556,'type.wordSem.word_loc'(X555)),X557),s(X556,happ(s(fun(X558,X556),X559),s(X558,X560)))))) )),
  inference(superposition,[],[f163,f179])).
fof(f4068,plain,(
  ( ! [X511,X516,X518,X512,X514,X510,X517,X513,X515] : (s('type.wordSem.result'(X510),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X510),X517),s('type.wordSem.word_loc'(X510),X518))) != s('type.wordSem.result'(X510),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X510),happ(s(fun(X513,'type.wordSem.word_loc'(X510)),'_dst_x0x1_2'(s(fun(X511,'type.wordSem.word_loc'(X510)),X512),s(fun(X513,X511),X514))),s(X513,X515))),s('type.wordSem.word_loc'(X510),X516))) | s('type.wordSem.word_loc'(X510),X517) = s('type.wordSem.word_loc'(X510),happ(s(fun(X511,'type.wordSem.word_loc'(X510)),X512),s(X511,happ(s(fun(X513,X511),X514),s(X513,X515)))))) )),
  inference(superposition,[],[f162,f179])).
fof(f4057,plain,(
  ( ! [X414,X416,X418,X413,X415,X417,X419] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X413,bool),happ(s(fun(X416,fun(X413,bool)),'_dst_x0x1_2'(s(fun(X414,fun(X413,bool)),X415),s(fun(X416,X414),X417))),s(X416,X418))))) | s(X413,X419) = s(X413,sK5(X419,happ(s(fun(X414,fun(X413,bool)),X415),s(X414,happ(s(fun(X416,X414),X417),s(X416,X418)))),X413))) )),
  inference(superposition,[],[f419,f179])).
fof(f4047,plain,(
  ( ! [X350,X352,X349,X351,X347,X348] : (s(bool,X352) != s(bool,happ(s(fun(X349,bool),'_dst_x0x1_2'(s(fun(X347,bool),X348),s(fun(X349,X347),X350))),s(X349,X351))) | s(bool,'T_0') = s(bool,happ(s(fun(X347,bool),X348),s(X347,happ(s(fun(X349,X347),X350),s(X349,X351))))) | s(bool,'F_0') = s(bool,X352)) )),
  inference(superposition,[],[f2056,f179])).
fof(f4046,plain,(
  ( ! [X344,X346,X341,X343,X345,X342] : (s(bool,X346) != s(bool,happ(s(fun(X343,bool),'_dst_x0x1_2'(s(fun(X341,bool),X342),s(fun(X343,X341),X344))),s(X343,X345))) | s(bool,'T_0') = s(bool,X346) | s(bool,'F_0') = s(bool,happ(s(fun(X341,bool),X342),s(X341,happ(s(fun(X343,X341),X344),s(X343,X345)))))) )),
  inference(superposition,[],[f2056,f179])).
fof(f4042,plain,(
  ( ! [X321,X318,X320,X322,X317,X319] : (s(bool,'F_0') != s(bool,happ(s(fun(X319,bool),'_dst_x0x1_2'(s(fun(X317,bool),X318),s(fun(X319,X317),X320))),s(X319,X321))) | s(bool,X322) = s(bool,happ(s(fun(X317,bool),X318),s(X317,happ(s(fun(X319,X317),X320),s(X319,X321))))) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f1774,f179])).
fof(f4166,plain,(
  ( ! [X277,X279,X275,X276,X278] : (s(bool,'T_0') != s(bool,happ(s(fun(X277,bool),'_dst_x0x1_2'(s(fun(X275,bool),X276),s(fun(X277,X275),X278))),s(X277,X279))) | s(bool,'T_0') = s(bool,happ(s(fun(X275,bool),X276),s(X275,happ(s(fun(X277,X275),X278),s(X277,X279)))))) )),
  inference(forward_demodulation,[],[f4035,f148])).
fof(f4035,plain,(
  ( ! [X280,X277,X279,X275,X281,X276,X278] : (s(bool,'const.bool.DATATYPE_1'(s(X280,X281))) != s(bool,happ(s(fun(X277,bool),'_dst_x0x1_2'(s(fun(X275,bool),X276),s(fun(X277,X275),X278))),s(X277,X279))) | s(bool,'T_0') = s(bool,happ(s(fun(X275,bool),X276),s(X275,happ(s(fun(X277,X275),X278),s(X277,X279)))))) )),
  inference(superposition,[],[f381,f179])).
fof(f4163,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70))) | s(bool,'F_0') = s(bool,happ(s(fun(X66,bool),X67),s(X66,happ(s(fun(X68,X66),X69),s(X68,X70)))))) )),
  inference(forward_demodulation,[],[f3932,f148])).
fof(f3932,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.bool.DATATYPE_1'(s(X71,X72))) = s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70))) | s(bool,'F_0') = s(bool,happ(s(fun(X66,bool),X67),s(X66,happ(s(fun(X68,X66),X69),s(X68,X70)))))) )),
  inference(superposition,[],[f179,f3606])).
fof(f3931,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),'_dst_x0x1_2'(s(fun(X60,bool),X61),s(fun(X62,X60),X63))),s(X62,X64))) | s(bool,X65) = s(bool,'const.bool.DATATYPE_1'(s(bool,X65))) | s(bool,X65) = s(bool,happ(s(fun(X60,bool),X61),s(X60,happ(s(fun(X62,X60),X63),s(X62,X64)))))) )),
  inference(superposition,[],[f179,f553])).
fof(f3929,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,X53) = s(bool,happ(s(fun(X50,bool),'_dst_x0x1_2'(s(fun(X48,bool),X49),s(fun(X50,X48),X51))),s(X50,X52))) | s(bool,X53) = s(bool,'const.bool.DATATYPE_1'(s(bool,X53))) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,happ(s(fun(X50,X48),X51),s(X50,X52)))))) )),
  inference(superposition,[],[f179,f553])).
fof(f3926,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X35) = s(bool,happ(s(fun(X32,bool),'_dst_x0x1_2'(s(fun(X30,bool),X31),s(fun(X32,X30),X33))),s(X32,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34)))))) )),
  inference(superposition,[],[f179,f211])).
fof(f3925,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'_dst_x0x1_2'(s(fun(X24,bool),X25),s(fun(X26,X24),X27))),s(X26,X28))) | s(bool,X29) = s(bool,happ(s(fun(X24,bool),X25),s(X24,happ(s(fun(X26,X24),X27),s(X26,X28))))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f179,f211])).
fof(f3924,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(X21,bool),'_dst_x0x1_2'(s(fun(X19,bool),X20),s(fun(X21,X19),X22))),s(X21,X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,happ(s(fun(X21,X19),X22),s(X21,X23)))))) )),
  inference(superposition,[],[f179,f130])).
fof(f3923,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X13,X11),X14),s(X13,happ(s(fun(X12,X13),X15),s(X12,happ(s(fun(X16,X12),X17),s(X16,X18))))))) = s(X11,happ(s(fun(X16,X11),'_dst_x0x1_2'(s(fun(X12,X11),'_dst_x0x1_2'(s(fun(X13,X11),X14),s(fun(X12,X13),X15))),s(fun(X16,X12),X17))),s(X16,X18)))) )),
  inference(superposition,[],[f179,f179])).
fof(f3922,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),happ(s(fun(X2,'type.pair.prod'(X0,X1)),X3),s(X2,X4))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s(fun(X2,'type.pair.prod'(X0,X1)),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f179,f169])).
fof(f3921,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(X19,happ(s(fun(X13,X19),X20),s(X13,happ(s(fun(X15,X13),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))))) = s(X19,happ(s(fun(X14,X19),'_dst_x0x1_2'(s(fun(X13,X19),X20),s(fun(X14,X13),'_dst_x0x1_2'(s(fun(X15,X13),X16),s(fun(X14,X15),X17))))),s(X14,X18)))) )),
  inference(superposition,[],[f179,f179])).
fof(f3920,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s(X11,happ(s(fun(X5,X11),X12),s(X5,happ(s(fun(X8,X5),'_dst_x0x1_2'(s(fun(X6,X5),X7),s(fun(X8,X6),X9))),s(X8,X10))))) = s(X11,happ(s(fun(X6,X11),'_dst_x0x1_2'(s(fun(X5,X11),X12),s(fun(X6,X5),X7))),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))))) )),
  inference(superposition,[],[f179,f179])).
fof(f3892,plain,(
  ( ! [X370,X369,X371] : (s(X371,f(s(bool,'T_0'))) != s(X371,g(s(bool,'T_0'))) | s(fun(bool,X371),X369) = s(fun(bool,X371),X370) | s(bool,'F_0') = s(bool,sK7(X369,X370,X371,bool))) )),
  inference(forward_demodulation,[],[f3758,f148])).
fof(f3758,plain,(
  ( ! [X372,X370,X373,X369,X371] : (s(X371,f(s(bool,'const.bool.DATATYPE_1'(s(X372,X373))))) != s(X371,g(s(bool,'const.bool.DATATYPE_1'(s(X372,X373))))) | s(fun(bool,X371),X369) = s(fun(bool,X371),X370) | s(bool,'F_0') = s(bool,sK7(X369,X370,X371,bool))) )),
  inference(superposition,[],[f173,f3606])).
fof(f3673,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f3606,f3606])).
fof(f3641,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X123) = s(bool,'const.bool.DATATYPE_1'(s(X125,X126))) | s(bool,'F_0') = s(bool,sK5(X123,X124,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X124)))) )),
  inference(superposition,[],[f3606,f419])).
fof(f3640,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X119) = s(bool,'const.bool.DATATYPE_1'(s(X121,X122))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120)))))) )),
  inference(superposition,[],[f3606,f172])).
fof(f3639,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X115) = s(bool,'const.bool.DATATYPE_1'(s(X116,X117))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X115)))))) )),
  inference(superposition,[],[f3606,f171])).
fof(f3637,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) = s(bool,'const.bool.DATATYPE_1'(s(X106,X107))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X104),s(fun('type.num.num',bool),X105),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3606,f152])).
fof(f3636,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (s(bool,X100) = s(bool,'const.bool.DATATYPE_1'(s(X102,X103))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X96),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X96),fun('type.wordSem.word_loc'(X96),bool)),X97),s(fun('type.wordSem.word_loc'(X96),fun('type.wordSem.word_loc'(X96),bool)),X98),s(bool,X99),s(bool,X100),s(bool,X101)))) )),
  inference(superposition,[],[f3606,f176])).
fof(f3635,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s(bool,X93) = s(bool,'const.bool.DATATYPE_1'(s(X94,X95))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X88),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X88),fun('type.wordSem.word_loc'(X88),bool)),X89),s(fun('type.wordSem.word_loc'(X88),fun('type.wordSem.word_loc'(X88),bool)),X90),s(bool,X91),s(bool,X92),s(bool,X93)))) )),
  inference(superposition,[],[f3606,f175])).
fof(f3634,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X83) = s(bool,'const.bool.DATATYPE_1'(s(X86,X87))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X80),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X80),fun('type.wordSem.word_loc'(X80),bool)),X81),s(fun('type.wordSem.word_loc'(X80),fun('type.wordSem.word_loc'(X80),bool)),X82),s(bool,X83),s(bool,X84),s(bool,X85)))) )),
  inference(superposition,[],[f3606,f174])).
fof(f3843,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) = s(bool,'const.bool.DATATYPE_1'(s(X73,X74))) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f3842,f195])).
fof(f3842,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X72) = s(bool,'const.bool.DATATYPE_1'(s(X73,X74))) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(forward_demodulation,[],[f3632,f148])).
fof(f3632,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) = s(bool,'const.bool.DATATYPE_1'(s(X73,X74))) | s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(bool,X72))) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f3606,f553])).
fof(f3606,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) = s(bool,'const.bool.DATATYPE_1'(s(X43,X44))) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f3579,f127])).
fof(f3579,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | s(bool,X45) = s(bool,'const.bool.DATATYPE_1'(s(X43,X44))) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f3302,f148])).
fof(f3604,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X39) = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38))) | s(bool,'F_0') = s(bool,X39) | s(X37,X40) = s(X37,sK5(X40,X38,X37))) )),
  inference(subsumption_resolution,[],[f3577,f127])).
fof(f3577,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | s(bool,X39) = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38))) | s(bool,'F_0') = s(bool,X39) | s(X37,X40) = s(X37,sK5(X40,X38,X37))) )),
  inference(superposition,[],[f3302,f419])).
fof(f3603,plain,(
  ( ! [X35,X36,X34] : (s(bool,X36) = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3576,f127])).
fof(f3576,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X36) = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f3302,f267])).
fof(f3598,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.DATATYPE_1'(s(bool,X22))) | s(bool,X20) = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f3571,f127])).
fof(f3571,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.DATATYPE_1'(s(bool,X22))) | s(bool,X20) = s(bool,X22)) )),
  inference(superposition,[],[f3302,f553])).
fof(f3592,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3565,f127])).
fof(f3565,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f3302,f211])).
fof(f3510,plain,(
  ( ! [X540,X542,X536,X538,X541,X543,X537,X539,X535,X544] : (s('type.wordSem.result'(X535),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X535),X542),s('type.wordSem.word_loc'(X535),X540))) != s('type.wordSem.result'(X535),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X535),X543),s('type.wordSem.word_loc'(X535),X544))) | s('type.wordSem.word_loc'(X535),X544) = s('type.wordSem.word_loc'(X535),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X536),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X536),fun('type.wordSem.word_loc'(X536),'type.wordSem.word_loc'(X535))),X537),s(fun('type.wordSem.word_loc'(X536),fun('type.wordSem.word_loc'(X536),'type.wordSem.word_loc'(X535))),X538),s('type.wordSem.word_loc'(X535),X539),s('type.wordSem.word_loc'(X535),X540),s('type.wordSem.word_loc'(X535),X541)))) )),
  inference(superposition,[],[f163,f176])).
fof(f3505,plain,(
  ( ! [X486,X492,X494,X488,X490,X485,X487,X493,X489,X491] : (s('type.wordSem.result'(X485),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X485),X490),s('type.wordSem.word_loc'(X485),X492))) != s('type.wordSem.result'(X485),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X485),X493),s('type.wordSem.word_loc'(X485),X494))) | s('type.wordSem.word_loc'(X485),X493) = s('type.wordSem.word_loc'(X485),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X486),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X486),fun('type.wordSem.word_loc'(X486),'type.wordSem.word_loc'(X485))),X487),s(fun('type.wordSem.word_loc'(X486),fun('type.wordSem.word_loc'(X486),'type.wordSem.word_loc'(X485))),X488),s('type.wordSem.word_loc'(X485),X489),s('type.wordSem.word_loc'(X485),X490),s('type.wordSem.word_loc'(X485),X491)))) )),
  inference(superposition,[],[f162,f176])).
fof(f3488,plain,(
  ( ! [X344,X346,X341,X343,X345,X340,X342] : (s(bool,X344) != s(bool,X346) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X340),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X340),fun('type.wordSem.word_loc'(X340),bool)),X341),s(fun('type.wordSem.word_loc'(X340),fun('type.wordSem.word_loc'(X340),bool)),X342),s(bool,X343),s(bool,X344),s(bool,X345))) | s(bool,'F_0') = s(bool,X346)) )),
  inference(superposition,[],[f2056,f176])).
fof(f3487,plain,(
  ( ! [X333,X335,X337,X339,X334,X336,X338] : (s(bool,X337) != s(bool,X339) | s(bool,'T_0') = s(bool,X339) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X333),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X333),fun('type.wordSem.word_loc'(X333),bool)),X334),s(fun('type.wordSem.word_loc'(X333),fun('type.wordSem.word_loc'(X333),bool)),X335),s(bool,X336),s(bool,X337),s(bool,X338)))) )),
  inference(superposition,[],[f2056,f176])).
fof(f3558,plain,(
  ( ! [X327,X329,X331,X332,X328,X330] : (s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X327),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X328),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X329),s(bool,X330),s(bool,X331),s(bool,X332))) | s(bool,'T_0') != s(bool,X331)) )),
  inference(forward_demodulation,[],[f3486,f148])).
fof(f3486,plain,(
  ( ! [X327,X329,X331,X332,X328,X330] : (s(bool,'T_0') != s(bool,X331) | s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X327),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X328),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X329),s(bool,X330),s(bool,X331),s(bool,X332))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X327),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X328),s(fun('type.wordSem.word_loc'(X327),fun('type.wordSem.word_loc'(X327),bool)),X329),s(bool,X330),s(bool,X331),s(bool,X332)))))) )),
  inference(superposition,[],[f1837,f176])).
fof(f3483,plain,(
  ( ! [X308,X310,X306,X309,X311,X305,X307] : (s(bool,'F_0') != s(bool,X309) | s(bool,X311) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X305),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X305),fun('type.wordSem.word_loc'(X305),bool)),X306),s(fun('type.wordSem.word_loc'(X305),fun('type.wordSem.word_loc'(X305),bool)),X307),s(bool,X308),s(bool,X309),s(bool,X310))) | s(bool,'T_0') = s(bool,X311)) )),
  inference(superposition,[],[f1774,f176])).
fof(f3555,plain,(
  ( ! [X300,X302,X304,X301,X303,X299] : (s(bool,'T_0') = s(bool,X303) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X299),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X299),fun('type.wordSem.word_loc'(X299),bool)),X300),s(fun('type.wordSem.word_loc'(X299),fun('type.wordSem.word_loc'(X299),bool)),X301),s(bool,X302),s(bool,X303),s(bool,X304)))) )),
  inference(forward_demodulation,[],[f3482,f148])).
fof(f3482,plain,(
  ( ! [X300,X302,X304,X301,X303,X299] : (s(bool,X303) = s(bool,'const.bool.DATATYPE_1'(s(bool,X303))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X299),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X299),fun('type.wordSem.word_loc'(X299),bool)),X300),s(fun('type.wordSem.word_loc'(X299),fun('type.wordSem.word_loc'(X299),bool)),X301),s(bool,X302),s(bool,X303),s(bool,X304)))) )),
  inference(superposition,[],[f1650,f176])).
fof(f3476,plain,(
  ( ! [X261,X263,X257,X259,X260,X262,X256,X258] : (s(bool,X260) != s(bool,'const.bool.DATATYPE_1'(s(X262,X263))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X256),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X257),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X258),s(bool,X259),s(bool,X260),s(bool,X261)))) )),
  inference(superposition,[],[f381,f176])).
fof(f3475,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X253),s(fun('type.num.num',bool),X255),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X249),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X249),fun('type.wordSem.word_loc'(X249),bool)),X250),s(fun('type.wordSem.word_loc'(X249),fun('type.wordSem.word_loc'(X249),bool)),X251),s(bool,X252),s(bool,X253),s(bool,X254)))) )),
  inference(superposition,[],[f267,f176])).
fof(f3444,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,'T_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.bool.DATATYPE_1'(s(bool,X61))) | s(bool,X61) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X55),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X56),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X57),s(bool,X58),s(bool,X59),s(bool,X60)))) )),
  inference(superposition,[],[f176,f553])).
fof(f3442,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X45) = s(bool,X47) | s(bool,X47) = s(bool,'const.bool.DATATYPE_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X41),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X42),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X43),s(bool,X44),s(bool,X45),s(bool,X46)))) )),
  inference(superposition,[],[f176,f553])).
fof(f3439,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X20),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X21),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X22),s(bool,X23),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f176,f211])).
fof(f3438,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,X19) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X13),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X14),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X15),s(bool,X16),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f176,f211])).
fof(f3437,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X7),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X8),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X9),s(bool,X10),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f176,f130])).
fof(f3303,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) != s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3246,f127])).
fof(f3246,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X31) != s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f214,f267])).
fof(f3302,plain,(
  ( ! [X18,X16] : (~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f3301,f152])).
fof(f3301,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3243,f127])).
fof(f3243,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f168,f267])).
fof(f3221,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X83),s(fun('type.num.num',bool),X85),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK5(X83,X84,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X84)))) )),
  inference(superposition,[],[f267,f419])).
fof(f3220,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X80),s(fun('type.num.num',bool),X82),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) )),
  inference(superposition,[],[f267,f172])).
fof(f3219,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X77),s(fun('type.num.num',bool),X78),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X75,bool),'const.pair.,_2'(s(X75,X76),s(bool,X77)))))) )),
  inference(superposition,[],[f267,f171])).
fof(f3217,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X68),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X68),s(fun('type.num.num',bool),X69),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f267,f152])).
fof(f3216,plain,(
  ( ! [X61,X66,X64,X62,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X66),s(fun('type.num.num',bool),X67),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X61),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X61),fun('type.wordSem.word_loc'(X61),bool)),X62),s(fun('type.wordSem.word_loc'(X61),fun('type.wordSem.word_loc'(X61),bool)),X63),s(bool,X64),s(bool,X65),s(bool,X66)))) )),
  inference(superposition,[],[f267,f175])).
fof(f3215,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X57),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X54),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X54),fun('type.wordSem.word_loc'(X54),bool)),X55),s(fun('type.wordSem.word_loc'(X54),fun('type.wordSem.word_loc'(X54),bool)),X56),s(bool,X57),s(bool,X58),s(bool,X59)))) )),
  inference(superposition,[],[f267,f174])).
fof(f3064,plain,(
  ( ! [X455,X461,X463,X457,X459,X454,X460,X462,X456,X458] : (s('type.wordSem.result'(X454),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X454),X461),s('type.wordSem.word_loc'(X454),X460))) != s('type.wordSem.result'(X454),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X454),X462),s('type.wordSem.word_loc'(X454),X463))) | s('type.wordSem.word_loc'(X454),X463) = s('type.wordSem.word_loc'(X454),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X455),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X455),fun('type.wordSem.word_loc'(X455),'type.wordSem.word_loc'(X454))),X456),s(fun('type.wordSem.word_loc'(X455),fun('type.wordSem.word_loc'(X455),'type.wordSem.word_loc'(X454))),X457),s('type.wordSem.word_loc'(X454),X458),s('type.wordSem.word_loc'(X454),X459),s('type.wordSem.word_loc'(X454),X460)))) )),
  inference(superposition,[],[f163,f175])).
fof(f3059,plain,(
  ( ! [X412,X408,X410,X405,X407,X413,X409,X411,X404,X406] : (s('type.wordSem.result'(X404),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X404),X410),s('type.wordSem.word_loc'(X404),X411))) != s('type.wordSem.result'(X404),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X404),X412),s('type.wordSem.word_loc'(X404),X413))) | s('type.wordSem.word_loc'(X404),X412) = s('type.wordSem.word_loc'(X404),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X405),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X405),fun('type.wordSem.word_loc'(X405),'type.wordSem.word_loc'(X404))),X406),s(fun('type.wordSem.word_loc'(X405),fun('type.wordSem.word_loc'(X405),'type.wordSem.word_loc'(X404))),X407),s('type.wordSem.word_loc'(X404),X408),s('type.wordSem.word_loc'(X404),X409),s('type.wordSem.word_loc'(X404),X410)))) )),
  inference(superposition,[],[f162,f175])).
fof(f3045,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (s(bool,X295) != s(bool,X296) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X290),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X290),fun('type.wordSem.word_loc'(X290),bool)),X291),s(fun('type.wordSem.word_loc'(X290),fun('type.wordSem.word_loc'(X290),bool)),X292),s(bool,X293),s(bool,X294),s(bool,X295))) | s(bool,'F_0') = s(bool,X296)) )),
  inference(superposition,[],[f2056,f175])).
fof(f3044,plain,(
  ( ! [X288,X285,X287,X283,X289,X284,X286] : (s(bool,X288) != s(bool,X289) | s(bool,'T_0') = s(bool,X289) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X283),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X283),fun('type.wordSem.word_loc'(X283),bool)),X284),s(fun('type.wordSem.word_loc'(X283),fun('type.wordSem.word_loc'(X283),bool)),X285),s(bool,X286),s(bool,X287),s(bool,X288)))) )),
  inference(superposition,[],[f2056,f175])).
fof(f3109,plain,(
  ( ! [X280,X282,X277,X279,X281,X278] : (s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X277),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X278),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X279),s(bool,X280),s(bool,X281),s(bool,X282))) | s(bool,'T_0') != s(bool,X282)) )),
  inference(forward_demodulation,[],[f3043,f148])).
fof(f3043,plain,(
  ( ! [X280,X282,X277,X279,X281,X278] : (s(bool,'T_0') != s(bool,X282) | s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X277),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X278),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X279),s(bool,X280),s(bool,X281),s(bool,X282))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X277),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X278),s(fun('type.wordSem.word_loc'(X277),fun('type.wordSem.word_loc'(X277),bool)),X279),s(bool,X280),s(bool,X281),s(bool,X282)))))) )),
  inference(superposition,[],[f1837,f175])).
fof(f3041,plain,(
  ( ! [X263,X265,X267,X262,X268,X264,X266] : (s(bool,'F_0') != s(bool,X267) | s(bool,X268) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X262),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X262),fun('type.wordSem.word_loc'(X262),bool)),X263),s(fun('type.wordSem.word_loc'(X262),fun('type.wordSem.word_loc'(X262),bool)),X264),s(bool,X265),s(bool,X266),s(bool,X267))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f1774,f175])).
fof(f3107,plain,(
  ( ! [X261,X257,X259,X260,X256,X258] : (s(bool,'T_0') = s(bool,X261) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X256),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X257),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X258),s(bool,X259),s(bool,X260),s(bool,X261)))) )),
  inference(forward_demodulation,[],[f3040,f148])).
fof(f3040,plain,(
  ( ! [X261,X257,X259,X260,X256,X258] : (s(bool,X261) = s(bool,'const.bool.DATATYPE_1'(s(bool,X261))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X256),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X257),s(fun('type.wordSem.word_loc'(X256),fun('type.wordSem.word_loc'(X256),bool)),X258),s(bool,X259),s(bool,X260),s(bool,X261)))) )),
  inference(superposition,[],[f1650,f175])).
fof(f3034,plain,(
  ( ! [X218,X216,X220,X215,X213,X219,X217,X214] : (s(bool,X218) != s(bool,'const.bool.DATATYPE_1'(s(X219,X220))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X213),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X213),fun('type.wordSem.word_loc'(X213),bool)),X214),s(fun('type.wordSem.word_loc'(X213),fun('type.wordSem.word_loc'(X213),bool)),X215),s(bool,X216),s(bool,X217),s(bool,X218)))) )),
  inference(superposition,[],[f381,f175])).
fof(f3006,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,'const.bool.DATATYPE_1'(s(bool,X61))) | s(bool,X61) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X55),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X56),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X57),s(bool,X58),s(bool,X59),s(bool,X60)))) )),
  inference(superposition,[],[f175,f553])).
fof(f3004,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.bool.DATATYPE_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X41),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X42),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X43),s(bool,X44),s(bool,X45),s(bool,X46)))) )),
  inference(superposition,[],[f175,f553])).
fof(f3001,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X20),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X21),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X22),s(bool,X23),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f175,f211])).
fof(f3000,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X13),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X14),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X15),s(bool,X16),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f175,f211])).
fof(f2999,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X7),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X8),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X9),s(bool,X10),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f175,f130])).
fof(f2784,plain,(
  ( ! [X70,X72,X71] : (s(bool,X70) != s(bool,X72) | s(bool,'T_0') = s(bool,sK5(X70,X71,bool)) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f2056,f419])).
fof(f2783,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f2056,f172])).
fof(f2782,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X62,bool),'const.pair.,_2'(s(X62,X63),s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f2056,f171])).
fof(f2780,plain,(
  ( ! [X57,X56,X55] : (s(bool,X55) != s(bool,X57) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X55),s(fun('type.num.num',bool),X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X57)) )),
  inference(superposition,[],[f2056,f152])).
fof(f2779,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,X51) != s(bool,X54) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X48),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X48),fun('type.wordSem.word_loc'(X48),bool)),X49),s(fun('type.wordSem.word_loc'(X48),fun('type.wordSem.word_loc'(X48),bool)),X50),s(bool,X51),s(bool,X52),s(bool,X53))) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f2056,f174])).
fof(f2760,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) != s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,sK5(X73,X74,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f2056,f419])).
fof(f2759,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X70) != s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),'const.pair.,_2'(s(bool,X70),s(X69,X71)))))) )),
  inference(superposition,[],[f2056,f172])).
fof(f2758,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X67) != s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X65,bool),'const.pair.,_2'(s(X65,X66),s(bool,X67)))))) )),
  inference(superposition,[],[f2056,f171])).
fof(f2756,plain,(
  ( ! [X59,X60,X58] : (s(bool,X58) != s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X58),s(fun('type.num.num',bool),X59),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2056,f152])).
fof(f2755,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(bool,X54) != s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X51),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X51),fun('type.wordSem.word_loc'(X51),bool)),X52),s(fun('type.wordSem.word_loc'(X51),fun('type.wordSem.word_loc'(X51),bool)),X53),s(bool,X54),s(bool,X55),s(bool,X56)))) )),
  inference(superposition,[],[f2056,f174])).
fof(f2680,plain,(
  ( ! [X366,X372,X374,X368,X370,X367,X373,X375,X369,X371] : (s('type.wordSem.result'(X366),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X366),X373),s('type.wordSem.word_loc'(X366),X370))) != s('type.wordSem.result'(X366),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X366),X374),s('type.wordSem.word_loc'(X366),X375))) | s('type.wordSem.word_loc'(X366),X375) = s('type.wordSem.word_loc'(X366),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X367),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X367),fun('type.wordSem.word_loc'(X367),'type.wordSem.word_loc'(X366))),X368),s(fun('type.wordSem.word_loc'(X367),fun('type.wordSem.word_loc'(X367),'type.wordSem.word_loc'(X366))),X369),s('type.wordSem.word_loc'(X366),X370),s('type.wordSem.word_loc'(X366),X371),s('type.wordSem.word_loc'(X366),X372)))) )),
  inference(superposition,[],[f163,f174])).
fof(f2675,plain,(
  ( ! [X325,X321,X323,X316,X318,X324,X320,X322,X317,X319] : (s('type.wordSem.result'(X316),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X316),X320),s('type.wordSem.word_loc'(X316),X323))) != s('type.wordSem.result'(X316),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X316),X324),s('type.wordSem.word_loc'(X316),X325))) | s('type.wordSem.word_loc'(X316),X324) = s('type.wordSem.word_loc'(X316),'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X317),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X317),fun('type.wordSem.word_loc'(X317),'type.wordSem.word_loc'(X316))),X318),s(fun('type.wordSem.word_loc'(X317),fun('type.wordSem.word_loc'(X317),'type.wordSem.word_loc'(X316))),X319),s('type.wordSem.word_loc'(X316),X320),s('type.wordSem.word_loc'(X316),X321),s('type.wordSem.word_loc'(X316),X322)))) )),
  inference(superposition,[],[f162,f174])).
fof(f2724,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X241),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X242),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X243),s(bool,X244),s(bool,X245),s(bool,X246))) | s(bool,'T_0') != s(bool,X244)) )),
  inference(forward_demodulation,[],[f2665,f148])).
fof(f2665,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (s(bool,'T_0') != s(bool,X244) | s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X241),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X242),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X243),s(bool,X244),s(bool,X245),s(bool,X246))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X241),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X242),s(fun('type.wordSem.word_loc'(X241),fun('type.wordSem.word_loc'(X241),bool)),X243),s(bool,X244),s(bool,X245),s(bool,X246)))))) )),
  inference(superposition,[],[f1837,f174])).
fof(f2663,plain,(
  ( ! [X231,X229,X227,X230,X228,X226,X232] : (s(bool,'F_0') != s(bool,X229) | s(bool,X232) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X226),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X226),fun('type.wordSem.word_loc'(X226),bool)),X227),s(fun('type.wordSem.word_loc'(X226),fun('type.wordSem.word_loc'(X226),bool)),X228),s(bool,X229),s(bool,X230),s(bool,X231))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f1774,f174])).
fof(f2722,plain,(
  ( ! [X222,X220,X225,X223,X221,X224] : (s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X220),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X220),fun('type.wordSem.word_loc'(X220),bool)),X221),s(fun('type.wordSem.word_loc'(X220),fun('type.wordSem.word_loc'(X220),bool)),X222),s(bool,X223),s(bool,X224),s(bool,X225)))) )),
  inference(forward_demodulation,[],[f2662,f148])).
fof(f2662,plain,(
  ( ! [X222,X220,X225,X223,X221,X224] : (s(bool,X223) = s(bool,'const.bool.DATATYPE_1'(s(bool,X223))) | s(bool,'F_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X220),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X220),fun('type.wordSem.word_loc'(X220),bool)),X221),s(fun('type.wordSem.word_loc'(X220),fun('type.wordSem.word_loc'(X220),bool)),X222),s(bool,X223),s(bool,X224),s(bool,X225)))) )),
  inference(superposition,[],[f1650,f174])).
fof(f2656,plain,(
  ( ! [X177,X182,X180,X178,X184,X183,X181,X179] : (s(bool,X180) != s(bool,'const.bool.DATATYPE_1'(s(X183,X184))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X177),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X177),fun('type.wordSem.word_loc'(X177),bool)),X178),s(fun('type.wordSem.word_loc'(X177),fun('type.wordSem.word_loc'(X177),bool)),X179),s(bool,X180),s(bool,X181),s(bool,X182)))) )),
  inference(superposition,[],[f381,f174])).
fof(f2631,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,'T_0') = s(bool,X58) | s(bool,X61) = s(bool,'const.bool.DATATYPE_1'(s(bool,X61))) | s(bool,X61) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X55),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X56),s(fun('type.wordSem.word_loc'(X55),fun('type.wordSem.word_loc'(X55),bool)),X57),s(bool,X58),s(bool,X59),s(bool,X60)))) )),
  inference(superposition,[],[f174,f553])).
fof(f2629,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X44) = s(bool,X47) | s(bool,X47) = s(bool,'const.bool.DATATYPE_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X41),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X42),s(fun('type.wordSem.word_loc'(X41),fun('type.wordSem.word_loc'(X41),bool)),X43),s(bool,X44),s(bool,X45),s(bool,X46)))) )),
  inference(superposition,[],[f174,f553])).
fof(f2626,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X20),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X21),s(fun('type.wordSem.word_loc'(X20),fun('type.wordSem.word_loc'(X20),bool)),X22),s(bool,X23),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f174,f211])).
fof(f2625,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X19) = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X13),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X14),s(fun('type.wordSem.word_loc'(X13),fun('type.wordSem.word_loc'(X13),bool)),X15),s(bool,X16),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f174,f211])).
fof(f2624,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X7),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X8),s(fun('type.wordSem.word_loc'(X7),fun('type.wordSem.word_loc'(X7),bool)),X9),s(bool,X10),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f174,f130])).
fof(f2504,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,X60) | s(bool,X62) = s(bool,sK5(X60,X61,bool)) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1774,f419])).
fof(f2503,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | s(bool,X59) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X56),'const.pair.,_2'(s(bool,X57),s(X56,X58))))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f1774,f172])).
fof(f2502,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'F_0') != s(bool,X54) | s(bool,X55) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X53),s(bool,X54))))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f1774,f171])).
fof(f2500,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') != s(bool,X45) | s(bool,X47) = s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),X46),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1774,f152])).
fof(f2440,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.wordSem.result'(X36),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X36),X39),s('type.wordSem.word_loc'(X36),X37))) != s('type.wordSem.result'(X36),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X36),X40),s('type.wordSem.word_loc'(X36),X41))) | s('type.wordSem.word_loc'(X36),X41) = s('type.wordSem.word_loc'(X36),sK5(X37,X38,'type.wordSem.word_loc'(X36))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.wordSem.word_loc'(X36),bool),X38)))) )),
  inference(superposition,[],[f163,f419])).
fof(f2439,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (s('type.wordSem.result'(X29),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X29),X33),s('type.wordSem.word_loc'(X29),X31))) != s('type.wordSem.result'(X29),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X29),X34),s('type.wordSem.word_loc'(X29),X35))) | s('type.wordSem.word_loc'(X29),X35) = s('type.wordSem.word_loc'(X29),'const.pair.FST_1'(s('type.pair.prod'('type.wordSem.word_loc'(X29),X30),'const.pair.,_2'(s('type.wordSem.word_loc'(X29),X31),s(X30,X32)))))) )),
  inference(superposition,[],[f163,f172])).
fof(f2438,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s('type.wordSem.result'(X22),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X22),X26),s('type.wordSem.word_loc'(X22),X25))) != s('type.wordSem.result'(X22),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X22),X27),s('type.wordSem.word_loc'(X22),X28))) | s('type.wordSem.word_loc'(X22),X28) = s('type.wordSem.word_loc'(X22),'const.pair.SND_1'(s('type.pair.prod'(X23,'type.wordSem.word_loc'(X22)),'const.pair.,_2'(s(X23,X24),s('type.wordSem.word_loc'(X22),X25)))))) )),
  inference(superposition,[],[f163,f171])).
fof(f2437,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s('type.wordSem.result'(X15),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X15),X20),s('type.wordSem.word_loc'(X15),X21))) != s('type.wordSem.result'(X15),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X15),X19),s('type.wordSem.word_loc'(X15),happ(s(fun('type.num.num','type.wordSem.word_loc'(X15)),X17),s('type.num.num',X18))))) | s('type.wordSem.word_loc'(X15),X21) = s('type.wordSem.word_loc'(X15),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(X15),X16),s(fun('type.num.num','type.wordSem.word_loc'(X15)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))))) )),
  inference(superposition,[],[f163,f157])).
fof(f2436,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (s('type.wordSem.result'(X9),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X9),X12),s('type.wordSem.word_loc'(X9),X10))) != s('type.wordSem.result'(X9),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X9),X13),s('type.wordSem.word_loc'(X9),X14))) | s('type.wordSem.word_loc'(X9),X14) = s('type.wordSem.word_loc'(X9),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(X9),X10),s(fun('type.num.num','type.wordSem.word_loc'(X9)),X11),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f163,f152])).
fof(f2435,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s('type.wordSem.result'(X3),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X3),X7),s('type.wordSem.word_loc'(X3),X8))) != s('type.wordSem.result'(X3),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X3),X6),s('type.wordSem.word_loc'(X3),'const.pair.FST_1'(s('type.pair.prod'('type.wordSem.word_loc'(X3),X4),X5))))) | s('type.wordSem.word_loc'(X3),X8) = s('type.wordSem.word_loc'(X3),happ(s(fun('type.pair.prod'('type.wordSem.word_loc'(X3),X4),'type.wordSem.word_loc'(X3)),'const.pair.FST_0'),s('type.pair.prod'('type.wordSem.word_loc'(X3),X4),X5)))) )),
  inference(superposition,[],[f163,f169])).
fof(f2333,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.wordSem.result'(X36),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X36),X37),s('type.wordSem.word_loc'(X36),X39))) != s('type.wordSem.result'(X36),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X36),X40),s('type.wordSem.word_loc'(X36),X41))) | s('type.wordSem.word_loc'(X36),X40) = s('type.wordSem.word_loc'(X36),sK5(X37,X38,'type.wordSem.word_loc'(X36))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.wordSem.word_loc'(X36),bool),X38)))) )),
  inference(superposition,[],[f162,f419])).
fof(f2332,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (s('type.wordSem.result'(X29),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X29),X31),s('type.wordSem.word_loc'(X29),X33))) != s('type.wordSem.result'(X29),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X29),X34),s('type.wordSem.word_loc'(X29),X35))) | s('type.wordSem.word_loc'(X29),X34) = s('type.wordSem.word_loc'(X29),'const.pair.FST_1'(s('type.pair.prod'('type.wordSem.word_loc'(X29),X30),'const.pair.,_2'(s('type.wordSem.word_loc'(X29),X31),s(X30,X32)))))) )),
  inference(superposition,[],[f162,f172])).
fof(f2331,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s('type.wordSem.result'(X22),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X22),X25),s('type.wordSem.word_loc'(X22),X26))) != s('type.wordSem.result'(X22),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X22),X27),s('type.wordSem.word_loc'(X22),X28))) | s('type.wordSem.word_loc'(X22),X27) = s('type.wordSem.word_loc'(X22),'const.pair.SND_1'(s('type.pair.prod'(X23,'type.wordSem.word_loc'(X22)),'const.pair.,_2'(s(X23,X24),s('type.wordSem.word_loc'(X22),X25)))))) )),
  inference(superposition,[],[f162,f171])).
fof(f2330,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s('type.wordSem.result'(X15),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X15),X20),s('type.wordSem.word_loc'(X15),X21))) != s('type.wordSem.result'(X15),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X15),happ(s(fun('type.num.num','type.wordSem.word_loc'(X15)),X17),s('type.num.num',X18))),s('type.wordSem.word_loc'(X15),X19))) | s('type.wordSem.word_loc'(X15),X20) = s('type.wordSem.word_loc'(X15),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(X15),X16),s(fun('type.num.num','type.wordSem.word_loc'(X15)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))))) )),
  inference(superposition,[],[f162,f157])).
fof(f2329,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (s('type.wordSem.result'(X9),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X9),X10),s('type.wordSem.word_loc'(X9),X12))) != s('type.wordSem.result'(X9),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X9),X13),s('type.wordSem.word_loc'(X9),X14))) | s('type.wordSem.word_loc'(X9),X13) = s('type.wordSem.word_loc'(X9),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(X9),X10),s(fun('type.num.num','type.wordSem.word_loc'(X9)),X11),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f162,f152])).
fof(f2328,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s('type.wordSem.result'(X3),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X3),X7),s('type.wordSem.word_loc'(X3),X8))) != s('type.wordSem.result'(X3),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X3),'const.pair.FST_1'(s('type.pair.prod'('type.wordSem.word_loc'(X3),X4),X5))),s('type.wordSem.word_loc'(X3),X6))) | s('type.wordSem.word_loc'(X3),X7) = s('type.wordSem.word_loc'(X3),happ(s(fun('type.pair.prod'('type.wordSem.word_loc'(X3),X4),'type.wordSem.word_loc'(X3)),'const.pair.FST_0'),s('type.pair.prod'('type.wordSem.word_loc'(X3),X4),X5)))) )),
  inference(superposition,[],[f162,f169])).
fof(f2216,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') = s(bool,sK5(X44,X45,bool)) | s(bool,'T_0') != s(bool,X44) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(forward_demodulation,[],[f2174,f148])).
fof(f2174,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | s(bool,sK5(X44,X45,bool)) = s(bool,'const.bool.DATATYPE_1'(s(bool,sK5(X44,X45,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f1837,f419])).
fof(f2215,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43)))))) )),
  inference(forward_demodulation,[],[f2173,f148])).
fof(f2173,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43))))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43)))))))) )),
  inference(superposition,[],[f1837,f172])).
fof(f2214,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))) )),
  inference(forward_demodulation,[],[f2172,f148])).
fof(f2172,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40))))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))))) )),
  inference(superposition,[],[f1837,f171])).
fof(f2212,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2170,f148])).
fof(f2170,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') != s(bool,X33) | s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.bool.DATATYPE_1'(s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1837,f152])).
fof(f2143,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,'const.bool.DATATYPE_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1837,f211])).
fof(f2142,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1))) | s(bool,X2) = s(bool,'const.bool.DATATYPE_1'(s(bool,X2))) | s(bool,X0) = s(bool,X2)) )),
  inference(superposition,[],[f1837,f553])).
fof(f2056,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f381,f1650])).
fof(f2102,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,sK5(X44,X45,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(forward_demodulation,[],[f2046,f148])).
fof(f2046,plain,(
  ( ! [X45,X44] : (s(bool,X44) = s(bool,'const.bool.DATATYPE_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,sK5(X44,X45,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f1650,f419])).
fof(f2101,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43))))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(forward_demodulation,[],[f2045,f148])).
fof(f2045,plain,(
  ( ! [X43,X41,X42] : (s(bool,X42) = s(bool,'const.bool.DATATYPE_1'(s(bool,X42))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43)))))) )),
  inference(superposition,[],[f1650,f172])).
fof(f2100,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(forward_demodulation,[],[f2044,f148])).
fof(f2044,plain,(
  ( ! [X39,X38,X40] : (s(bool,X40) = s(bool,'const.bool.DATATYPE_1'(s(bool,X40))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))) )),
  inference(superposition,[],[f1650,f171])).
fof(f2098,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(forward_demodulation,[],[f2042,f148])).
fof(f2042,plain,(
  ( ! [X33,X34] : (s(bool,X33) = s(bool,'const.bool.DATATYPE_1'(s(bool,X33))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1650,f152])).
fof(f129,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.wordSem.result_|49||49||split|1',aHL_FALSITY)).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X32] : (s(bool,'F_0') = s(bool,X32) <=> ~p(s(bool,X32)))),
  file('thm.wordSem.result_|49||49||split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f267,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f152,f130])).
fof(f302,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f171,f130])).
fof(f341,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f172,f130])).
fof(f833,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16))) | s(bool,'T_0') = s(bool,sK5(X15,X16,bool))) )),
  inference(superposition,[],[f419,f130])).
fof(f1327,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK7(X16,X17,X18,bool))) )),
  inference(superposition,[],[f173,f130])).
fof(f1650,plain,(
  ( ! [X51] : (s(bool,X51) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(trivial_inequality_removal,[],[f1594])).
fof(f1594,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X51) = s(bool,'const.bool.DATATYPE_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f195,f553])).
fof(f1609,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'T_0') != s(bool,f(s(X107,sK7(X108,X109,bool,X107)))) | s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110))) | s(bool,X110) = s(bool,g(s(X107,sK7(X108,X109,bool,X107))))) )),
  inference(superposition,[],[f173,f553])).
fof(f1545,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X120) != s(bool,f(s(X117,sK7(X118,X119,bool,X117)))) | s(fun(X117,bool),X118) = s(fun(X117,bool),X119) | s(bool,X120) = s(bool,'const.bool.DATATYPE_1'(s(bool,X120))) | s(bool,'T_0') = s(bool,g(s(X117,sK7(X118,X119,bool,X117))))) )),
  inference(superposition,[],[f173,f553])).
fof(f1428,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X37),s('type.num.num',X38))) | s(bool,X39) = s(bool,'const.bool.DATATYPE_1'(s(bool,X39))) | s(bool,X39) = s(bool,'const.ind_type.FCONS_3'(s(bool,X36),s(fun('type.num.num',bool),X37),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38)))))) )),
  inference(superposition,[],[f553,f157])).
fof(f1405,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X52) = s(bool,happ(s(fun('type.num.num',bool),X50),s('type.num.num',X51))) | s(bool,X52) = s(bool,'const.bool.DATATYPE_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X49),s(fun('type.num.num',bool),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f553,f157])).
fof(f1426,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),X31))) | s(bool,X32) = s(bool,'const.bool.DATATYPE_1'(s(bool,X32))) | s(bool,X32) = s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X30),X31)))) )),
  inference(superposition,[],[f553,f169])).
fof(f1403,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),X44))) | s(bool,X45) = s(bool,'const.bool.DATATYPE_1'(s(bool,X45))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X43),X44)))) )),
  inference(superposition,[],[f553,f169])).
fof(f1616,plain,(
  ( ! [X134,X132,X135,X133] : (s(X134,f(s(bool,'T_0'))) != s(X134,g(s(bool,'T_0'))) | s(fun(bool,X134),X132) = s(fun(bool,X134),X133) | s(bool,X135) = s(bool,'const.bool.DATATYPE_1'(s(bool,X135))) | s(bool,X135) = s(bool,sK7(X132,X133,X134,bool))) )),
  inference(superposition,[],[f173,f553])).
fof(f1865,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,'const.bool.DATATYPE_1'(s(bool,X29))) | s(bool,X27) = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f1585,f127])).
fof(f1585,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X27) != s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,'const.bool.DATATYPE_1'(s(bool,X29))) | s(bool,X27) = s(bool,X29)) )),
  inference(superposition,[],[f214,f553])).
fof(f1566,plain,(
  ( ! [X45,X44] : (s(bool,X44) = s(bool,'const.bool.DATATYPE_1'(s(bool,X44))) | s(bool,X45) = s(bool,'const.bool.DATATYPE_1'(s(bool,X45))) | s(bool,X44) = s(bool,X45)) )),
  inference(superposition,[],[f148,f553])).
fof(f1552,plain,(
  ( ! [X144,X142,X145,X143] : (s(X144,f(s(bool,X145))) != s(X144,g(s(bool,X145))) | s(fun(bool,X144),X142) = s(fun(bool,X144),X143) | s(bool,X145) = s(bool,'const.bool.DATATYPE_1'(s(bool,X145))) | s(bool,'T_0') = s(bool,sK7(X142,X143,X144,bool))) )),
  inference(superposition,[],[f173,f553])).
fof(f1837,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,X56) | s(bool,X56) = s(bool,'const.bool.DATATYPE_1'(s(bool,X56)))) )),
  inference(subsumption_resolution,[],[f1527,f195])).
fof(f1527,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,X56) | s(bool,X56) = s(bool,'const.bool.DATATYPE_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f195,f553])).
fof(f1774,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1773,f127])).
fof(f1773,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f1437,f148])).
fof(f1437,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | ~p(s(bool,'const.bool.DATATYPE_1'(s(bool,X11)))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f135,f553])).
fof(f1431,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,X50) = s(bool,'const.bool.DATATYPE_1'(s(bool,X50))) | s(bool,X50) = s(bool,sK5(X48,X49,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f553,f419])).
fof(f1430,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,'const.bool.DATATYPE_1'(s(bool,X47))) | s(bool,X47) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X45),s(X44,X46)))))) )),
  inference(superposition,[],[f553,f172])).
fof(f1429,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,'const.bool.DATATYPE_1'(s(bool,X43))) | s(bool,X43) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X40,bool),'const.pair.,_2'(s(X40,X41),s(bool,X42)))))) )),
  inference(superposition,[],[f553,f171])).
fof(f1427,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,'const.bool.DATATYPE_1'(s(bool,X35))) | s(bool,X35) = s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f553,f152])).
fof(f1410,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.DATATYPE_1'(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f553,f211])).
fof(f1409,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.DATATYPE_1'(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.DATATYPE_1'(s(bool,X3))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f553,f553])).
fof(f1408,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,X63) | s(bool,X63) = s(bool,'const.bool.DATATYPE_1'(s(bool,X63))) | s(bool,'T_0') = s(bool,sK5(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f553,f419])).
fof(f1407,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,X60) | s(bool,X60) = s(bool,'const.bool.DATATYPE_1'(s(bool,X60))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X58),s(X57,X59)))))) )),
  inference(superposition,[],[f553,f172])).
fof(f1406,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X55) = s(bool,X56) | s(bool,X56) = s(bool,'const.bool.DATATYPE_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X53,bool),'const.pair.,_2'(s(X53,X54),s(bool,X55)))))) )),
  inference(superposition,[],[f553,f171])).
fof(f1404,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X48) | s(bool,X48) = s(bool,'const.bool.DATATYPE_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X46),s(fun('type.num.num',bool),X47),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f553,f152])).
fof(f1325,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK7(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f173,f211])).
fof(f1323,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK7(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f173,f211])).
fof(f1078,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) | s(bool,X2) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) )),
  inference(global_subsumption,[],[f126,f125,f124,f127,f129,f130,f132,f131,f136,f135,f148,f152,f153,f154,f156,f155,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f267,f302,f341,f381,f392,f393,f394,f415,f416,f417,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f833,f897,f850,f851,f852,f853,f854,f856,f861,f862,f863,f865,f866,f898,f899,f903,f932,f951,f972,f995,f970,f837,f836,f835,f834,f938,f940,f942,f952,f982,f996])).
fof(f996,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) | s(bool,X2) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f990,f127])).
fof(f990,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) | s(bool,X2) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f952,f211])).
fof(f952,plain,(
  ( ! [X37,X38] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),X38)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f897,f169])).
fof(f951,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK5(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f156,f169])).
fof(f899,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))) | s(bool,'T_0') = s(bool,X45) | s(X43,X46) = s(X43,sK5(X46,X44,X43))) )),
  inference(subsumption_resolution,[],[f878,f127])).
fof(f878,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | s(bool,X45) != s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))) | s(bool,'T_0') = s(bool,X45) | s(X43,X46) = s(X43,sK5(X46,X44,X43))) )),
  inference(superposition,[],[f214,f419])).
fof(f865,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,happ(s(fun(X79,bool),X80),s(X79,sK5(X82,sK5(X80,X81,fun(X79,bool)),X79))))) | p(s(bool,'|dollar|exists_1'(s(fun(X79,bool),sK5(X80,X81,fun(X79,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X79,bool),bool),X81)))) )),
  inference(superposition,[],[f156,f419])).
fof(f863,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X73,X74))) | s(bool,'T_0') = s(bool,sK5(X71,X72,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X72)))) )),
  inference(superposition,[],[f381,f419])).
fof(f831,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X10))) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f419,f211])).
fof(f829,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,X5) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f419,f211])).
fof(f828,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f820,f127])).
fof(f820,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f156,f211])).
fof(f819,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK5(X12,X13,bool))) )),
  inference(superposition,[],[f156,f130])).
fof(f817,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f156,f211])).
fof(f815,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK5(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f156,f211])).
fof(f827,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X11)))) | ~p(s(bool,happ(s(fun(X9,bool),X11),s(X9,sK5(X13,'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9)))))) )),
  inference(forward_demodulation,[],[f814,f172])).
fof(f814,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X11),s(X9,sK5(X13,'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9))))) | p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12))))))))) )),
  inference(superposition,[],[f156,f172])).
fof(f826,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X7)))) | ~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,sK5(X8,'const.pair.SND_1'(s('type.pair.prod'(X5,fun(X4,bool)),'const.pair.,_2'(s(X5,X6),s(fun(X4,bool),X7)))),X4)))))) )),
  inference(forward_demodulation,[],[f813,f171])).
fof(f813,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,sK5(X8,'const.pair.SND_1'(s('type.pair.prod'(X5,fun(X4,bool)),'const.pair.,_2'(s(X5,X6),s(fun(X4,bool),X7)))),X4))))) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'const.pair.SND_1'(s('type.pair.prod'(X5,fun(X4,bool)),'const.pair.,_2'(s(X5,X6),s(fun(X4,bool),X7))))))))) )),
  inference(superposition,[],[f156,f171])).
fof(f825,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X3,'const.ind_type.FCONS_3'(s(fun(X0,bool),X1),s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0')),X0)))))) )),
  inference(forward_demodulation,[],[f812,f152])).
fof(f812,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X3,'const.ind_type.FCONS_3'(s(fun(X0,bool),X1),s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0')),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'const.ind_type.FCONS_3'(s(fun(X0,bool),X1),s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f156,f152])).
fof(f805,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f568,f127])).
fof(f568,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f214,f211])).
fof(f804,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f567,f127])).
fof(f567,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f168,f211])).
fof(f553,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,'const.bool.DATATYPE_1'(s(bool,X28))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f148,f211])).
fof(f472,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X24),s(X23,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f211,f172])).
fof(f471,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X19,bool),'const.pair.,_2'(s(X19,X20),s(bool,X21))))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f211,f171])).
fof(f470,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X18) = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f211,f152])).
fof(f460,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f211,f211])).
fof(f447,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X31),s(X30,X32)))))) )),
  inference(superposition,[],[f211,f172])).
fof(f446,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(superposition,[],[f211,f171])).
fof(f445,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f211,f152])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f418,f129])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f155,f130])).
fof(f394,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,'const.bool.DATATYPE_1'(s(X22,X23))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X20),s(X19,X21)))))) )),
  inference(superposition,[],[f381,f172])).
fof(f393,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X16) != s(bool,'const.bool.DATATYPE_1'(s(X17,X18))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))))) )),
  inference(superposition,[],[f381,f171])).
fof(f392,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X10) != s(bool,'const.bool.DATATYPE_1'(s(X12,X13))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X10),s(fun('type.num.num',bool),X11),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f381,f152])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f372,f127])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f214,f148])).
fof(f247,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f243,f129])).
fof(f243,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f165,f130])).
fof(f214,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f135,f130])).
fof(f211,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f130,f130])).
fof(f195,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f129,f131])).
fof(f179,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X33,X34,X35] : ! [X36,X37,X38] : s(X33,happ(s(fun(X34,X33),X36),s(X34,happ(s(fun(X35,X34),X37),s(X35,X38))))) = s(X33,happ(s(fun(X35,X33),'_dst_x0x1_2'(s(fun(X34,X33),X36),s(fun(X35,X34),X37))),s(X35,X38)))),
  file('thm.wordSem.result_|49||49||split|1','_dst_x0x1_ax')).
fof(f178,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14,X6] : ! [X7,X8,X13,X15,X16,X17,X18] : s(X14,happ(s(fun('type.wordSem.word_loc'(X6),X14),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s('type.wordSem.word_loc'(X6),X7))),s('type.wordSem.word_loc'(X6),X8))) = s(X14,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s(X14,X16),s(X14,X17),s(X14,X18)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_case_def|split|0')).
fof(f176,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X14] : ! [X13,X15,X16,X17,X18] : s(X14,X17) = s(X14,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s(X14,X16),s(X14,X17),s(X14,X18)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_case_def|split|3')).
fof(f175,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X14] : ! [X13,X15,X16,X17,X18] : s(X14,X18) = s(X14,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s(X14,X16),s(X14,X17),s(X14,X18)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_case_def|split|4')).
fof(f174,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X14] : ! [X13,X15,X16,X17,X18] : s(X14,X16) = s(X14,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X13),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X14)),X15),s(X14,X16),s(X14,X17),s(X14,X18)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_case_def|split|2')).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.wordSem.result_|49||49||split|1',aHL_EXT)).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X20,X14] : ! [X21,X22] : s(X14,X21) = s(X14,'const.pair.FST_1'(s('type.pair.prod'(X14,X20),'const.pair.,_2'(s(X14,X21),s(X20,X22)))))),
  file('thm.wordSem.result_|49||49||split|1','thm.pair.FST')).
fof(f171,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X20] : ! [X21,X22] : s(X20,X22) = s(X20,'const.pair.SND_1'(s('type.pair.prod'(X14,X20),'const.pair.,_2'(s(X14,X21),s(X20,X22)))))),
  file('thm.wordSem.result_|49||49||split|1','thm.pair.SND')).
fof(f170,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK6(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X20,X14] : ! [X23] : ? [X13] : ! [X24,X25,X26] : s(X20,happ(s(fun('type.ind_type.recspace'(X14),X20),X13),s('type.ind_type.recspace'(X14),'const.ind_type.CONSTR_3'(s('type.num.num',X24),s(X14,X25),s(fun('type.num.num','type.ind_type.recspace'(X14)),X26))))) = s(X20,happ(s(fun(fun('type.num.num',X20),X20),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X14)),fun(fun('type.num.num',X20),X20)),happ(s(fun(X14,fun(fun('type.num.num','type.ind_type.recspace'(X14)),fun(fun('type.num.num',X20),X20))),happ(s(fun('type.num.num',fun(X14,fun(fun('type.num.num','type.ind_type.recspace'(X14)),fun(fun('type.num.num',X20),X20)))),X23),s('type.num.num',X24))),s(X14,X25))),s(fun('type.num.num','type.ind_type.recspace'(X14)),X26))),s(fun('type.num.num',X20),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X14),X20),X13),s(fun('type.num.num','type.ind_type.recspace'(X14)),X26)))))),
  file('thm.wordSem.result_|49||49||split|1','thm.ind_type.CONSTR_REC')).
fof(f169,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X14,X20] : ! [X11] : s(X14,'const.pair.FST_1'(s('type.pair.prod'(X14,X20),X11))) = s(X14,happ(s(fun('type.pair.prod'(X14,X20),X14),'const.pair.FST_0'),s('type.pair.prod'(X14,X20),X11)))),
  file('thm.wordSem.result_|49||49||split|1','ARITY_0_1_const.pair.FST')).
fof(f165,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X28,X29] : ((p(s(bool,X28)) => p(s(bool,X29))) => ((p(s(bool,X29)) => p(s(bool,X28))) => s(bool,X28) = s(bool,X29)))),
  file('thm.wordSem.result_|49||49||split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f168,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f164,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X3) | s('type.wordSem.word_loc'(X0),X2) != s('type.wordSem.word_loc'(X0),X4) | s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f188,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3)))))) )),
  inference(forward_demodulation,[],[f161,f160])).
fof(f161,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X13,X7,X8] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),X7))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),X8)))))))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_size_def|split|0')).
fof(f160,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X2))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X0,'type.num.num'),X1),s('type.wordSem.word_loc'(X0),X3)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X13,X7,X8] : s('type.num.num','const.wordSem.result_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),X7))),s('type.num.num','const.wordSem.word_loc_size_2'(s(fun(X6,'type.num.num'),X13),s('type.wordSem.word_loc'(X6),X8)))))))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_size_def|split|1')).
fof(f159,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X14] : ! [X11,X12,X19] : s(X14,'const.ind_type.FCONS_3'(s(X14,X11),s(fun('type.num.num',X14),X12),s('type.num.num',X19))) = s(X14,happ(s(fun('type.num.num',X14),happ(s(fun(fun('type.num.num',X14),fun('type.num.num',X14)),happ(s(fun(X14,fun(fun('type.num.num',X14),fun('type.num.num',X14))),'const.ind_type.FCONS_0'),s(X14,X11))),s(fun('type.num.num',X14),X12))),s('type.num.num',X19)))),
  file('thm.wordSem.result_|49||49||split|1','ARITY_0_3_const.ind_type.FCONS')).
fof(f158,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14] : ! [X11,X12,X19] : s('type.ind_type.recspace'(X14),'const.ind_type.CONSTR_3'(s('type.num.num',X11),s(X14,X12),s(fun('type.num.num','type.ind_type.recspace'(X14)),X19))) = s('type.ind_type.recspace'(X14),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X14)),'type.ind_type.recspace'(X14)),happ(s(fun(X14,fun(fun('type.num.num','type.ind_type.recspace'(X14)),'type.ind_type.recspace'(X14))),'const.ind_type.CONSTR_1'(s('type.num.num',X11))),s(X14,X12))),s(fun('type.num.num','type.ind_type.recspace'(X14)),X19)))),
  file('thm.wordSem.result_|49||49||split|1','ARITY_1_3_const.ind_type.CONSTR')).
fof(f157,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X14] : ! [X27,X13,X30] : s(X14,happ(s(fun('type.num.num',X14),X13),s('type.num.num',X30))) = s(X14,'const.ind_type.FCONS_3'(s(X14,X27),s(fun('type.num.num',X14),X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))),
  file('thm.wordSem.result_|49||49||split|1','thm.ind_type.FCONS|split|1')).
fof(f156,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Exception_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Exception_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),happ(s(fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0)),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),'type.wordSem.result'(X0))),'const.wordSem.Exception_0'),s('type.wordSem.word_loc'(X0),X1))),s('type.wordSem.word_loc'(X0),X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X11,X12] : s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X11),s('type.wordSem.word_loc'(X6),X12))) = s('type.wordSem.result'(X6),happ(s(fun('type.wordSem.word_loc'(X6),'type.wordSem.result'(X6)),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),'type.wordSem.result'(X6))),'const.wordSem.Exception_0'),s('type.wordSem.word_loc'(X6),X11))),s('type.wordSem.word_loc'(X6),X12)))),
  file('thm.wordSem.result_|49||49||split|1','ARITY_0_2_const.wordSem.Exception')).
fof(f152,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : ! [X27,X13] : s(X14,X27) = s(X14,'const.ind_type.FCONS_3'(s(X14,X27),s(fun('type.num.num',X14),X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.wordSem.result_|49||49||split|1','thm.ind_type.FCONS|split|0')).
fof(f148,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X14] : ! [X21] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X14,X21)))),
  file('thm.wordSem.result_|49||49||split|1','thm.bool.DATATYPE_TAG_THM')).
fof(f131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X32] : (s(bool,'T_0') = s(bool,X32) <=> p(s(bool,X32)))),
  file('thm.wordSem.result_|49||49||split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f130,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.wordSem.result_|49||49||split|1',aHL_BOOL_CASES)).
fof(f127,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.wordSem.result_|49||49||split|1',aHL_TRUTH)).
fof(f7568,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267,f7546,f7555,f7559,f7567])).
fof(f7560,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267,f7546,f7555,f7559])).
fof(f7556,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267,f7546,f7555])).
fof(f7551,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7550,f127])).
fof(f7550,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f7549,f148])).
fof(f7549,plain,(
  ( ! [X15] : (~p(s(bool,'const.bool.DATATYPE_1'(s(bool,X15))))) ) | $spl28),
  inference(subsumption_resolution,[],[f7548,f7115])).
fof(f7548,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,'const.bool.DATATYPE_1'(s(bool,X15))))) ) | $spl28),
  inference(forward_demodulation,[],[f7206,f7115])).
fof(f7206,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,X15) | ~p(s(bool,'const.bool.DATATYPE_1'(s(bool,X15))))) ) | $spl28),
  inference(superposition,[],[f135,f6869])).
fof(f7547,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267,f7546])).
fof(f7545,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543,f7267])).
fof(f7544,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271,f7543])).
fof(f7534,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503,f7271])).
fof(f7520,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507,f7503])).
fof(f7516,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7515,f127])).
fof(f7515,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f7514,f7115])).
fof(f7514,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,happ(s(fun(X200,bool),'_dst_x0x1_2'(s(fun(X201,bool),X202),s(fun(X200,X201),X203))),s(X200,X204))))) ) | $spl28),
  inference(subsumption_resolution,[],[f7397,f7115])).
fof(f7397,plain,(
  ( ! [X204,X202,X200,X203,X201] : (s(bool,'T_0') != s(bool,happ(s(fun(X201,bool),X202),s(X201,happ(s(fun(X200,X201),X203),s(X200,X204))))) | ~p(s(bool,happ(s(fun(X200,bool),'_dst_x0x1_2'(s(fun(X201,bool),X202),s(fun(X200,X201),X203))),s(X200,X204))))) ) | $spl28),
  inference(backward_demodulation,[],[f7115,f4102])).
fof(f7513,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7512,f127])).
fof(f7512,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f7511,f7115])).
fof(f7511,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,happ(s(fun(X207,X205),X208),s(X207,X209))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f7389,f7115])).
fof(f7389,plain,(
  ( ! [X208,X206,X209,X207,X205] : (s(bool,'T_0') != s(bool,happ(s(fun(X207,bool),'_dst_x0x1_2'(s(fun(X205,bool),X206),s(fun(X207,X205),X208))),s(X207,X209))) | ~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,happ(s(fun(X207,X205),X208),s(X207,X209))))))) ) | $spl28),
  inference(backward_demodulation,[],[f7115,f4023])).
fof(f7508,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f1403,f1426,f1405,f1428,f1545,f1609,f1650,f1327,f833,f341,f302,f267,f136,f135,f129,f2098,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2500,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2675,f2680,f2755,f2756,f2758,f2759,f2760,f2779,f2780,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3059,f3064,f3215,f3216,f3217,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4166,f4042,f4046,f4047,f4057,f4068,f4073,f4178,f4121,f4125,f4126,f4136,f4147,f4152,f4220,f4221,f4222,f4223,f4224,f4225,f4227,f4228,f4229,f4267,f4302,f4303,f4304,f4307,f4308,f4309,f4441,f4533,f4534,f4535,f5060,f4665,f5127,f5166,f5167,f5168,f5169,f5170,f5171,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5819,f5820,f5821,f5822,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6175,f6180,f6181,f6929,f6318,f6869,f3590,f862,f861,f856,f854,f853,f852,f851,f850,f970,f972,f4056,f4135,f1218,f1219,f3496,f3495,f3052,f3051,f2670,f2669,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f5816,f5767,f5165,f4299,f4219,f3214,f2778,f2754,f2499,f2097,f982,f942,f940,f938,f4113,f4098,f4134,f4055,f4129,f4112,f4111,f4106,f4104,f4103,f4102,f4101,f4100,f4050,f4034,f4033,f4032,f4027,f4025,f4024,f4023,f4022,f4021,f4019,f897,f932,f903,f3887,f5823,f5774,f5172,f4306,f4226,f3218,f2781,f2757,f2501,f2099,f1234,f1194,f1192,f1190,f6306,f6163,f4653,f4429,f3889,f1332,f1330,f1328,f7115,f7507])).
fof(f7506,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7300,f127])).
fof(f7300,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(backward_demodulation,[],[f7115,f129])).
fof(f7298,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7297,f127])).
fof(f7297,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f7121,f148])).
fof(f7121,plain,(
  ~p(s(bool,'const.bool.DATATYPE_1'(s(bool,'F_0')))) | $spl28),
  inference(unit_resulting_resolution,[],[f6869,f135])).
fof(f7296,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7295,f127])).
fof(f7295,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f7124,f148])).
fof(f7124,plain,(
  ~p(s(bool,'const.bool.DATATYPE_1'(s(bool,'F_0')))) | $spl28),
  inference(unit_resulting_resolution,[],[f195,f6869,f214])).
fof(f7116,plain,(
  $false | $spl28),
  inference(unit_resulting_resolution,[],[f195,f6869,f381])).
fof(f7071,plain,(
  $false | ($spl22 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f7070,f7068])).
fof(f7068,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X72,X73)))) ) | ($spl24 | $spl34)),
  inference(global_subsumption,[],[f6865,f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f862,f861,f856,f854,f851,f850,f970,f972,f1218,f1219,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f938,f940,f982,f1403,f1426,f1190,f1192,f1234,f1405,f1428,f897,f932,f903,f1328,f1330,f1545,f1609,f1650,f1332,f1327,f1194,f942,f853,f852,f833,f341,f302,f267,f136,f135,f129,f2097,f2098,f2099,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2499,f2500,f2501,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2669,f2670,f2675,f2680,f2754,f2755,f2756,f2757,f2758,f2759,f2760,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3051,f3052,f3059,f3064,f3214,f3215,f3216,f3217,f3218,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3495,f3496,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3590,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3887,f3889,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4019,f4021,f4022,f4023,f4024,f4025,f4027,f4032,f4033,f4034,f4166,f4042,f4046,f4047,f4050,f4055,f4056,f4057,f4068,f4073,f4098,f4100,f4101,f4102,f4103,f4104,f4106,f4111,f4112,f4113,f4178,f4121,f4125,f4126,f4129,f4134,f4135,f4136,f4147,f4152,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4226,f4227,f4228,f4229,f4267,f4299,f4302,f4303,f4304,f4306,f4307,f4308,f4309,f4429,f4441,f4533,f4534,f4535,f5060,f4653,f4665,f5127,f5165,f5166,f5167,f5168,f5169,f5170,f5171,f5172,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5774,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5816,f5819,f5820,f5821,f5822,f5823,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6163,f6175,f6180,f6181,f6929,f6306,f6318,f6888,f7049,f7045,f7052,f7055,f7059,f7060,f7040,f7063,f7066])).
fof(f7066,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | $spl34),
  inference(subsumption_resolution,[],[f7007,f129])).
fof(f7007,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f155])).
fof(f7063,plain,(
  ( ! [X43,X44,X42] : (s(bool,X44) = s(bool,happ(s(fun('type.pair.prod'(bool,X42),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X42),X43)))) ) | $spl34),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f862,f861,f856,f854,f851,f850,f970,f972,f1218,f1219,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f938,f940,f982,f1403,f1426,f1190,f1192,f1234,f1405,f1428,f897,f932,f903,f1328,f1330,f1545,f1609,f1650,f1332,f1327,f1194,f942,f853,f852,f833,f341,f302,f267,f136,f135,f129,f2097,f2098,f2099,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2499,f2500,f2501,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2669,f2670,f2675,f2680,f2754,f2755,f2756,f2757,f2758,f2759,f2760,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3051,f3052,f3059,f3064,f3214,f3215,f3216,f3217,f3218,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3495,f3496,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3590,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3887,f3889,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4019,f4021,f4022,f4023,f4024,f4025,f4027,f4032,f4033,f4034,f4166,f4042,f4046,f4047,f4050,f4055,f4056,f4057,f4068,f4073,f4098,f4100,f4101,f4102,f4103,f4104,f4106,f4111,f4112,f4113,f4178,f4121,f4125,f4126,f4129,f4134,f4135,f4136,f4147,f4152,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4226,f4227,f4228,f4229,f4267,f4299,f4302,f4303,f4304,f4306,f4307,f4308,f4309,f4429,f4441,f4533,f4534,f4535,f5060,f4653,f4665,f5127,f5165,f5166,f5167,f5168,f5169,f5170,f5171,f5172,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5774,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5816,f5819,f5820,f5821,f5822,f5823,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6163,f6175,f6180,f6181,f6929,f6306,f6318,f6888,f7049,f7045,f7052,f7055,f7059,f7060,f7040])).
fof(f7040,plain,(
  ( ! [X43,X44,X42] : (s(bool,X44) = s(bool,happ(s(fun('type.pair.prod'(bool,X42),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X42),X43))) | s(bool,'T_0') = s(bool,X44)) ) | $spl34),
  inference(trivial_inequality_removal,[],[f6998])).
fof(f6998,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X44) = s(bool,happ(s(fun('type.pair.prod'(bool,X42),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X42),X43))) | s(bool,'T_0') = s(bool,X44)) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f2499])).
fof(f7060,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (s(bool,X268) = s(bool,happ(s(fun(X263,bool),X264),s(X263,happ(s(fun(X265,X263),X266),s(X265,X267)))))) ) | $spl34),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f862,f861,f856,f854,f851,f850,f970,f972,f1218,f1219,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f938,f940,f982,f1403,f1426,f1190,f1192,f1234,f1405,f1428,f897,f932,f903,f1328,f1330,f1545,f1609,f1650,f1332,f1327,f1194,f942,f853,f852,f833,f341,f302,f267,f136,f135,f129,f2097,f2098,f2099,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2499,f2500,f2501,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2669,f2670,f2675,f2680,f2754,f2755,f2756,f2757,f2758,f2759,f2760,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3051,f3052,f3059,f3064,f3214,f3215,f3216,f3217,f3218,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3495,f3496,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3590,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3887,f3889,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4019,f4021,f4022,f4023,f4024,f4025,f4027,f4032,f4033,f4034,f4166,f4042,f4046,f4047,f4050,f4055,f4056,f4057,f4068,f4073,f4098,f4100,f4101,f4102,f4103,f4104,f4106,f4111,f4112,f4113,f4178,f4121,f4125,f4126,f4129,f4134,f4135,f4136,f4147,f4152,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4226,f4227,f4228,f4229,f4267,f4299,f4302,f4303,f4304,f4306,f4307,f4308,f4309,f4429,f4441,f4533,f4534,f4535,f5060,f4653,f4665,f5127,f5165,f5166,f5167,f5168,f5169,f5170,f5171,f5172,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5774,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5816,f5819,f5820,f5821,f5822,f5823,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6163,f6175,f6180,f6181,f6929,f6306,f6318,f6888,f7049,f7045,f7052,f7055,f7059])).
fof(f7059,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (s(bool,X268) = s(bool,happ(s(fun(X263,bool),X264),s(X263,happ(s(fun(X265,X263),X266),s(X265,X267))))) | s(bool,'T_0') = s(bool,X268)) ) | $spl34),
  inference(subsumption_resolution,[],[f6977,f129])).
fof(f6977,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (p(s(bool,'F_0')) | s(bool,X268) = s(bool,happ(s(fun(X263,bool),X264),s(X263,happ(s(fun(X265,X263),X266),s(X265,X267))))) | s(bool,'T_0') = s(bool,X268)) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f4033])).
fof(f7055,plain,(
  ( ! [X198,X196,X199,X197,X195] : (s(bool,'T_0') != s(bool,happ(s(fun(X195,bool),X196),s(X195,happ(s(fun(X197,X195),X198),s(X197,X199)))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6970,f129])).
fof(f6970,plain,(
  ( ! [X198,X196,X199,X197,X195] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(X195,bool),X196),s(X195,happ(s(fun(X197,X195),X198),s(X197,X199)))))) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f4021])).
fof(f7052,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X51) = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) ) | $spl34),
  inference(global_subsumption,[],[f126,f125,f124,f127,f130,f132,f131,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f863,f865,f899,f951,f995,f952,f996,f1078,f1323,f1325,f1730,f1714,f1702,f1697,f1404,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1427,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1552,f1566,f1865,f1651,f1616,f862,f861,f856,f854,f851,f850,f970,f972,f1218,f1219,f898,f866,f837,f836,f835,f834,f417,f416,f415,f155,f938,f940,f982,f1403,f1426,f1190,f1192,f1234,f1405,f1428,f897,f932,f903,f1328,f1330,f1545,f1609,f1650,f1332,f1327,f1194,f942,f853,f852,f833,f341,f302,f267,f136,f135,f129,f2097,f2098,f2099,f2100,f2101,f2102,f2056,f2142,f2143,f2212,f2214,f2215,f2216,f2328,f2329,f2330,f2331,f2332,f2333,f2435,f2436,f2437,f2438,f2439,f2440,f2499,f2500,f2501,f2502,f2503,f2504,f2624,f2625,f2626,f2629,f2631,f2656,f2722,f2663,f2724,f2669,f2670,f2675,f2680,f2754,f2755,f2756,f2757,f2758,f2759,f2760,f2778,f2779,f2780,f2781,f2782,f2783,f2784,f2999,f3000,f3001,f3004,f3006,f3034,f3107,f3041,f3109,f3044,f3045,f3051,f3052,f3059,f3064,f3214,f3215,f3216,f3217,f3218,f3219,f3220,f3221,f3302,f3303,f3437,f3438,f3439,f3442,f3444,f3475,f3476,f3555,f3483,f3558,f3487,f3488,f3495,f3496,f3505,f3510,f3592,f3598,f3603,f3604,f3606,f3590,f3843,f3634,f3635,f3636,f3637,f3639,f3640,f3641,f3673,f3887,f3889,f3892,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3929,f3931,f4163,f4019,f4021,f4022,f4023,f4024,f4025,f4027,f4032,f4033,f4034,f4166,f4042,f4046,f4047,f4050,f4055,f4056,f4057,f4068,f4073,f4098,f4100,f4101,f4102,f4103,f4104,f4106,f4111,f4112,f4113,f4178,f4121,f4125,f4126,f4129,f4134,f4135,f4136,f4147,f4152,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4226,f4227,f4228,f4229,f4267,f4299,f4302,f4303,f4304,f4306,f4307,f4308,f4309,f4429,f4441,f4533,f4534,f4535,f5060,f4653,f4665,f5127,f5165,f5166,f5167,f5168,f5169,f5170,f5171,f5172,f5173,f5174,f5175,f5609,f5610,f5611,f5612,f5614,f5615,f5616,f6529,f6528,f6506,f6503,f6660,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5774,f5775,f5776,f5777,f5778,f5780,f6494,f6493,f5816,f5819,f5820,f5821,f5822,f5823,f5824,f5825,f5826,f6405,f6871,f6400,f6887,f6893,f6894,f6163,f6175,f6180,f6181,f6929,f6306,f6318,f6888,f7049,f7045])).
fof(f7045,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X51) = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl34),
  inference(trivial_inequality_removal,[],[f6957])).
fof(f6957,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X51) = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f2501])).
fof(f7049,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl34),
  inference(condensation,[],[f7048])).
fof(f7048,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK7(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl34),
  inference(subsumption_resolution,[],[f6942,f130])).
fof(f6942,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK7(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl34),
  inference(backward_demodulation,[],[f6888,f1330])).
fof(f7070,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.bool.DATATYPE_1'(s(X1,X2)))) ) | ($spl22 | $spl24 | $spl34)),
  inference(condensation,[],[f7069])).
fof(f7069,plain,(
  ( ! [X80,X81,X79,X77] : (s(bool,X77) = s(bool,'const.bool.DATATYPE_1'(s(X79,X80))) | s(bool,X77) = s(bool,X81)) ) | ($spl22 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f6620,f7068])).
fof(f6872,plain,(
  ( ! [X123,X124,X122] : (s(bool,'T_0') != s(bool,X122) | s(bool,X122) = s(bool,'const.bool.DATATYPE_1'(s(X123,X124)))) ) | $spl30),
  inference(cnf_transformation,[],[f6872_D])).
fof(f6872_D,plain,(
  ( ! [X123,X124,X122] : (s(bool,'T_0') != s(bool,X122) | s(bool,X122) = s(bool,'const.bool.DATATYPE_1'(s(X123,X124)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6888,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(cnf_transformation,[],[f6888_D])).
fof(f6888_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f6887,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(X2,X3)))) )),
  inference(condensation,[],[f6886])).
fof(f6886,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X182) | s(bool,X181) = s(bool,X183) | s(bool,'F_0') = s(bool,X183) | s(bool,X182) = s(bool,'const.bool.DATATYPE_1'(s(X184,X185)))) )),
  inference(subsumption_resolution,[],[f6114,f4534])).
fof(f6114,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,X182) | s(bool,X181) = s(bool,X183) | s(bool,'F_0') = s(bool,X183) | s(bool,X182) = s(bool,'const.bool.DATATYPE_1'(s(X184,X185))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f4534,f3843])).
fof(f6869,plain,(
  ( ! [X110] : (s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110)))) ) | $spl28),
  inference(cnf_transformation,[],[f6869_D])).
fof(f6869_D,plain,(
  ( ! [X110] : (s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6873,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,X131) != s(bool,X132) | s(bool,X131) = s(bool,'const.bool.DATATYPE_1'(s(X133,X134))) | s(bool,'F_0') = s(bool,X132)) ) | $spl32),
  inference(cnf_transformation,[],[f6873_D])).
fof(f6873_D,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,X131) != s(bool,X132) | s(bool,X131) = s(bool,'const.bool.DATATYPE_1'(s(X133,X134))) | s(bool,'F_0') = s(bool,X132)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6876,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,X141) != s(bool,X142) | s(bool,X140) = s(bool,'const.bool.DATATYPE_1'(s(bool,X140))) | s(bool,'F_0') = s(bool,X142) | s(bool,X141) = s(bool,'const.bool.DATATYPE_1'(s(X143,X144)))) ) | $spl16),
  inference(subsumption_resolution,[],[f6106,f1837])).
fof(f6106,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,X141) != s(bool,X142) | s(bool,X140) = s(bool,'const.bool.DATATYPE_1'(s(bool,X140))) | s(bool,'F_0') = s(bool,X142) | s(bool,X141) = s(bool,'const.bool.DATATYPE_1'(s(X143,X144))) | s(bool,'T_0') = s(bool,X140)) ) | $spl16),
  inference(superposition,[],[f2266,f3843])).
fof(f2266,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,'const.bool.DATATYPE_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X3)) ) | $spl16),
  inference(superposition,[],[f1833,f1650])).
fof(f1738,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl10),
  inference(cnf_transformation,[],[f1738_D])).
fof(f1738_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f6400,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X130) | s(bool,'F_0') = s(bool,X132) | s(bool,X131) = s(bool,'const.bool.DATATYPE_1'(s(X133,X134)))) )),
  inference(duplicate_literal_removal,[],[f6104])).
fof(f6104,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X130) | s(bool,'F_0') = s(bool,X132) | s(bool,X131) = s(bool,'const.bool.DATATYPE_1'(s(X133,X134))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f2056,f3843])).
fof(f6871,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,X122) | s(bool,X121) = s(bool,'const.bool.DATATYPE_1'(s(bool,X121))) | s(bool,X122) = s(bool,'const.bool.DATATYPE_1'(s(X123,X124)))) )),
  inference(subsumption_resolution,[],[f6102,f1837])).
fof(f6102,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,X122) | s(bool,X121) = s(bool,'const.bool.DATATYPE_1'(s(bool,X121))) | s(bool,X122) = s(bool,'const.bool.DATATYPE_1'(s(X123,X124))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f1837,f3843])).
fof(f6867,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X98),bool),'const.pair.FST_0'))) | s(bool,X97) = s(bool,'const.bool.DATATYPE_1'(s(X99,X100)))) ) | $spl26),
  inference(cnf_transformation,[],[f6867_D])).
fof(f6867_D,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X98),bool),'const.pair.FST_0'))) | s(bool,X97) = s(bool,'const.bool.DATATYPE_1'(s(X99,X100)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f6870,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X117) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X118),bool),'const.pair.FST_0'))) | s(bool,X116) = s(bool,'const.bool.DATATYPE_1'(s(bool,X116))) | s(bool,X117) = s(bool,'const.bool.DATATYPE_1'(s(X119,X120)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6101,f1837])).
fof(f6101,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X117) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X118),bool),'const.pair.FST_0'))) | s(bool,X116) = s(bool,'const.bool.DATATYPE_1'(s(bool,X116))) | s(bool,X117) = s(bool,'const.bool.DATATYPE_1'(s(X119,X120))) | s(bool,'T_0') = s(bool,X116)) ) | $spl18),
  inference(superposition,[],[f1834,f3843])).
fof(f6865,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X72,X73))) | s(bool,X71) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75)))) ) | $spl24),
  inference(cnf_transformation,[],[f6865_D])).
fof(f6865_D,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X72,X73))) | s(bool,X71) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6868,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,X111) != s(bool,'const.bool.DATATYPE_1'(s(X112,X113))) | s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110))) | s(bool,X111) = s(bool,'const.bool.DATATYPE_1'(s(X114,X115)))) ) | $spl16),
  inference(subsumption_resolution,[],[f6100,f1837])).
fof(f6100,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,X111) != s(bool,'const.bool.DATATYPE_1'(s(X112,X113))) | s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110))) | s(bool,X111) = s(bool,'const.bool.DATATYPE_1'(s(X114,X115))) | s(bool,'T_0') = s(bool,X110)) ) | $spl16),
  inference(superposition,[],[f1833,f3843])).
fof(f6402,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X98),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.bool.DATATYPE_1'(s(X99,X100)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6097])).
fof(f6097,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X98),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X96) | s(bool,X97) = s(bool,'const.bool.DATATYPE_1'(s(X99,X100))) | s(bool,'T_0') = s(bool,X96)) ) | $spl6),
  inference(superposition,[],[f1158,f3843])).
fof(f1158,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X20),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f1137,f127])).
fof(f1137,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X20),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f214,f1102])).
fof(f1102,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f997,f132])).
fof(f6405,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X72,X73))) | s(bool,'T_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75)))) )),
  inference(duplicate_literal_removal,[],[f6092])).
fof(f6092,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,X71) != s(bool,'const.bool.DATATYPE_1'(s(X72,X73))) | s(bool,'T_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f381,f3843])).
fof(f6620,plain,(
  ( ! [X80,X83,X81,X79,X77,X82] : (s(bool,X77) = s(bool,'const.bool.DATATYPE_1'(s(X79,X80))) | s(bool,X77) = s(bool,X81) | s(bool,X81) = s(bool,'const.bool.DATATYPE_1'(s(X82,X83)))) ) | $spl22),
  inference(cnf_transformation,[],[f6620_D])).
fof(f6620_D,plain,(
  ( ! [X80,X83,X81,X79,X77,X82] : (s(bool,X77) = s(bool,'const.bool.DATATYPE_1'(s(X79,X80))) | s(bool,X77) = s(bool,X81) | s(bool,X81) = s(bool,'const.bool.DATATYPE_1'(s(X82,X83)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6493,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'T_0') = s(bool,X67) | s(bool,X68) = s(bool,'const.bool.DATATYPE_1'(s(X69,X70))) | s(bool,X66) = s(bool,X68) | s(bool,X66) = s(bool,'const.bool.DATATYPE_1'(s(X71,X72)))) )),
  inference(duplicate_literal_removal,[],[f5806])).
fof(f5806,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'T_0') = s(bool,X67) | s(bool,X68) = s(bool,'const.bool.DATATYPE_1'(s(X69,X70))) | s(bool,X66) = s(bool,X68) | s(bool,X66) = s(bool,'const.bool.DATATYPE_1'(s(X71,X72))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f3843,f3843])).
fof(f6503,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,X92) = s(bool,X93) | s(bool,X93) = s(bool,'const.bool.DATATYPE_1'(s(X94,X95))) | s(bool,'T_0') = s(bool,X91) | s(bool,X92) = s(bool,'const.bool.DATATYPE_1'(s(X96,X97)))) )),
  inference(duplicate_literal_removal,[],[f5750])).
fof(f5750,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,X92) = s(bool,X93) | s(bool,X93) = s(bool,'const.bool.DATATYPE_1'(s(X94,X95))) | s(bool,'T_0') = s(bool,X91) | s(bool,X92) = s(bool,'const.bool.DATATYPE_1'(s(X96,X97))) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f3843,f3843])).
fof(f6613,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X34) = s(bool,'const.bool.DATATYPE_1'(s(X36,X37))) | s(bool,X34) = s(bool,X38) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38)))) ) | $spl20),
  inference(cnf_transformation,[],[f6613_D])).
fof(f6613_D,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X34) = s(bool,'const.bool.DATATYPE_1'(s(X36,X37))) | s(bool,X34) = s(bool,X38) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6506,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X27,X28))) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25)))) )),
  inference(duplicate_literal_removal,[],[f5738])).
fof(f5738,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X27,X28))) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f3843,f553])).
fof(f6528,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,X78) | s(bool,X77) = s(bool,'const.bool.DATATYPE_1'(s(X79,X80))) | s(bool,X81) = s(bool,'const.bool.DATATYPE_1'(s(X82,X83))) | s(bool,X77) = s(bool,X81)) )),
  inference(duplicate_literal_removal,[],[f5703])).
fof(f5703,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,X78) | s(bool,X77) = s(bool,'const.bool.DATATYPE_1'(s(X79,X80))) | s(bool,'T_0') = s(bool,X78) | s(bool,X81) = s(bool,'const.bool.DATATYPE_1'(s(X82,X83))) | s(bool,X77) = s(bool,X81)) )),
  inference(superposition,[],[f3843,f3843])).
fof(f6529,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,'const.bool.DATATYPE_1'(s(X36,X37))) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38))) | s(bool,X34) = s(bool,X38)) )),
  inference(duplicate_literal_removal,[],[f5695])).
fof(f5695,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,'const.bool.DATATYPE_1'(s(X36,X37))) | s(bool,'T_0') = s(bool,X35) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38))) | s(bool,X34) = s(bool,X38)) )),
  inference(superposition,[],[f3843,f553])).
fof(f1917,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f126,f125,f124,f127,f129,f130,f132,f131,f136,f135,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f267,f302,f341,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f833,f863,f865,f899,f951,f995,f970,f837,f836,f835,f834,f938,f940,f942,f952,f982,f996,f1078,f897,f932,f903,f862,f861,f856,f854,f853,f852,f851,f850,f866,f417,f416,f415,f972,f155,f898,f1190,f1192,f1194,f1234,f1218,f1219,f1323,f1325,f1327,f1328,f1330,f1332,f1730,f1714,f1702,f1697,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1426,f1427,f1428,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1545,f1552,f1566,f1865,f1651,f1650,f1609,f1616,f1738,f1915,f1913])).
fof(f1913,plain,(
  ( ! [X109,X107,X110,X108] : (s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110))) | s(bool,X110) = s(bool,g(s(X107,sK7(X108,X109,bool,X107))))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f1884])).
fof(f1884,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,X110) = s(bool,'const.bool.DATATYPE_1'(s(bool,X110))) | s(bool,X110) = s(bool,g(s(X107,sK7(X108,X109,bool,X107))))) ) | $spl10),
  inference(backward_demodulation,[],[f1738,f1609])).
fof(f1915,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f1871,f1738])).
fof(f1871,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f1738,f135])).
fof(f1916,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f126,f125,f124,f127,f129,f130,f132,f131,f136,f135,f148,f152,f153,f154,f156,f157,f158,f159,f160,f188,f164,f163,f162,f168,f165,f169,f170,f171,f172,f173,f174,f175,f176,f177,f178,f179,f195,f211,f214,f247,f267,f302,f341,f381,f392,f393,f394,f419,f445,f446,f447,f460,f470,f471,f472,f553,f804,f805,f825,f826,f827,f815,f817,f819,f828,f829,f831,f833,f863,f865,f899,f951,f995,f970,f837,f836,f835,f834,f938,f940,f942,f952,f982,f996,f1078,f897,f932,f903,f862,f861,f856,f854,f853,f852,f851,f850,f866,f417,f416,f415,f972,f155,f898,f1190,f1192,f1194,f1234,f1218,f1219,f1323,f1325,f1327,f1328,f1330,f1332,f1730,f1714,f1702,f1697,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1688,f1687,f1426,f1427,f1428,f1429,f1430,f1431,f1774,f1667,f1661,f1837,f1545,f1552,f1566,f1865,f1651,f1650,f1609,f1616,f1738,f1915])).
fof(f1914,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1870,f127])).
fof(f1870,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f1738,f129])).
fof(f1764,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1)))) ) | $spl14),
  inference(cnf_transformation,[],[f1764_D])).
fof(f1764_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1651,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.bool.DATATYPE_1'(s(bool,X48))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(duplicate_literal_removal,[],[f1592])).
fof(f1592,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,'const.bool.DATATYPE_1'(s(bool,X48))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f130,f553])).
fof(f1834,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X42) = s(bool,'const.bool.DATATYPE_1'(s(bool,X42)))) ) | $spl18),
  inference(cnf_transformation,[],[f1834_D])).
fof(f1834_D,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X42) = s(bool,'const.bool.DATATYPE_1'(s(bool,X42)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1660,plain,(
  ( ! [X43,X41,X42] : (s(bool,X42) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.bool.DATATYPE_1'(s(bool,X42)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1521])).
fof(f1521,plain,(
  ( ! [X43,X41,X42] : (s(bool,X42) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.bool.DATATYPE_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41)) ) | $spl6),
  inference(superposition,[],[f1158,f553])).
fof(f1833,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,'const.bool.DATATYPE_1'(s(X39,X40))) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38)))) ) | $spl16),
  inference(cnf_transformation,[],[f1833_D])).
fof(f1833_D,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,'const.bool.DATATYPE_1'(s(X39,X40))) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1661,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,'const.bool.DATATYPE_1'(s(X39,X40))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38)))) )),
  inference(duplicate_literal_removal,[],[f1520])).
fof(f1520,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,'const.bool.DATATYPE_1'(s(X39,X40))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,'const.bool.DATATYPE_1'(s(bool,X38))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f381,f553])).
fof(f1667,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1)))) )),
  inference(duplicate_literal_removal,[],[f1506])).
fof(f1506,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f130,f553])).
fof(f1687,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.DATATYPE_1'(s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) )),
  inference(duplicate_literal_removal,[],[f1424])).
fof(f1424,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.DATATYPE_1'(s(bool,X21))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f553,f130])).
fof(f1739,plain,(
  ( ! [X23,X25] : (s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(bool,X23))) | s(bool,X23) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25)))) ) | $spl12),
  inference(cnf_transformation,[],[f1739_D])).
fof(f1739_D,plain,(
  ( ! [X23,X25] : (s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(bool,X23))) | s(bool,X23) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1688,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,'const.bool.DATATYPE_1'(s(bool,X15))) | s(bool,X13) = s(bool,X15) | s(bool,X13) = s(bool,'const.bool.DATATYPE_1'(s(bool,X13)))) )),
  inference(duplicate_literal_removal,[],[f1423])).
fof(f1423,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,'const.bool.DATATYPE_1'(s(bool,X15))) | s(bool,X13) = s(bool,X15) | s(bool,X13) = s(bool,'const.bool.DATATYPE_1'(s(bool,X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f553,f553])).
fof(f1697,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,X16) = s(bool,'const.bool.DATATYPE_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,'const.bool.DATATYPE_1'(s(bool,X15)))) )),
  inference(duplicate_literal_removal,[],[f1392])).
fof(f1392,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,X16) = s(bool,'const.bool.DATATYPE_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,'const.bool.DATATYPE_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f553,f553])).
fof(f1702,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1387])).
fof(f1387,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.DATATYPE_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f553,f130])).
fof(f1714,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,'const.bool.DATATYPE_1'(s(bool,X20))) | s(bool,X22) = s(bool,'const.bool.DATATYPE_1'(s(bool,X22))) | s(bool,X20) = s(bool,X22)) )),
  inference(duplicate_literal_removal,[],[f1372])).
fof(f1372,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,'const.bool.DATATYPE_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.DATATYPE_1'(s(bool,X22))) | s(bool,X20) = s(bool,X22)) )),
  inference(superposition,[],[f553,f553])).
fof(f1730,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25))) | s(bool,X23) = s(bool,X25)) )),
  inference(duplicate_literal_removal,[],[f1342])).
fof(f1342,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,'const.bool.DATATYPE_1'(s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.DATATYPE_1'(s(bool,X25))) | s(bool,X23) = s(bool,X25)) )),
  inference(superposition,[],[f553,f553])).
fof(f1077,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1008,f129])).
fof(f1008,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f195,f998,f165])).
fof(f1076,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1014,f129])).
fof(f1014,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f195,f998,f165])).
fof(f1019,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f148,f998,f131])).
fof(f1017,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f195,f998,f136])).
fof(f1015,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f129,f195,f998,f165])).
fof(f1013,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f998,f195,f998,f165])).
fof(f1009,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f129,f195,f998,f165])).
fof(f1007,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f998,f195,f998,f165])).
fof(f1003,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f195,f195,f998,f247])).
fof(f999,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f127,f998])).
fof(f997,plain,(
  ( ! [X14] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0'))))) ) | $spl6),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  ( ! [X14] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0'))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f998,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) | $spl8),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f995,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f952,f172])).
fof(f191,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f190])).
fof(f190,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f189,f183])).
fof(f189,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK2))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f182,f185])).
fof(f181,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | $spl0),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f183,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK3) | $spl2),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f124,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK3) | s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4)))),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  (s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),sK3) & s('type.wordSem.word_loc'(sK0),sK2) = s('type.wordSem.word_loc'(sK0),sK4))) & (s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK3) | s('type.wordSem.word_loc'(sK0),sK2) != s('type.wordSem.word_loc'(sK0),sK4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4))) & (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X3) | s('type.wordSem.word_loc'(X0),X2) != s('type.wordSem.word_loc'(X0),X4)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4))) & (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) != s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) | (s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X3) | s('type.wordSem.word_loc'(X0),X2) != s('type.wordSem.word_loc'(X0),X4))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) <~> (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) <=> (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X1),s('type.wordSem.word_loc'(X0),X2))) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4))) <=> (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X3) & s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),X4)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))) = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X9),s('type.wordSem.word_loc'(X6),X10))) <=> (s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X9) & s('type.wordSem.word_loc'(X6),X8) = s('type.wordSem.word_loc'(X6),X10)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X7),s('type.wordSem.word_loc'(X6),X8))) = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X9),s('type.wordSem.word_loc'(X6),X10))) <=> (s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X9) & s('type.wordSem.word_loc'(X6),X8) = s('type.wordSem.word_loc'(X6),X10)))),
  file('thm.wordSem.result_|49||49||split|1','thm.wordSem.result_|49||49||split|1')).
fof(f185,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) = s('type.wordSem.word_loc'(sK0),sK4) | $spl4),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) = s('type.wordSem.word_loc'(sK0),sK4) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f125,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) = s('type.wordSem.word_loc'(sK0),sK4) | s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4)))),
  inference(cnf_transformation,[],[f107])).
fof(f182,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) | $spl1),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f184,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK3) | $spl3),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f186,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) != s('type.wordSem.word_loc'(sK0),sK4) | $spl5),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) != s('type.wordSem.word_loc'(sK0),sK4) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f126,plain,(
  s('type.wordSem.word_loc'(sK0),sK2) != s('type.wordSem.word_loc'(sK0),sK4) | s('type.wordSem.word_loc'(sK0),sK1) != s('type.wordSem.word_loc'(sK0),sK3) | s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK1),s('type.wordSem.word_loc'(sK0),sK2))) != s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK3),s('type.wordSem.word_loc'(sK0),sK4)))),
  inference(cnf_transformation,[],[f107])).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 107
% Passive clauses: 1135
% Generated clauses: 16570
% Final active clauses: 80
% Final passive clauses: 603
% Input formulas: 37
% Initial clauses: 56
% 
% Duplicate literals: 575
% Trivial inequalities: 144
% Fw subsumption resolutions: 803
% Fw demodulations: 2430
% Bw demodulations: 360
% Condensations: 1477
% Global subsumptions: 19
% 
% Simple tautologies: 144
% Equational tautologies: 79
% Forward subsumptions: 8804
% Fw demodulations to eq. taut.: 908
% Bw demodulations to eq. taut.: 50
% 
% Unit resulting resolution: 1420
% Factoring: 1
% Forward superposition: 4688
% Backward superposition: 3858
% Self superposition: 193
% Equality factoring: 92
% Equality resolution: 23
% 
% Split clauses: 27
% Split components: 55
% SAT solver clauses: 1246
% SAT solver unit clauses: 84
% SAT solver binary clauses: 447
% 
% Sat splits: 27
% Sat splitting refutations: 36
% 
% Memory used [KB]: 14456
% Time elapsed: 0.787 s
% ------------------------------
% ------------------------------
% Success in time 2.14 s
