Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (57287)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 921
% Passive clauses: 11598
% Generated clauses: 23379
% Final active clauses: 919
% Final passive clauses: 10308
% Input formulas: 133
% Initial clauses: 642
% 
% Duplicate literals: 339
% Trivial inequalities: 46
% Fw subsumption resolutions: 1374
% Fw demodulations: 1143
% Bw demodulations: 1
% 
% Simple tautologies: 312
% Equational tautologies: 60
% Forward subsumptions: 8301
% 
% Binary resolution: 13383
% Factoring: 1786
% Forward superposition: 3048
% Backward superposition: 1411
% Self superposition: 13
% Equality factoring: 4
% Equality resolution: 35
% 
% Split clauses: 317
% Split components: 635
% SAT solver clauses: 584
% SAT solver unit clauses: 19
% SAT solver binary clauses: 98
% 
% Sat splits: 317
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 584
% TWLsolver calls for satisfiability: 479
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (57316)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 600
% Passive clauses: 3242
% Generated clauses: 11352
% Final active clauses: 593
% Final passive clauses: 2414
% Input formulas: 133
% Initial clauses: 642
% 
% Duplicate literals: 100
% Trivial inequalities: 8
% Fw subsumption resolutions: 546
% Fw demodulations: 746
% Bw demodulations: 233
% Condensations: 49
% Global subsumptions: 22
% 
% Simple tautologies: 228
% Equational tautologies: 21
% Forward subsumptions: 2933
% 
% Binary resolution: 3198
% Factoring: 149
% Forward superposition: 1771
% Backward superposition: 578
% Self superposition: 11
% Equality factoring: 4
% Equality resolution: 17
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 3972
% SAT solver unit clauses: 657
% SAT solver binary clauses: 252
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 6
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3884
% TWLsolver calls for satisfiability: 10693
% 
% Memory used [KB]: 15991
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f15195,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f793,f1448,f1447,f792,f1450,f791,f1452,f1446,f12104,f12128,f12134,f15124])).
fof(f15124,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f15123,f12541])).
fof(f12541,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK246('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56))))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f12537,f1448])).
fof(f12537,plain,(
  s('type.exhSem.v',sK58) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK246('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56)))) | $spl0),
  inference(unit_resulting_resolution,[],[f12533,f1085])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK246(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK245(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK246(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK245(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK246(X2,X1,X0))))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP21(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK245,sK246])],[f642])).
fof(f642,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | ? [X5,X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) & (! [X3,X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP21(X0,X1,X2)))),
  inference(rectify,[],[f641])).
fof(f641,plain,(
  ! [X0,X2,X1] : ((~sP21(X0,X2,X1) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))) & (! [X3,X4] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))) | sP21(X0,X2,X1)))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X2,X1] : (sP21(X0,X2,X1) <=> ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP21])])).
fof(f12533,plain,(
  sP21(sK56,sK58,'const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))) | $spl0),
  inference(unit_resulting_resolution,[],[f9494,f5104,f7048,f1002,f1003,f12141,f1074])).
fof(f1074,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK244(X2,X1,X0)))) | sP19(X0,X1,X2) | sP18(X0,X1,X2) | sP20(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK243(X2,X1,X0)))) | sP21(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK243(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK243(X2,X1,X0))))) | sP20(X0,X1,X2) | sP18(X0,X1,X2) | sP19(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK244(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK244(X2,X1,X0))))) | sP21(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP20(X0,X1,X2) & ~sP18(X0,X1,X2) & ~sP19(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP21(X0,X1,X2)) | sP22(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK243,sK244])],[f639])).
fof(f639,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | ? [X6] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X6))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X6)))) | sP20(X0,X1,X2) | sP18(X0,X1,X2) | sP19(X0,X1,X2) | ? [X5] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP21(X0,X1,X2)) & ((! [X4] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X4)))) & ~sP20(X0,X1,X2) & ~sP18(X0,X1,X2) & ~sP19(X0,X1,X2) & ! [X3] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X3))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X3)))) & ~sP21(X0,X1,X2)) | sP22(X0,X1,X2)))),
  inference(rectify,[],[f638])).
fof(f638,plain,(
  ! [X0,X2,X1] : ((~sP22(X0,X2,X1) | ? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP20(X0,X2,X1) | sP18(X0,X2,X1) | sP19(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP21(X0,X2,X1)) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP20(X0,X2,X1) & ~sP18(X0,X2,X1) & ~sP19(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP21(X0,X2,X1)) | sP22(X0,X2,X1)))),
  inference(flattening,[],[f637])).
fof(f637,plain,(
  ! [X0,X2,X1] : ((~sP22(X0,X2,X1) | (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP20(X0,X2,X1) | sP18(X0,X2,X1) | sP19(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP21(X0,X2,X1))) & ((! [X19] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) & ~sP20(X0,X2,X1) & ~sP18(X0,X2,X1) & ~sP19(X0,X2,X1) & ! [X5] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) & ~sP21(X0,X2,X1)) | sP22(X0,X2,X1)))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X2,X1] : (sP22(X0,X2,X1) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | sP20(X0,X2,X1) | sP18(X0,X2,X1) | sP19(X0,X2,X1) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | sP21(X0,X2,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f12141,plain,(
  sP22(sK56,sK58,'const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))) | $spl0),
  inference(unit_resulting_resolution,[],[f1446,f1102])).
fof(f1102,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | sP22(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f653])).
fof(f653,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | sP22(X0,X2,X1)) & (~sP22(X0,X2,X1) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> sP22(X0,X2,X1))),
  inference(definition_folding,[],[f279,f437,f436,f435,f434,f433])).
fof(f433,plain,(
  ! [X0,X2,X1] : (sP18(X0,X2,X1) <=> ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f434,plain,(
  ! [X0,X2,X1] : (sP19(X0,X2,X1) <=> ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP19])])).
fof(f435,plain,(
  ! [X0,X2,X1] : (sP20(X0,X2,X1) <=> ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP20])])).
fof(f279,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> (? [X19] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X19))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X19)))) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | ? [X5] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X5))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X5)))) | ? [X3,X4] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X3),s('type.list.list'('type.exhSem.v'),X4)))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X15,X16,X17] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conSem.v',X16),s('type.exhSem.v',X17)))) <=> (? [X18] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X18))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X18)))) | ? [X19,X7,X9] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X19),s('type.list.list'('type.conSem.v'),X7))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X19))),s('type.list.list'('type.exhSem.v'),X9))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X9))))) | ? [X20,X21,X22,X23,X24] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.string.char'),X21),s('type.conLang.exp',X22))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.string.char'),X21),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.conLang.exp',X22))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X23,X20,X21,X25,X24] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25),s('type.list.list'('type.string.char'),X21))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X25))),s('type.list.list'('type.string.char'),X21))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X24),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X20),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X23))))) | ? [X18] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X18))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X18)))) | ? [X7,X9] : (s('type.conSem.v',X16) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X7))) & s('type.exhSem.v',X17) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X9))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X9)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.dec_to_exhProof.v_rel_cases|split|0')).
fof(f1003,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X104,X52] : ~s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X52))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X104)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_distinct|split|14')).
fof(f1002,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ~s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X104,X52] : ~s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X52))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X104)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_distinct|split|4')).
fof(f7048,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,'const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2)))) )),
  inference(unit_resulting_resolution,[],[f1413,f1097])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK255(X2,X1,X0)),s('type.list.list'('type.string.char'),sK256(X2,X1,X0)),s('type.conLang.exp',sK257(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK255(X2,X1,X0)),s('type.list.list'('type.string.char'),sK256(X2,X1,X0)),s('type.conLang.exp',sK257(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK258(X2,X1,X0)),s('type.list.list'('type.string.char'),sK256(X2,X1,X0)),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK259(X2,X1,X0)),s('type.conLang.exp',sK257(X2,X1,X0)))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK259(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK255(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK258(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP18(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK255,sK256,sK257,sK258,sK259])],[f651])).
fof(f651,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.string.char'),X9),s('type.conLang.exp',X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11),s('type.list.list'('type.string.char'),X9),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.conLang.exp',X10))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X11)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.string.char'),X4),s('type.conLang.exp',X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.string.char'),X4),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.conLang.exp',X5))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP18(X0,X1,X2)))),
  inference(rectify,[],[f650])).
fof(f650,plain,(
  ! [X0,X2,X1] : ((~sP18(X0,X2,X1) | ? [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14)))))) & (! [X11,X12,X13,X14,X15] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.string.char'),X12),s('type.conLang.exp',X13))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14),s('type.list.list'('type.string.char'),X12),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.conLang.exp',X13))))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X14))))) | sP18(X0,X2,X1)))),
  inference(nnf_transformation,[],[f433])).
fof(f1413,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) != s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2,X3] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.string.char'),X1),s('type.conLang.exp',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X17,X16,X15,X52] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X52))) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X15),s('type.list.list'('type.string.char'),X16),s('type.conLang.exp',X17)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_distinct|split|11')).
fof(f5104,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,'const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2)))) )),
  inference(unit_resulting_resolution,[],[f1385,f1088])).
fof(f1088,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK247(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK248(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f646])).
fof(f646,plain,(
  ! [X0,X1,X2] : ((~sP20(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK247(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK248(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),sK247(X2,X1,X0)))),s('type.list.list'('type.exhSem.v'),sK249(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK248(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK249(X2,X1,X0))))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP20(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK247,sK248,sK249])],[f645])).
fof(f645,plain,(
  ! [X0,X1,X2] : ((~sP20(X0,X1,X2) | ? [X6,X7,X8] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6),s('type.list.list'('type.conSem.v'),X7))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X6))),s('type.list.list'('type.exhSem.v'),X8))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X8)))))) & (! [X3,X4,X5] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3),s('type.list.list'('type.conSem.v'),X4))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X3))),s('type.list.list'('type.exhSem.v'),X5))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X4),s('type.list.list'('type.exhSem.v'),X5))))) | sP20(X0,X1,X2)))),
  inference(rectify,[],[f644])).
fof(f644,plain,(
  ! [X0,X2,X1] : ((~sP20(X0,X2,X1) | ? [X16,X17,X18] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18)))))) & (! [X16,X17,X18] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X17),s('type.list.list'('type.exhSem.v'),X18))))) | sP20(X0,X2,X1)))),
  inference(nnf_transformation,[],[f435])).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2))) != s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.conSem.v'),X0)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X16,X15,X52] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X52))) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X15),s('type.list.list'('type.conSem.v'),X16)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_distinct|split|8')).
fof(f9494,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,'const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X2)))) )),
  inference(unit_resulting_resolution,[],[f1416,f1092])).
fof(f1092,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK251(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK253(X2,X1,X0)),s('type.list.list'('type.string.char'),sK252(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f649])).
fof(f649,plain,(
  ! [X0,X1,X2] : ((~sP19(X0,X1,X2) | (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK251(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK253(X2,X1,X0)),s('type.list.list'('type.string.char'),sK252(X2,X1,X0)))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK250(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK254(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),sK253(X2,X1,X0)))),s('type.list.list'('type.string.char'),sK252(X2,X1,X0)))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK254(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK251(X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK250(X2,X1,X0))))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP19(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK250,sK251,sK252,sK253,sK254])],[f648])).
fof(f648,plain,(
  ! [X0,X1,X2] : ((~sP19(X0,X1,X2) | ? [X8,X9,X10,X11,X12] : (s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11),s('type.list.list'('type.string.char'),X10))) & s('type.exhSem.v',X1) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X11))),s('type.list.list'('type.string.char'),X10))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))) & (! [X3,X4,X5,X6,X7] : (s('type.conSem.v',X2) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6),s('type.list.list'('type.string.char'),X5))) | s('type.exhSem.v',X1) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X6))),s('type.list.list'('type.string.char'),X5))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X7),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3))))) | sP19(X0,X1,X2)))),
  inference(rectify,[],[f647])).
fof(f647,plain,(
  ! [X0,X2,X1] : ((~sP19(X0,X2,X1) | ? [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) & p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6)))))) & (! [X6,X7,X8,X9,X10] : (s('type.conSem.v',X1) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9),s('type.list.list'('type.string.char'),X8))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X9))),s('type.list.list'('type.string.char'),X8))) | ~p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0)))) | ~p(s(bool,'const.dec_to_exhProof.env_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X6))))) | sP19(X0,X2,X1)))),
  inference(nnf_transformation,[],[f434])).
fof(f1416,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2,X3] : s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) != s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X3))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X1),s('type.list.list'('type.string.char'),X0)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X17,X16,X15,X52] : ~s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X52))) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X15),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X16),s('type.list.list'('type.string.char'),X17)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_distinct|split|13')).
fof(f15123,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK246('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56))))) | $spl0),
  inference(backward_demodulation,[],[f14957,f12536])).
fof(f12536,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK245('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56)),s('type.list.list'('type.exhSem.v'),sK246('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56))))) | $spl0),
  inference(unit_resulting_resolution,[],[f12533,f1086])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK245(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK246(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f14957,plain,(
  s('type.list.list'('type.conSem.v'),sK57) = s('type.list.list'('type.conSem.v'),sK245('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56)) | $spl0),
  inference(unit_resulting_resolution,[],[f12538,f974])).
fof(f974,plain,(
  ( ! [X0,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))) | s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1)) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1] : ((s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) != s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))) | s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1)) & (s('type.list.list'('type.conSem.v'),X0) != s('type.list.list'('type.conSem.v'),X1) | s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1)))))),
  inference(nnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X0))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))) <=> s('type.list.list'('type.conSem.v'),X0) = s('type.list.list'('type.conSem.v'),X1))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X52,X104] : (s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X52))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X104))) <=> s('type.list.list'('type.conSem.v'),X52) = s('type.list.list'('type.conSem.v'),X104))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.conSem.v_|49||49||split|5')).
fof(f12538,plain,(
  s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK245('const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57)),sK58,sK56)))) | $spl0),
  inference(unit_resulting_resolution,[],[f12533,f1084])).
fof(f1084,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | s('type.conSem.v',X2) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK245(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f12134,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f12133,f794])).
fof(f794,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3',aHL_TRUTH)).
fof(f12133,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f12126,f1511])).
fof(f1511,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59))) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f845])).
fof(f845,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f12126,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) | ($spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f12125])).
fof(f12125,plain,(
  s('type.exhSem.v',sK58) != s('type.exhSem.v',sK58) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) | ($spl2 | $spl4)),
  inference(superposition,[],[f1448,f1450])).
fof(f12128,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f12127,f794])).
fof(f12127,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f12118,f1511])).
fof(f12118,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f1450,f1448])).
fof(f12104,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f12103,f1453])).
fof(f1453,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f797,f794])).
fof(f797,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f48])).
fof(f48,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.bool.NOT_CLAUSES|split|2')).
fof(f12103,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f12102,f1520])).
fof(f1520,plain,(
  s(bool,'F_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58))) | $spl1),
  inference(unit_resulting_resolution,[],[f1447,f849])).
fof(f849,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f494])).
fof(f494,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.bool.EQ_CLAUSES|split|3')).
fof(f12102,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f12101,f1450])).
fof(f12101,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59)))))) | $spl6),
  inference(subsumption_resolution,[],[f12086,f794])).
fof(f12086,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59)))))) | $spl6),
  inference(superposition,[],[f1027,f1511])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X2))))))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X2)))))))),
  inference(ennf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) => p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X2)))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7,X9] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X9)))) => p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X7))),s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X9)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.dec_to_exhProof.v_rel_rules|split|5')).
fof(f1446,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) | $spl0),
  inference(cnf_transformation,[],[f1446_D])).
fof(f1446_D,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1452,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) | $spl6),
  inference(cnf_transformation,[],[f1452_D])).
fof(f1452_D,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f791,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58))))),
  inference(cnf_transformation,[],[f483])).
fof(f483,plain,(
  (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) | (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),sK59)))) & s('type.exhSem.v',sK58) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59))))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) | ! [X3] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),X3)))) | s('type.exhSem.v',sK58) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK56,sK57,sK58,sK59])],[f482])).
fof(f482,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) | ? [X4] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X4)))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X4))))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) | ! [X3] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X3)))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))))))),
  inference(rectify,[],[f481])).
fof(f481,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) | ? [X3] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X3)))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) | ! [X3] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X3)))) | s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))))))),
  inference(nnf_transformation,[],[f337])).
fof(f337,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) <~> ? [X3] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X3)))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X1))),s('type.exhSem.v',X2)))) <=> ? [X3] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X3)))) & s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X7))),s('type.exhSem.v',X8)))) <=> ? [X9] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X9)))) & s('type.exhSem.v',X8) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X7))),s('type.exhSem.v',X8)))) <=> ? [X9] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X9)))) & s('type.exhSem.v',X8) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X9)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3','thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|3')).
fof(f1450,plain,(
  s('type.exhSem.v',sK58) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59))) | $spl4),
  inference(cnf_transformation,[],[f1450_D])).
fof(f1450_D,plain,(
  s('type.exhSem.v',sK58) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f792,plain,(
  s('type.exhSem.v',sK58) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK59))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58))))),
  inference(cnf_transformation,[],[f483])).
fof(f1447,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) | $spl1),
  inference(cnf_transformation,[],[f1447_D])).
fof(f1447_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1448,plain,(
  ( ! [X3] : (s('type.exhSem.v',sK58) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),X3))))) ) | $spl2),
  inference(cnf_transformation,[],[f1448_D])).
fof(f1448_D,plain,(
  ( ! [X3] : (s('type.exhSem.v',sK58) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),X3))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f793,plain,(
  ( ! [X3] : (s('type.exhSem.v',sK58) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X3))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.list.list'('type.conSem.v'),sK57),s('type.list.list'('type.exhSem.v'),X3)))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK56),s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),sK57))),s('type.exhSem.v',sK58))))) )),
  inference(cnf_transformation,[],[f483])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 906
% Passive clauses: 2952
% Generated clauses: 14420
% Final active clauses: 809
% Final passive clauses: 1891
% Input formulas: 133
% Initial clauses: 642
% 
% Duplicate literals: 38
% Trivial inequalities: 7
% Fw subsumption resolutions: 98
% Fw demodulations: 2653
% Bw demodulations: 22
% Condensations: 12
% 
% Simple tautologies: 94
% Equational tautologies: 40
% Forward subsumptions: 7495
% Fw demodulations to eq. taut.: 1024
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 4749
% Factoring: 1
% Forward superposition: 3516
% Backward superposition: 2608
% Self superposition: 22
% Equality factoring: 4
% Equality resolution: 24
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 2964
% SAT solver unit clauses: 1067
% SAT solver binary clauses: 1080
% 
% Sat splits: 3
% Sat splitting refutations: 4
% 
% Memory used [KB]: 18549
% Time elapsed: 0.498 s
% ------------------------------
% ------------------------------
% Success in time 1.842 s
