Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f2611,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2609,f101,f2362,f370,f368,f2595,f105,f163,f1244,f1035,f409,f415,f106,f159,f157,f403,f158,f2126,f1980,f363,f100,f156,f2133,f113,f169,f1976,f1984,f161,f1989,f1807,f1968,f1969,f1967,f623,f625,f1966,f1815,f350,f1784,f103,f167,f1255,f361,f1217,f454,f364,f104,f165,f155,f151,f170,f179,f182,f180,f188,f172,f1228,f1418,f944,f1252,f339,f359,f531,f1854,f357,f1813,f402,f360,f418,f1307,f1795,f945,f1787,f946,f1805,f102,f160,f162,f164,f166,f154,f1979,f2074,f2078,f2001,f349,f336,f354,f353,f2128,f1843,f2213,f366,f2215,f2599])).
fof(f2599,plain,(
  $false | ($spl1 | $spl16 | $spl40 | $spl108 | $spl110)),
  inference(subsumption_resolution,[],[f2598,f170])).
fof(f2598,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl40 | $spl108 | $spl110)),
  inference(forward_demodulation,[],[f2597,f2215])).
fof(f2597,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl1 | $spl40 | $spl108)),
  inference(forward_demodulation,[],[f2596,f2078])).
fof(f2596,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'T_0')))) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f155,f361])).
fof(f2215,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0'))) | $spl110),
  inference(cnf_transformation,[],[f2215_D])).
fof(f2215_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0'))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f366,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2213,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) ) | ($spl15 | $spl102)),
  inference(subsumption_resolution,[],[f2204,f176])).
fof(f176,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f169,f114])).
fof(f114,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relationProps.clos_relationProps_4|48|_prove_ltb',aHL_BOOL_CASES)).
fof(f2204,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) ) | $spl102),
  inference(superposition,[],[f1843,f174])).
fof(f174,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f114,f114])).
fof(f1843,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) | $spl102),
  inference(cnf_transformation,[],[f1843_D])).
fof(f1843_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2128,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0')))) | $spl8),
  inference(resolution,[],[f162,f88])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X1)) & ~p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f52])).
fof(f52,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f353,plain,(
  s(bool,'T_0') = s(bool,sK4) | $spl30),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  s(bool,'T_0') = s(bool,sK4) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f354,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,X2))),s(bool,sK3)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,X2))),s(bool,sK3)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f336,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,X2))),s(bool,sK3)))) | s(bool,'T_0') = s(bool,sK4) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f154,f174])).
fof(f349,plain,(
  ( ! [X0] : (s(bool,sK4) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  ( ! [X0] : (s(bool,sK4) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2001,plain,(
  ( ! [X0] : (s(bool,sK4) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl16)),
  inference(resolution,[],[f160,f245])).
fof(f245,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f170,f174])).
fof(f2078,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl108),
  inference(cnf_transformation,[],[f2078_D])).
fof(f2078_D,plain,(
  s(bool,'F_0') = s(bool,sK4) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2074,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2027])).
fof(f2027,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f349,f114])).
fof(f1979,plain,(
  $false | ($spl7 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f1978,f170])).
fof(f1978,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl30)),
  inference(forward_demodulation,[],[f160,f353])).
fof(f154,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) | $spl0),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f166,plain,(
  sP0(sK4,sK3,sK5) | $spl12),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  sP0(sK4,sK3,sK5) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f164,plain,(
  sP1(sK4,sK3,sK5) | $spl10),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  sP1(sK4,sK3,sK5) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f162,plain,(
  sP2(sK4,sK3,sK5) | $spl8),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  sP2(sK4,sK3,sK5) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f160,plain,(
  ~p(s(bool,sK4)) | $spl7),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  ~p(s(bool,sK4)) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f102,plain,(
  ~p(s(bool,sK4)) | sP2(sK4,sK3,sK5) | sP1(sK4,sK3,sK5) | sP0(sK4,sK3,sK5) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) | sP0(sK4,sK3,sK5) | sP1(sK4,sK3,sK5) | sP2(sK4,sK3,sK5) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,sK3)) & ~p(s(bool,sK4)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) | (~sP0(sK4,sK3,sK5) & ~sP1(sK4,sK3,sK5) & ~sP2(sK4,sK3,sK5) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,sK3)) | p(s(bool,sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5])],[f61])).
fof(f61,plain,(
  ? [X0,X1,X2] : ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) | sP0(X1,X0,X2) | sP1(X1,X0,X2) | sP2(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) | (~sP0(X1,X0,X2) & ~sP1(X1,X0,X2) & ~sP2(X1,X0,X2) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2] : ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) | (sP0(X1,X0,X2) | sP1(X1,X0,X2) | sP2(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1))))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) | (~sP0(X1,X0,X2) & ~sP1(X1,X0,X2) & ~sP2(X1,X0,X2) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ? [X0,X1,X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) <~> (sP0(X1,X0,X2) | sP1(X1,X0,X2) | sP2(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(definition_folding,[],[f45,f49,f48,f47])).
fof(f47,plain,(
  ! [X1,X0,X2] : (sP0(X1,X0,X2) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f48,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f45,plain,(
  ? [X0,X1,X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) <~> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,X1))),s(bool,X0)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,X7))),s(bool,X6)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,X7))),s(bool,X6)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)))))),
  file('thm.clos_relationProps.clos_relationProps_4|48|_prove_ltb','thm.clos_relationProps.clos_relationProps_4|48|_prove_ltb')).
fof(f1805,plain,(
  $false | ($spl8 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f1804,f170])).
fof(f1804,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl30)),
  inference(forward_demodulation,[],[f1801,f353])).
fof(f1801,plain,(
  ~p(s(bool,sK4)) | $spl8),
  inference(resolution,[],[f162,f90])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f53])).
fof(f946,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) | $spl81),
  inference(cnf_transformation,[],[f946_D])).
fof(f946_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f1787,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl1 | $spl30 | $spl40)),
  inference(forward_demodulation,[],[f1786,f353])).
fof(f1786,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'T_0')))) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f155,f361])).
fof(f945,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) | $spl80),
  inference(cnf_transformation,[],[f945_D])).
fof(f945_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1795,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) | $spl12),
  inference(resolution,[],[f166,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f58])).
fof(f58,plain,(
  ! [X1,X0,X2] : ((~sP0(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP0(X1,X0,X2)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X1,X0,X2] : ((~sP0(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) | sP0(X1,X0,X2)))),
  inference(nnf_transformation,[],[f47])).
fof(f1307,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl40 | $spl54)),
  inference(superposition,[],[f418,f361])).
fof(f418,plain,(
  ( ! [X0] : (s(bool,sK3) = s(bool,X0)) ) | $spl54),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  ( ! [X0] : (s(bool,sK3) = s(bool,X0)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f360,plain,(
  s(bool,'T_0') != s(bool,sK3) | $spl41),
  inference(cnf_transformation,[],[f360_D])).
fof(f360_D,plain,(
  s(bool,'T_0') != s(bool,sK3) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f402,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK3) | s(bool,sK3) = s(bool,X0)) ) | $spl36),
  inference(equality_factoring,[],[f357])).
fof(f1813,plain,(
  $false | ($spl10 | $spl16 | $spl40)),
  inference(subsumption_resolution,[],[f1812,f170])).
fof(f1812,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl40)),
  inference(forward_demodulation,[],[f1809,f361])).
fof(f1809,plain,(
  ~p(s(bool,sK3)) | $spl10),
  inference(resolution,[],[f164,f93])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X1)) | ~p(s(bool,X0)) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f55])).
fof(f55,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1)) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f48])).
fof(f357,plain,(
  ( ! [X0] : (s(bool,sK3) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(cnf_transformation,[],[f357_D])).
fof(f357_D,plain,(
  ( ! [X0] : (s(bool,sK3) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1854,plain,(
  ( ! [X0] : (s(bool,sK3) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl41),
  inference(trivial_inequality_removal,[],[f1849])).
fof(f1849,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,sK3) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl41),
  inference(superposition,[],[f360,f174])).
fof(f531,plain,(
  $false | ($spl15 | $spl16 | $spl46)),
  inference(subsumption_resolution,[],[f480,f170])).
fof(f480,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl46)),
  inference(backward_demodulation,[],[f366,f169])).
fof(f359,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'T_0')))) | $spl38),
  inference(cnf_transformation,[],[f359_D])).
fof(f359_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'T_0')))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f339,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'T_0')))) | s(bool,sK3) = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f154,f174])).
fof(f1252,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl30 | $spl38)),
  inference(forward_demodulation,[],[f359,f353])).
fof(f944,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0'))) | $spl78),
  inference(cnf_transformation,[],[f944_D])).
fof(f944_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1418,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'T_0'))) | ($spl15 | $spl80)),
  inference(resolution,[],[f945,f176])).
fof(f1228,plain,(
  $false | ($spl16 | $spl18)),
  inference(resolution,[],[f172,f170])).
fof(f172,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f188,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | ($spl15 | $spl22)),
  inference(superposition,[],[f169,f180])).
fof(f180,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f182,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl25),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f179,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f114])).
fof(f170,plain,(
  p(s(bool,'T_0')) | $spl16),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f151,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : ! [X11] : (s(X10,X11) = s(X10,X11) <=> p(s(bool,'T_0')))),
  file('thm.clos_relationProps.clos_relationProps_4|48|_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f155,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) | $spl1),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f165,plain,(
  ~sP1(sK4,sK3,sK5) | $spl11),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  ~sP1(sK4,sK3,sK5) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f104,plain,(
  ~sP1(sK4,sK3,sK5) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f364,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) | $spl44),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f454,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) | ($spl0 | $spl50)),
  inference(backward_demodulation,[],[f415,f154])).
fof(f1217,plain,(
  $false | ($spl11 | $spl15 | $spl16 | $spl25 | $spl30 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f1216,f169])).
fof(f1216,plain,(
  p(s(bool,'F_0')) | ($spl11 | $spl16 | $spl25 | $spl30 | $spl44 | $spl50)),
  inference(forward_demodulation,[],[f1212,f415])).
fof(f1212,plain,(
  p(s(bool,sK3)) | ($spl11 | $spl16 | $spl25 | $spl30 | $spl44)),
  inference(resolution,[],[f1205,f165])).
fof(f1205,plain,(
  ( ! [X0] : (sP1(sK4,X0,sK5) | p(s(bool,X0))) ) | ($spl16 | $spl25 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f1182,f182])).
fof(f1182,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP1(sK4,X0,sK5) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl16 | $spl30 | $spl44)),
  inference(resolution,[],[f1076,f606])).
fof(f606,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f598,f114])).
fof(f598,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl30 | $spl44)),
  inference(forward_demodulation,[],[f364,f353])).
fof(f1076,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X41),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X40)) | sP1(sK4,X40,X41)) ) | ($spl16 | $spl30)),
  inference(subsumption_resolution,[],[f1068,f170])).
fof(f1068,plain,(
  ( ! [X41,X40] : (~p(s(bool,'T_0')) | sP1(sK4,X40,X41) | p(s(bool,X40)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X41),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl30),
  inference(superposition,[],[f95,f353])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP1(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f361,plain,(
  s(bool,'T_0') = s(bool,sK3) | $spl40),
  inference(cnf_transformation,[],[f361_D])).
fof(f361_D,plain,(
  s(bool,'T_0') = s(bool,sK3) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1255,plain,(
  s(bool,'T_0') = s(bool,sK3) | ($spl4 | $spl15)),
  inference(resolution,[],[f159,f176])).
fof(f167,plain,(
  ~sP0(sK4,sK3,sK5) | $spl13),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  ~sP0(sK4,sK3,sK5) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f103,plain,(
  ~sP0(sK4,sK3,sK5) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f1784,plain,(
  $false | ($spl13 | $spl16 | $spl30 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f1783,f170])).
fof(f1783,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl16 | $spl30 | $spl40 | $spl78)),
  inference(forward_demodulation,[],[f1782,f353])).
fof(f1782,plain,(
  ~p(s(bool,sK4)) | ($spl13 | $spl16 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f1781,f170])).
fof(f1781,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK4)) | ($spl13 | $spl16 | $spl40 | $spl78)),
  inference(forward_demodulation,[],[f1777,f361])).
fof(f1777,plain,(
  ~p(s(bool,sK3)) | ~p(s(bool,sK4)) | ($spl13 | $spl16 | $spl78)),
  inference(resolution,[],[f1774,f167])).
fof(f1774,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,sK5) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl16 | $spl78)),
  inference(subsumption_resolution,[],[f1766,f170])).
fof(f1766,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP0(X0,X1,sK5)) ) | $spl78),
  inference(superposition,[],[f99,f944])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f59])).
fof(f350,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,sK3)))) | $spl29),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,sK3)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1815,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,sK3)))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f155,f353])).
fof(f1966,plain,(
  $false | ($spl16 | $spl29 | $spl50 | $spl56)),
  inference(subsumption_resolution,[],[f1965,f170])).
fof(f1965,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl50 | $spl56)),
  inference(forward_demodulation,[],[f1964,f625])).
fof(f1964,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl29 | $spl50)),
  inference(forward_demodulation,[],[f350,f415])).
fof(f625,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0'))) | $spl56),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0'))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f623,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) ) | ($spl15 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f614,f176])).
fof(f614,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f598,f174])).
fof(f1967,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | $spl107),
  inference(cnf_transformation,[],[f1967_D])).
fof(f1967_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1969,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl30 | $spl45)),
  inference(forward_demodulation,[],[f363,f353])).
fof(f1968,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | $spl106),
  inference(cnf_transformation,[],[f1968_D])).
fof(f1968_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1807,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'T_0'))),s(bool,'F_0')))) | $spl10),
  inference(resolution,[],[f164,f92])).
fof(f92,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'T_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f1989,plain,(
  s(bool,'T_0') = s(bool,sK4) | ($spl6 | $spl15)),
  inference(resolution,[],[f161,f176])).
fof(f161,plain,(
  p(s(bool,sK4)) | $spl6),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  p(s(bool,sK4)) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1984,plain,(
  p(s(bool,sK4)) | $spl10),
  inference(resolution,[],[f164,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f56])).
fof(f1976,plain,(
  $false | ($spl12 | $spl15 | $spl50)),
  inference(subsumption_resolution,[],[f1975,f169])).
fof(f1975,plain,(
  p(s(bool,'F_0')) | ($spl12 | $spl50)),
  inference(forward_demodulation,[],[f1973,f415])).
fof(f1973,plain,(
  p(s(bool,sK3)) | $spl12),
  inference(resolution,[],[f166,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f59])).
fof(f169,plain,(
  ~p(s(bool,'F_0')) | $spl15),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f113,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relationProps.clos_relationProps_4|48|_prove_ltb',aHL_FALSITY)).
fof(f2133,plain,(
  $false | ($spl8 | $spl15 | $spl50)),
  inference(subsumption_resolution,[],[f2132,f169])).
fof(f2132,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl50)),
  inference(forward_demodulation,[],[f2130,f415])).
fof(f2130,plain,(
  p(s(bool,sK3)) | $spl8),
  inference(resolution,[],[f162,f89])).
fof(f89,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f53])).
fof(f156,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | $spl2),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f100,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | sP2(sK4,sK3,sK5) | sP1(sK4,sK3,sK5) | sP0(sK4,sK3,sK5) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f363,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) | $spl45),
  inference(cnf_transformation,[],[f363_D])).
fof(f363_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1980,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,'F_0')))) | ($spl1 | $spl50)),
  inference(forward_demodulation,[],[f155,f415])).
fof(f2126,plain,(
  $false | ($spl2 | $spl45 | $spl108)),
  inference(subsumption_resolution,[],[f2125,f156])).
fof(f2125,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl45 | $spl108)),
  inference(forward_demodulation,[],[f363,f2078])).
fof(f158,plain,(
  ~p(s(bool,sK3)) | $spl5),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  ~p(s(bool,sK3)) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f403,plain,(
  ( ! [X9] : (~p(s(bool,sK3)) | s(bool,'T_0') = s(bool,X9)) ) | ($spl15 | $spl36)),
  inference(duplicate_literal_removal,[],[f396])).
fof(f396,plain,(
  ( ! [X9] : (~p(s(bool,sK3)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl15 | $spl36)),
  inference(superposition,[],[f176,f357])).
fof(f157,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | $spl3),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f159,plain,(
  p(s(bool,sK3)) | $spl4),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  p(s(bool,sK3)) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f106,plain,(
  p(s(bool,sK4)) | p(s(bool,sK3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f415,plain,(
  s(bool,'F_0') = s(bool,sK3) | $spl50),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s(bool,'F_0') = s(bool,sK3) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f409,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,sK3) | s(bool,'T_0') = s(bool,X9)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f382])).
fof(f382,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,sK3) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl36),
  inference(superposition,[],[f357,f114])).
fof(f1035,plain,(
  $false | ($spl3 | $spl25 | $spl32 | $spl50)),
  inference(subsumption_resolution,[],[f1000,f182])).
fof(f1000,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl32 | $spl50)),
  inference(resolution,[],[f710,f157])).
fof(f710,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,X2))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl32 | $spl50)),
  inference(forward_demodulation,[],[f354,f415])).
fof(f1244,plain,(
  ( ! [X4] : (s(bool,sK3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f539,f170])).
fof(f539,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,sK3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl5),
  inference(superposition,[],[f158,f174])).
fof(f163,plain,(
  ~sP2(sK4,sK3,sK5) | $spl9),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  ~sP2(sK4,sK3,sK5) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f105,plain,(
  ~sP2(sK4,sK3,sK5) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f2595,plain,(
  $false | ($spl4 | $spl9 | $spl15 | $spl16 | $spl108 | $spl110)),
  inference(subsumption_resolution,[],[f2594,f169])).
fof(f2594,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl9 | $spl16 | $spl108 | $spl110)),
  inference(forward_demodulation,[],[f2591,f2078])).
fof(f2591,plain,(
  p(s(bool,sK4)) | ($spl4 | $spl9 | $spl16 | $spl110)),
  inference(resolution,[],[f2584,f163])).
fof(f2584,plain,(
  ( ! [X0] : (sP2(X0,sK3,sK5) | p(s(bool,X0))) ) | ($spl4 | $spl16 | $spl110)),
  inference(subsumption_resolution,[],[f2573,f170])).
fof(f2573,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP2(X0,sK3,sK5) | p(s(bool,X0))) ) | ($spl4 | $spl110)),
  inference(superposition,[],[f2137,f2215])).
fof(f2137,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X3),s(bool,'F_0'))),s(bool,'T_0')))) | sP2(X2,sK3,X3) | p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f159,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1,X2) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X2),s(bool,'F_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f368,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))) | $spl48),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f370,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))) | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f348,f169])).
fof(f348,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))) | $spl0),
  inference(superposition,[],[f154,f114])).
fof(f2362,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,'T_0'))) | ($spl40 | $spl48 | $spl108)),
  inference(forward_demodulation,[],[f2134,f361])).
fof(f2134,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,'F_0'))),s(bool,sK3))) | ($spl48 | $spl108)),
  inference(forward_demodulation,[],[f368,f2078])).
fof(f101,plain,(
  ~p(s(bool,sK3)) | sP2(sK4,sK3,sK5) | sP1(sK4,sK3,sK5) | sP0(sK4,sK3,sK5) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),sK5),s(bool,sK4))),s(bool,sK3))))),
  inference(cnf_transformation,[],[f62])).
fof(f2609,plain,(
  $false | ($spl12 | $spl15 | $spl108)),
  inference(subsumption_resolution,[],[f2608,f169])).
fof(f2608,plain,(
  p(s(bool,'F_0')) | ($spl12 | $spl108)),
  inference(forward_demodulation,[],[f2605,f2078])).
fof(f2605,plain,(
  p(s(bool,sK4)) | $spl12),
  inference(resolution,[],[f166,f98])).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f59])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 122
% Passive clauses: 498
% Generated clauses: 2521
% Final active clauses: 39
% Final passive clauses: 132
% Input formulas: 21
% Initial clauses: 63
% 
% Duplicate literals: 242
% Trivial inequalities: 21
% Fw subsumption resolutions: 287
% Fw demodulations: 167
% Bw demodulations: 29
% 
% Simple tautologies: 75
% Equational tautologies: 26
% Forward subsumptions: 969
% 
% Binary resolution: 148
% Forward superposition: 951
% Backward superposition: 335
% Self superposition: 46
% Equality factoring: 40
% 
% Split clauses: 112
% Split components: 235
% SAT solver clauses: 303
% SAT solver unit clauses: 7
% SAT solver binary clauses: 45
% SAT solver learnt clauses: 8
% SAT solver learnt literals: 12
% 
% Sat splits: 112
% Sat splitting refutations: 28
% 
% TWLsolver clauses: 295
% TWLsolver calls for satisfiability: 89
% 
% Memory used [KB]: 1407
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
% Success in time 0.037 s
