Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (54532)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 382
% Passive clauses: 2461
% Generated clauses: 20241
% Final active clauses: 313
% Final passive clauses: 1874
% Input formulas: 109
% Initial clauses: 207
% 
% Duplicate literals: 513
% Trivial inequalities: 224
% Fw subsumption resolutions: 1369
% Fw demodulations: 3568
% Bw demodulations: 59
% 
% Simple tautologies: 353
% Equational tautologies: 44
% Forward subsumptions: 11500
% Fw demodulations to eq. taut.: 9
% 
% Binary resolution: 508
% Factoring: 5
% Forward superposition: 7428
% Backward superposition: 6263
% Self superposition: 124
% Equality factoring: 72
% Equality resolution: 27
% 
% Split clauses: 195
% Split components: 391
% SAT solver clauses: 394
% SAT solver unit clauses: 6
% SAT solver binary clauses: 56
% 
% Sat splits: 195
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 393
% TWLsolver calls for satisfiability: 129
% 
% Memory used [KB]: 13432
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (54563)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 351
% Passive clauses: 1518
% Generated clauses: 16258
% Final active clauses: 245
% Final passive clauses: 561
% Input formulas: 109
% Initial clauses: 207
% 
% Duplicate literals: 41
% Trivial inequalities: 97
% Fw subsumption resolutions: 610
% Fw demodulations: 3105
% Bw demodulations: 661
% Condensations: 45
% Global subsumptions: 19
% 
% Simple tautologies: 149
% Equational tautologies: 13
% Forward subsumptions: 9100
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 559
% Factoring: 9
% Forward superposition: 5078
% Backward superposition: 3898
% Self superposition: 56
% Equality factoring: 8
% Equality resolution: 21
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 1672
% SAT solver unit clauses: 302
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 40
% SAT solver learnt literals: 17
% 
% Sat splits: 19
% Sat splitting refutations: 36
% 
% TWLsolver clauses: 1594
% TWLsolver calls for satisfiability: 1240
% 
% Memory used [KB]: 11897
% Time elapsed: 0.715 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f10526,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1682,f1690,f1683,f1693,f1684,f1694,f1686,f1697,f1695,f1687,f1699,f1688,f1703,f2518,f2581,f1701,f4316,f3047,f4323,f2336,f4325,f1734,f1692,f4329,f10525])).
fof(f10525,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f10524,f4440])).
fof(f4440,plain,(
  s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) | $spl12),
  inference(backward_demodulation,[],[f4330,f2464])).
fof(f2464,plain,(
  s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0')))),
  inference(backward_demodulation,[],[f2403,f557])).
fof(f557,plain,(
  s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0')))),
  inference(backward_demodulation,[],[f349,f343])).
fof(f343,plain,(
  s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) != s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) != s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) = s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) = s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.inferProps.inferProps_|49||48|9_prove_ltb')).
fof(f349,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.arithmetic.ONE')).
fof(f2403,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f378,f364])).
fof(f364,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|1')).
fof(f378,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7)))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_add|split|6')).
fof(f4330,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0') | $spl12),
  inference(backward_demodulation,[],[f1699,f347])).
fof(f347,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f50])).
fof(f50,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.prim_rec.PRE|split|0')).
fof(f10524,plain,(
  s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))) | $spl12),
  inference(forward_demodulation,[],[f10523,f506])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ! [X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10,X15] : ! [X16] : s('type.list.list'(X15),'const.list.NIL_0') = s('type.list.list'(X15),'const.list.MAP_2'(s(fun(X10,X15),X16),s('type.list.list'(X10),'const.list.NIL_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.list.MAP|split|0')).
fof(f10523,plain,(
  s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.list.NIL_0'))))) | $spl12),
  inference(superposition,[],[f4351,f4331])).
fof(f4331,plain,(
  s('type.list.list'('type.num.num'),'const.list.NIL_0') = s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(backward_demodulation,[],[f4330,f348])).
fof(f348,plain,(
  s('type.list.list'('type.num.num'),'const.list.NIL_0') = s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f96])).
fof(f96,axiom,(
  s('type.list.list'('type.num.num'),'const.list.NIL_0') = s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.rich_list.COUNT_LIST_def|split|0')).
fof(f4351,plain,(
  ( ! [X0] : (s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))))) ) | $spl12),
  inference(backward_demodulation,[],[f4330,f561])).
fof(f561,plain,(
  ( ! [X0] : (s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f388,f374])).
fof(f374,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X20] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X20)))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','ARITY_0_1_const.num.SUC')).
fof(f388,plain,(
  ( ! [X0] : (s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X0))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X0))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X7] : s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.num.num',X7))))) = s('type.list.list'('type.num.num'),'const.list.CONS_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.num.num'),'const.list.MAP_2'(s(fun('type.num.num','type.num.num'),'const.num.SUC_0'),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X7)))))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.rich_list.COUNT_LIST_def|split|1')).
fof(f4329,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl24),
  inference(cnf_transformation,[],[f4329_D])).
fof(f4329_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1692,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f1692_D])).
fof(f1692_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1734,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f404,f1052])).
fof(f1052,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f368,f366])).
fof(f366,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.arithmetic.ADD|split|0')).
fof(f368,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_add|split|1')).
fof(f404,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X7))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|17')).
fof(f4325,plain,(
  ( ! [X20] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))))) ) | $spl22),
  inference(cnf_transformation,[],[f4325_D])).
fof(f4325_D,plain,(
  ( ! [X20] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2336,plain,(
  ( ! [X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f404,f377])).
fof(f377,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7)))))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|9')).
fof(f4323,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24)))))))) ) | $spl20),
  inference(cnf_transformation,[],[f4323_D])).
fof(f4323_D,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24)))))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3047,plain,(
  ( ! [X24] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24)))))))) )),
  inference(superposition,[],[f404,f382])).
fof(f382,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7)))))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|10')).
fof(f4316,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f4315,f344])).
fof(f344,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb',aHL_TRUTH)).
fof(f4315,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f4314,f365])).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|26')).
fof(f4314,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl3),
  inference(backward_demodulation,[],[f4274,f1800])).
fof(f1800,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f1692,f408])).
fof(f408,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|30')).
fof(f4274,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f458,f1781])).
fof(f1781,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f1767,f368])).
fof(f1767,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f369,f405])).
fof(f405,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f276])).
fof(f369,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7)))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_add|split|0')).
fof(f458,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X7,X6] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|31')).
fof(f1701,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl14),
  inference(cnf_transformation,[],[f1701_D])).
fof(f1701_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2581,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl19),
  inference(cnf_transformation,[],[f2581_D])).
fof(f2581_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2518,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f404,f2463])).
fof(f2463,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(backward_demodulation,[],[f2403,f349])).
fof(f1703,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl17),
  inference(cnf_transformation,[],[f1703_D])).
fof(f1703_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1688,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f404,f349])).
fof(f1699,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1687,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f404,f347])).
fof(f1695,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(cnf_transformation,[],[f1695_D])).
fof(f1695_D,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1697,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl11),
  inference(cnf_transformation,[],[f1697_D])).
fof(f1697_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1686,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f404,f559])).
fof(f559,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f379,f349])).
fof(f379,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|11')).
fof(f1694,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) ) | $spl6),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1684,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) )),
  inference(superposition,[],[f404,f357])).
fof(f357,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|6')).
fof(f1693,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) | $spl4),
  inference(cnf_transformation,[],[f1693_D])).
fof(f1693_D,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1683,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) )),
  inference(superposition,[],[f404,f363])).
fof(f363,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.arithmetic.MULT|split|0')).
fof(f1690,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1682,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f404,f362])).
fof(f362,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.inferProps.inferProps_|49||48|9_prove_ltb','thm.numeral.numeral_distrib|split|4')).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 389
% Passive clauses: 1006
% Generated clauses: 10177
% Final active clauses: 260
% Final passive clauses: 413
% Input formulas: 109
% Initial clauses: 207
% 
% Duplicate literals: 4
% Trivial inequalities: 24
% Fw subsumption resolutions: 56
% Fw demodulations: 1477
% Bw demodulations: 246
% Condensations: 4
% 
% Simple tautologies: 47
% Equational tautologies: 20
% Forward subsumptions: 7428
% Fw demodulations to eq. taut.: 99
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 1697
% Factoring: 1
% Forward superposition: 3254
% Backward superposition: 3124
% Self superposition: 56
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 1030
% SAT solver unit clauses: 410
% SAT solver binary clauses: 412
% 
% Sat splits: 10
% Sat splitting refutations: 2
% 
% Memory used [KB]: 7675
% Time elapsed: 0.325 s
% ------------------------------
% ------------------------------
% Success in time 1.63 s
